<?xml version="1.0"?>
<section xml:id="classes_adding-other-methods-to-a-class">
  <title>Adding Other Methods to a Class</title>
  <p>The key advantage of using a class like <c>Point</c> rather than something like a simple
            tuple <c>(7, 6)</c> now becomes apparent.  We can add methods to
            the <c>Point</c> class that are sensible operations for points.  Had we chosen to use a
            tuple to represent the point, we would not have this capability.
            Creating a class like <c>Point</c> brings an exceptional
            amount of <q>organizational power</q> to our programs, and to our thinking.
            We can group together the sensible operations, and the kinds of data
            they apply to, and each instance of the class can have its own state.</p>
  <p>A <term>method</term> behaves like a function but it is invoked on a specific
            instance.  For example, with a list bound to variable L, <c>L.append(7)</c> calls the function append, with the list itself as the first parameter and 7 as the second parameter.   Methods are accessed using dot notation. This is why <c>L.append(7)</c> has 2 parameters even though you may think it only has one: the list stored in the variable <c>L</c> is the first parameter value and 7 is the second.</p>
  <p>Let's add two simple methods to allow a point to give us information about its state.  The <c>getX</c> method, when invoked, will return the value of the x coordinate.</p>
  <p>The implementation of this method is straight forward since we already know how
            to write functions that return values.  One thing to notice is that even though the <c>getX</c> method does not need any other parameter information to do its work, there is still one formal parameter, <c>self</c>.  As we stated earlier, all methods defined in a class that operate on objects of that class will have <c>self</c> as their first parameter.  Again, this serves as a reference to the object itself which in turn gives access to the state data inside the object.</p>
  <program xml:id="chp13_classes4" interactive="activecode" language="python">
    <input>
class Point:
    """ Point class for representing and manipulating x,y coordinates. """

    def __init__(self, initX, initY):

        self.x = initX
        self.y = initY

    def getX(self):
        return self.x

    def getY(self):
        return self.y


p = Point(7,6)
print(p.getX())
print(p.getY())
        </input>
  </program>
  <p>Note that the <c>getX</c> method simply returns the value of the instance variable x from the object self.  In other words, the implementation of the method is to go to the state of the object itself and get the value of <c>x</c>.  Likewise, the <c>getY</c> method looks almost the same.</p>
  <p>Let's add another method, <c>distanceFromOrigin</c>, to see better how methods
            work.  This method will again not need any additional information to do its work, beyond the data stored in the instance variables.
            It will perform a more complex task.</p>
  <program xml:id="chp13_classes5" interactive="activecode" language="python">
    <input>
class Point:
    """ Point class for representing and manipulating x,y coordinates. """

    def __init__(self, initX, initY):

        self.x = initX
        self.y = initY

    def getX(self):
        return self.x

    def getY(self):
        return self.y

    def distanceFromOrigin(self):
        return ((self.x ** 2) + (self.y ** 2)) ** 0.5


p = Point(7,6)
print(p.distanceFromOrigin())
        </input>
  </program>
  <p>Notice that the call of <c>distanceFromOrigin</c> does not <em>explicitly</em>
            supply an argument to match the <c>self</c> parameter.  This is true of all method calls. The definition will always seem to
            have one additional parameter as compared to the invocation.</p>
  <p>
    <term>Check Your Understanding</term>
  </p>
  <p>
    <ol label="1">
      <li>
        <p>Create a class called <c>Animal</c> that accepts two numbers as inputs and assigns them respectively to two instance variables: <c>arms</c> and <c>legs</c>. Create an instance method called <c>limbs</c> that, when called, returns the total number of limbs the animal has. To the variable name <c>spider</c>, assign an instance of <c>Animal</c> that has 4 arms and 4 legs. Call the limbs method on the <c>spider</c> instance and save the result to the variable name <c>spidlimbs</c>.</p>
      </li>
    </ol>
  </p>
  <program xml:id="ac_chp13_classes_01" interactive="activecode" language="python">
    <input>
=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(spider.arms, 4, "Testing that spider was assigned the correct number of arms.")
      self.assertEqual(spider.legs, 4, "Testing that spider was assigned the correct number of legs.")
      self.assertEqual(spidlimbs, 8, "Testing that spidlimbs was assigned correctly.")

myTests().main()
        </input>
  </program>
</section>
