..  Copyright (C)  Paul Resnick.  Permission is granted to copy, distribute
    and/or modify this document under the terms of the GNU Free Documentation
    License, Version 1.3 or any later version published by the Free Software
    Foundation; with Invariant Sections being Forward, Prefaces, and
    Contributor List, no Front-Cover Texts, and no Back-Cover Texts.  A copy of
    the license is included in the section entitled "GNU Free Documentation
    License".

.. qnum::
   :prefix: requests-13-
   :start: 1

Project - OMDB and TasteDive
============================

.. activecode:: assess_ac_24_1_1_1
    :autograde: unittest
    :language: python
    :practice: T

    This project will take you through the process of mashing up data from two different APIs to make movie recommendations. The TasteDive API lets you provide a movie (or bands, TV shows, etc.) as a query input, and returns a set of related items. The OMDB API lets you provide a movie title as a query input and get back data about the movie, including scores from various review sites (Rotten Tomatoes, IMDB, etc.).

    You will put those two together. You will use TasteDive to get related movies for a whole list of titles. You'll combine the resulting lists of related movies, and sort them according to their Rotten Tomatoes scores (which will require making API calls to the OMDB API.)

    To avoid problems with rate limits and site accessibility, we have provided a cache file with results for all the queries you need to make to both OMDB and TasteDive. Just use ``requests_with_caching.get()`` rather than ``requests.get()``. You're welcome to try other queries, but if you do, you will need to get an api key from OMDB.

    Your first task will be to fetch data from TasteDive. The documentation for the API is at https://tastedive.com/read/api.

    Define a function, called ``get_movies_from_tastedive``. It should take one input parameter, a string that is the name of a movie or music artist. The function should return the 5 TasteDive results that are associated with that string; be sure to only get movies, not other kinds of media. It will be a python dictionary with just one key, 'Similar'.

    Try invoking your function with the input "Black Panther".

    HINT: If, when you invoke the function, the results aren't found in the cache, you haven't passed the exact same key-value pairs in the params dictionary as you passed to ``requests_with_caching.get()``.

    ~~~~


    =====

    from unittest.gui import TestCaseGui

    class myTests(TestCaseGui):

      def test_get_movies_from_tastedive_typical(self):
        results = get_movies_from_tastedive("Bridesmaids")['Similar']["Info"][0]["Name"]
        expected = "Bridesmaids"
        self.assertEqual(results, expected, "Testing that the results for Bridesmaids is the expected result.")
        self.assertEqual(type(get_movies_from_tastedive("Black Panther")), type({}), "Testing that the correct python type is returned.")

      def test_get_movies_from_tastedive_limits(self):
        movies = get_movies_from_tastedive("Black Panther")['Similar']['Results']
        self.assertEqual(len(movies), 5, "Testing that get_movies_from_tastedive returns only 5 results")


      def test_get_movies_from_tastedive_restrictions(self):
        results = get_movies_from_tastedive("Tony Bennett")
        non_movies = filter((lambda x: x['Type'] != "movie"), results['Similar']["Results"])
        self.assertEqual(len(list(non_movies)), 0, "Testing that get_movies_from_tastedive retrieves only movies and not music.")



    myTests().main()

.. activecode:: assess_ac_24_1_1_2
    :autograde: unittest
    :language: python
    :practice: T

    Please copy the completed function from above into this active code window. Next, you will need to write a function that extracts just the list of movie titles from a dictionary returned by ``get_movies_from_tastedive``. Call it ``extract_movie_titles``.
    ~~~~


    =====

    from unittest.gui import TestCaseGui

    class myTests(TestCaseGui):

      def test_extract_movie_titles(self):
        print('Running test checking for results from\n   extract_movie_titles(get_movies_from_tastedive("Tony Bennett"))')
        results = ['The Startup Kids', Charlie Chaplin', 'Venus In Fur', 'Loving', 'The African Queen']
        self.assertEqual(extract_movie_titles(get_movies_from_tastedive("Tony Bennett")), results, "Testing that correct results are returned for 'Tony Bennett'.")
        print('Running test checking for results from\n   extract_movie_titles(get_movies_from_tastedive("Black Panther"))')
        sample_from_cache = extract_movie_titles(get_movies_from_tastedive("Black Panther"))
        self.assertEqual(type(sample_from_cache), type([]), "Testing that correct results are returned for 'Black Panther'." )
        
        

    myTests().main()

.. activecode:: assess_ac_24_1_1_3
    :autograde: unittest
    :language: python
    :practice: T

    Please copy the completed functions from the two code windows above into this active code window. Next, you'll write a function, called ``get_related_titles``. It takes *a list of movie titles* as input. It gets five related movies for each from TasteDive, extracts the titles for all of them, and combines them all into a single list. Don't include the same movie twice.
    ~~~~


    =====

    from unittest.gui import TestCaseGui

    class myTests(TestCaseGui):

      def test_related_titles(self):
        expected_results_from_cache = ['Captain Marvel', 'Avengers: Infinity War', 'Ant-Man And The Wasp', 'The Fate Of The Furious', 'Deadpool 2', 'Inhumans', 'Venom', 'American Assassin', 'Black Panther']
        actual_results_from_cache = get_related_titles(["Black Panther", "Captain Marvel"])
        self.assertEqual(get_related_titles([]), [], 'Testing that the correct response is returned when no titles are included.')
        self.assertEqual(actual_results_from_cache, expected_results_from_cache, 'Testing that the correct response is returned when searching for Black Panther and Captain Marvel.')
    
        

    myTests().main()

.. activecode:: assess_ac_24_1_1_4
    :autograde: unittest
    :language: python
    :practice: T

    Your next task will be to fetch data from OMDB. The documentation for the API is at https://www.omdbapi.com/

    Define a function called ``get_movie_data``. It takes in one parameter which is a string that should represent the title of a movie you want to search. The function should return a dictionary with information about that movie.

    Again, use ``requests_with_caching.get()``. For the queries on movies that are already in the cache, you won't need an api key. You will need to provide the following keys: ``t`` and ``r``. (If you want to run queries for other movies, you'll have to get an api key from OMDB, and run the code in a full python environment).

    ~~~~


    =====

    from unittest.gui import TestCaseGui

    class myTests(TestCaseGui):

      def test_get_movie_data(self):
        self.assertEqual(type(get_movie_data("Venom")), type({}), "Testing that the correct python type is returned.")
        self.assertEqual(get_movie_data("Baby Mama")["Title"], "Baby Mama", "Testing that the results match the query.")
        
        

    myTests().main()

.. activecode:: assess_ac_24_1_1_5
    :autograde: unittest
    :language: python
    :practice: T

    Please copy the completed function from above into this active code window. Now write a function called ``get_movie_rating``. It takes an OMDB dictionary result for one movie and extracts the Rotten Tomatoes rating as an integer. For example, if given the OMDB dictionary for "Black Panther", it would return 97. If there is no Rotten Tomatoes rating, return 0.
    ~~~~


    =====

    from unittest.gui import TestCaseGui

    class myTests(TestCaseGui):

      def test_type(self):
        self.assertEqual(type(get_movie_rating(get_movie_data("Deadpool 2"))), type(9), "Testing that a dictionary is returned.")

      def test_output(self):
        self.assertEqual(get_movie_rating(get_movie_data("Venom")), 0, "Testing that the code is acurate for Venom (no rating).")
        self.assertEqual(get_movie_rating(get_movie_data("Deadpool 2")), 83, "Testing that the code for 'Deadpool 2'.")

    myTests().main()

.. activecode:: assess_ac_24_1_1_6
    :autograde: unittest
    :language: python
    :practice: T

    Now, you'll put it all together. Don't forget to copy all of the functions that you have previously defined into this code window. Define a function ``get_sorted_recommendations``. It takes a list of movie titles as an input. It returns a sorted list of related movie titles as output, up to five related movies for each input movie title. The movies should be sorted in descending order by their Rotten Tomatoes rating, as returned by the ``get_movie_rating`` function.
    ~~~~


    =====

    from unittest.gui import TestCaseGui

    class myTests(TestCaseGui):

      def test_output(self):
        sample_actual_recommendations = get_sorted_recommendations(["Bridesmaids", "Sherlock Holmes"])
        sample_expected_recommendations = ['Date Night', 'Baby Mama', 'The Five-Year Engagement', 'Sherlock Holmes: A Game Of Shadows', 'Bachelorette', 'The A-Team', 'Bad Teacher', 'Angels & Demons', 'Prince Of Persia: The Sands Of Time', 'Pirates Of The Caribbean: On Stranger Tides']
        self.assertEqual(sample_actual_recommendations, sample_expected_recommendations, "Testing that actual value returned is the expected value returned.")
        self.assertEqual(type(sample_actual_recommendations), type([]), "Testing that the correct python type is returned.")



    myTests().main()

.. datafile:: permanent_cache.txt

    {
      "https://tastedive.com/api/similarlimit-5_q-Tony Bennett": "{\"Similar\": {\"Info\": [{\"Name\": \"Tony Bennett\", \"Type\": \"music\"}], \"Results\": [{\"Name\": \"Doris Day\", \"Type\": \"music\"}, {\"Name\": \"Tommy Dorsey\", \"Type\": \"music\"}, {\"Name\": \"Boss\", \"Type\": \"music\"}, {\"Name\": \"Natalie Cole\", \"Type\": \"music\"}, {\"Name\": \"Artie Shaw\", \"Type\": \"music\"}]}}",
      "https://api.datamuse.com/wordsmax-3_rel_rhy-muse": "[{\"word\":\"lose\",\"score\":4334,\"numSyllables\":1},{\"word\":\"choose\",\"score\":3823,\"numSyllables\":1},{\"word\":\"abuse\",\"score\":2482,\"numSyllables\":2}]",
      "http://www.omdbapi.com/r-json_t-The Heat": "{\"Title\":\"The Heat\",\"Year\":\"2013\",\"Rated\":\"R\",\"Released\":\"28 Jun 2013\",\"Runtime\":\"117 min\",\"Genre\":\"Action, Comedy, Crime\",\"Director\":\"Paul Feig\",\"Writer\":\"Katie Dippold\",\"Actors\":\"Sandra Bullock, Melissa McCarthy, Demi\u00e1n Bichir, Marlon Wayans\",\"Plot\":\"An uptight FBI Special Agent is paired with a foul-mouthed Boston cop to take down a ruthless drug lord.\",\"Language\":\"English, Spanish, Bulgarian\",\"Country\":\"USA\",\"Awards\":\"7 wins & 16 nominations.\",\"Poster\":\"https://m.media-amazon.com/images/M/MV5BMjA2MDQ2ODM3MV5BMl5BanBnXkFtZTcwNDUzMTQ3OQ@@._V1_SX300.jpg\",\"Ratings\":[{\"Source\":\"Internet Movie Database\",\"Value\":\"6.6/10\"},{\"Source\":\"Rotten Tomatoes\",\"Value\":\"65%\"},{\"Source\":\"Metacritic\",\"Value\":\"60/100\"}],\"Metascore\":\"60\",\"imdbRating\":\"6.6\",\"imdbVotes\":\"149,324\",\"imdbID\":\"tt2404463\",\"Type\":\"movie\",\"DVD\":\"15 Oct 2013\",\"BoxOffice\":\"$158,300,000\",\"Production\":\"20th Century Fox\",\"Website\":\"http://www.theheatmovie.com/#/trailer\",\"Response\":\"True\"}",
      "http://www.omdbapi.com/r-json_t-Yah\u015fi Bat\u0131": "{\"Title\":\"Yahsi Bati - The Ottoman Cowboys\",\"Year\":\"2009\",\"Rated\":\"N/A\",\"Released\":\"01 Jan 2010\",\"Runtime\":\"112 min\",\"Genre\":\"Comedy, Western\",\"Director\":\"\u00d6mer Faruk Sorak\",\"Writer\":\"Cem Yilmaz\",\"Actors\":\"Cem Yilmaz, Ozan G\u00fcven, Zafer Alg\u00f6z, Demet Evgar\",\"Plot\":\"In 1881, two Ottoman Secret Agents travel to the USA, at the Sultan's request, to deliver a valuable diamond as a gift for the President.\",\"Language\":\"Turkish, English\",\"Country\":\"Turkey\",\"Awards\":\"N/A\",\"Poster\":\"https://images-na.ssl-images-amazon.com/images/M/MV5BZDlkZjJiZTItZDAxYS00YzAzLTg3ZTMtMWU5OGEwMjI4ODEzXkEyXkFqcGdeQXVyMjExNjgyMTc@._V1_SX300.jpg\",\"Ratings\":[{\"Source\":\"Internet Movie Database\",\"Value\":\"7.3/10\"}],\"Metascore\":\"N/A\",\"imdbRating\":\"7.3\",\"imdbVotes\":\"20,310\",\"imdbID\":\"tt1567448\",\"Type\":\"movie\",\"DVD\":\"N/A\",\"BoxOffice\":\"N/A\",\"Production\":\"N/A\",\"Website\":\"N/A\",\"Response\":\"True\"}",
      "https://tastedive.com/api/similarlimit-5_q-Black Panther_type-movies": "{\"Similar\": {\"Info\": [{\"Name\": \"Black Panther\", \"Type\": \"movie\"}], \"Results\": [{\"Name\": \"Captain Marvel\", \"Type\": \"movie\"}, {\"Name\": \"Avengers: Infinity War\", \"Type\": \"movie\"}, {\"Name\": \"Ant-Man And The Wasp\", \"Type\": \"movie\"}, {\"Name\": \"The Fate Of The Furious\", \"Type\": \"movie\"}, {\"Name\": \"Deadpool 2\", \"Type\": \"movie\"}]}}",
      "https://tastedive.com/api/similarlimit-5_q-Sherlock Holmes_type-movies": "{\"Similar\": {\"Info\": [{\"Name\": \"Sherlock Holmes\", \"Type\": \"movie\"}], \"Results\": [{\"Name\": \"Sherlock Holmes: A Game Of Shadows\", \"Type\": \"movie\"}, {\"Name\": \"Yah\\u015fi Bat\\u0131\", \"Type\": \"movie\"}, {\"Name\": \"Eyyvah Eyvah\", \"Type\": \"movie\"}, {\"Name\": \"Pirates Of The Caribbean: On Stranger Tides\", \"Type\": \"movie\"}, {\"Name\": \"Prince Of Persia: The Sands Of Time\", \"Type\": \"movie\"}]}}",
      "https://api.flickr.com/services/rest/format-json_media-photos_method-flickr.photos.search_nojsoncallback-1_per_page-5_tag_mode-all_tags-river,mountains": "{\"photos\":{\"page\":1,\"pages\":20771,\"perpage\":5,\"total\":\"103852\",\"photo\":[{\"id\":\"44536834930\",\"owner\":\"63277546@N03\",\"secret\":\"cd1f2b7dc6\",\"server\":\"4853\",\"farm\":5,\"title\":\"Rapids at Manido Falls\",\"ispublic\":1,\"isfriend\":0,\"isfamily\":0},{\"id\":\"44535748420\",\"owner\":\"161486920@N04\",\"secret\":\"fea2fe0a81\",\"server\":\"4875\",\"farm\":5,\"title\":\"Iceland \\u2013 Glacier from Above in Color\",\"ispublic\":1,\"isfriend\":0,\"isfamily\":0},{\"id\":\"46301987692\",\"owner\":\"161486920@N04\",\"secret\":\"44ab129448\",\"server\":\"4819\",\"farm\":5,\"title\":\"Iceland \\u2013 Glacier from Above in Black and White\",\"ispublic\":1,\"isfriend\":0,\"isfamily\":0},{\"id\":\"44012321050\",\"owner\":\"157619092@N06\",\"secret\":\"978dcb0085\",\"server\":\"4888\",\"farm\":5,\"title\":\"Flam, Aurland, Norway\",\"ispublic\":1,\"isfriend\":0,\"isfamily\":0},{\"id\":\"46351760591\",\"owner\":\"144097436@N04\",\"secret\":\"4bc5294842\",\"server\":\"4889\",\"farm\":5,\"title\":\"Jablanica Fluss\",\"ispublic\":1,\"isfriend\":0,\"isfamily\":0}]},\"stat\":\"ok\"}",
      "https://tastedive.com/api/similarq-Black Panther_type-movies": "{\"Similar\": {\"Info\": [{\"Name\": \"Black Panther\", \"Type\": \"movie\"}], \"Results\": [{\"Name\": \"Captain Marvel\", \"Type\": \"movie\"}, {\"Name\": \"Avengers: Infinity War\", \"Type\": \"movie\"}, {\"Name\": \"Ant-Man And The Wasp\", \"Type\": \"movie\"}, {\"Name\": \"The Fate Of The Furious\", \"Type\": \"movie\"}, {\"Name\": \"Deadpool 2\", \"Type\": \"movie\"}, {\"Name\": \"Jumanji: Welcome To The Jungle\", \"Type\": \"movie\"}, {\"Name\": \"Ready Player One\", \"Type\": \"movie\"}, {\"Name\": \"Ocean's 8\", \"Type\": \"movie\"}, {\"Name\": \"Solo: A Star Wars Story\", \"Type\": \"movie\"}, {\"Name\": \"Tag\", \"Type\": \"movie\"}, {\"Name\": \"Thor: Ragnarok\", \"Type\": \"movie\"}, {\"Name\": \"Venom\", \"Type\": \"movie\"}, {\"Name\": \"Pitch Perfect 3\", \"Type\": \"movie\"}, {\"Name\": \"Independence Day: Resurgence\", \"Type\": \"movie\"}, {\"Name\": \"The Divergent Series: Allegiant\", \"Type\": \"movie\"}, {\"Name\": \"Daddy's Home 2\", \"Type\": \"movie\"}, {\"Name\": \"The Greatest Showman\", \"Type\": \"movie\"}, {\"Name\": \"A Quiet Place\", \"Type\": \"movie\"}, {\"Name\": \"Game Night\", \"Type\": \"movie\"}, {\"Name\": \"Power Rangers\", \"Type\": \"movie\"}]}}",
      "https://api.datamuse.com/wordsmax-3_rel_rhy-twitter": "[{\"word\":\"bitter\",\"score\":1632,\"numSyllables\":2},{\"word\":\"litter\",\"score\":1607,\"numSyllables\":2},{\"word\":\"glitter\",\"score\":721,\"numSyllables\":2}]",
      "https://api.datamuse.com/wordsmax-3_rel_rhy-flicker": "[{\"word\":\"liquor\",\"score\":1166,\"numSyllables\":2},{\"word\":\"snicker\",\"score\":941,\"numSyllables\":2},{\"word\":\"vicar\",\"score\":905,\"numSyllables\":2}]",
      "https://itunes.apple.com/searchentity-podcast_term-Ann Arbor": "{\n \"resultCount\":27,\n \"results\": [\n{\"wrapperType\":\"track\", \"kind\":\"podcast\", \"collectionId\":156621464, \"trackId\":156621464, \"artistName\":\"Vineyard Church of Ann Arbor\", \"collectionName\":\"Vineyard Church of Ann Arbor Sermon Podcast\", \"trackName\":\"Vineyard Church of Ann Arbor Sermon Podcast\", \"collectionCensoredName\":\"Vineyard Church of Ann Arbor Sermon Podcast\", \"trackCensoredName\":\"Vineyard Church of Ann Arbor Sermon Podcast\", \"collectionViewUrl\":\"https://itunes.apple.com/us/podcast/vineyard-church-of-ann-arbor-sermon-podcast/id156621464?mt=2&uo=4\", \"feedUrl\":\"http://annarborvineyard.libsyn.com/rss\", \"trackViewUrl\":\"https://itunes.apple.com/us/podcast/vineyard-church-of-ann-arbor-sermon-podcast/id156621464?mt=2&uo=4\", \"artworkUrl30\":\"https://is3-ssl.mzstatic.com/image/thumb/Music117/v4/47/8b/1f/478b1f4c-e396-94a2-f5f7-6c66c71bf773/source/30x30bb.jpg\", \"artworkUrl60\":\"https://is3-ssl.mzstatic.com/image/thumb/Music117/v4/47/8b/1f/478b1f4c-e396-94a2-f5f7-6c66c71bf773/source/60x60bb.jpg\", \"artworkUrl100\":\"https://is3-ssl.mzstatic.com/image/thumb/Music117/v4/47/8b/1f/478b1f4c-e396-94a2-f5f7-6c66c71bf773/source/100x100bb.jpg\", \"collectionPrice\":0.00, \"trackPrice\":0.00, \"trackRentalPrice\":0, \"collectionHdPrice\":0, \"trackHdPrice\":0, \"trackHdRentalPrice\":0, \"releaseDate\":\"2018-12-10T15:00:00Z\", \"collectionExplicitness\":\"cleaned\", \"trackExplicitness\":\"cleaned\", \"trackCount\":300, \"country\":\"USA\", \"currency\":\"USD\", \"primaryGenreName\":\"Christianity\", \"contentAdvisoryRating\":\"Clean\", \"artworkUrl600\":\"https://is3-ssl.mzstatic.com/image/thumb/Music117/v4/47/8b/1f/478b1f4c-e396-94a2-f5f7-6c66c71bf773/source/600x600bb.jpg\", \"genreIds\":[\"1439\", \"26\", \"1314\"], \"genres\":[\"Christianity\", \"Podcasts\", \"Religion & Spirituality\"]}, \n{\"wrapperType\":\"track\", \"kind\":\"podcast\", \"collectionId\":1123046320, \"trackId\":1123046320, \"artistName\":\"annarbor242\", \"collectionName\":\"2|42 Community Church - Ann Arbor\", \"trackName\":\"2|42 Community Church - Ann Arbor\", \"collectionCensoredName\":\"2|42 Community Church - Ann Arbor\", \"trackCensoredName\":\"2|42 Community Church - Ann Arbor\", \"collectionViewUrl\":\"https://itunes.apple.com/us/podcast/2-42-community-church-ann-arbor/id1123046320?mt=2&uo=4\", \"feedUrl\":\"http://annarborpodcast.242community.com/feed.xml\", \"trackViewUrl\":\"https://itunes.apple.com/us/podcast/2-42-community-church-ann-arbor/id1123046320?mt=2&uo=4\", \"artworkUrl30\":\"https://is1-ssl.mzstatic.com/image/thumb/Music62/v4/b1/e0/a1/b1e0a19f-afc1-7a65-ce2e-8c4dd753939c/source/30x30bb.jpg\", \"artworkUrl60\":\"https://is1-ssl.mzstatic.com/image/thumb/Music62/v4/b1/e0/a1/b1e0a19f-afc1-7a65-ce2e-8c4dd753939c/source/60x60bb.jpg\", \"artworkUrl100\":\"https://is1-ssl.mzstatic.com/image/thumb/Music62/v4/b1/e0/a1/b1e0a19f-afc1-7a65-ce2e-8c4dd753939c/source/100x100bb.jpg\", \"collectionPrice\":0.00, \"trackPrice\":0.00, \"trackRentalPrice\":0, \"collectionHdPrice\":0, \"trackHdPrice\":0, \"trackHdRentalPrice\":0, \"releaseDate\":\"2018-12-10T17:10:00Z\", \"collectionExplicitness\":\"cleaned\", \"trackExplicitness\":\"cleaned\", \"trackCount\":100, \"country\":\"USA\", \"currency\":\"USD\", \"primaryGenreName\":\"Christianity\", \"contentAdvisoryRating\":\"Clean\", \"artworkUrl600\":\"https://is1-ssl.mzstatic.com/image/thumb/Music62/v4/b1/e0/a1/b1e0a19f-afc1-7a65-ce2e-8c4dd753939c/source/600x600bb.jpg\", \"genreIds\":[\"1439\", \"26\", \"1314\"], \"genres\":[\"Christianity\", \"Podcasts\", \"Religion & Spirituality\"]}, \n{\"wrapperType\":\"track\", \"kind\":\"podcast\", \"collectionId\":118253222, \"trackId\":118253222, \"artistName\":\"New Life Church - Ann Arbor\", \"collectionName\":\"Sermons \u2013 NewLifeA2.org\", \"trackName\":\"Sermons \u2013 NewLifeA2.org\", \"collectionCensoredName\":\"Sermons \u2013 NewLifeA2.org\", \"trackCensoredName\":\"Sermons \u2013 NewLifeA2.org\", \"collectionViewUrl\":\"https://itunes.apple.com/us/podcast/sermons-newlifea2-org/id118253222?mt=2&uo=4\", \"feedUrl\":\"http://feeds.feedburner.com/newlifea2/\", \"trackViewUrl\":\"https://itunes.apple.com/us/podcast/sermons-newlifea2-org/id118253222?mt=2&uo=4\", \"artworkUrl30\":\"https://is3-ssl.mzstatic.com/image/thumb/Music4/v4/e3/03/b8/e303b87f-43e4-e7c4-f94e-3225e289421a/source/30x30bb.jpg\", \"artworkUrl60\":\"https://is3-ssl.mzstatic.com/image/thumb/Music4/v4/e3/03/b8/e303b87f-43e4-e7c4-f94e-3225e289421a/source/60x60bb.jpg\", \"artworkUrl100\":\"https://is3-ssl.mzstatic.com/image/thumb/Music4/v4/e3/03/b8/e303b87f-43e4-e7c4-f94e-3225e289421a/source/100x100bb.jpg\", \"collectionPrice\":0.00, \"trackPrice\":0.00, \"trackRentalPrice\":0, \"collectionHdPrice\":0, \"trackHdPrice\":0, \"trackHdRentalPrice\":0, \"releaseDate\":\"2018-12-16T15:01:00Z\", \"collectionExplicitness\":\"cleaned\", \"trackExplicitness\":\"cleaned\", \"trackCount\":209, \"country\":\"USA\", \"currency\":\"USD\", \"primaryGenreName\":\"Christianity\", \"contentAdvisoryRating\":\"Clean\", \"artworkUrl600\":\"https://is3-ssl.mzstatic.com/image/thumb/Music4/v4/e3/03/b8/e303b87f-43e4-e7c4-f94e-3225e289421a/source/600x600bb.jpg\", \"genreIds\":[\"1439\", \"26\", \"1314\"], \"genres\":[\"Christianity\", \"Podcasts\", \"Religion & Spirituality\"]}, \n{\"wrapperType\":\"track\", \"kind\":\"podcast\", \"collectionId\":446900910, \"trackId\":446900910, \"artistName\":\"Harvest Mission Community Church\", \"collectionName\":\"Harvest Mission Community Church (Ann Arbor) Sermons\", \"trackName\":\"Harvest Mission Community Church (Ann Arbor) Sermons\", \"collectionCensoredName\":\"Harvest Mission Community Church (Ann Arbor) Sermons\", \"trackCensoredName\":\"Harvest Mission Community Church (Ann Arbor) Sermons\", \"collectionViewUrl\":\"https://itunes.apple.com/us/podcast/harvest-mission-community-church-ann-arbor-sermons/id446900910?mt=2&uo=4\", \"feedUrl\":\"http://annarbor.hmcc.net/podcast/\", \"trackViewUrl\":\"https://itunes.apple.com/us/podcast/harvest-mission-community-church-ann-arbor-sermons/id446900910?mt=2&uo=4\", \"artworkUrl30\":\"https://is4-ssl.mzstatic.com/image/thumb/Music62/v4/fc/82/99/fc829985-3f12-7bb2-89e5-3964e82907bf/source/30x30bb.jpg\", \"artworkUrl60\":\"https://is4-ssl.mzstatic.com/image/thumb/Music62/v4/fc/82/99/fc829985-3f12-7bb2-89e5-3964e82907bf/source/60x60bb.jpg\", \"artworkUrl100\":\"https://is4-ssl.mzstatic.com/image/thumb/Music62/v4/fc/82/99/fc829985-3f12-7bb2-89e5-3964e82907bf/source/100x100bb.jpg\", \"collectionPrice\":0.00, \"trackPrice\":0.00, \"trackRentalPrice\":0, \"collectionHdPrice\":0, \"trackHdPrice\":0, \"trackHdRentalPrice\":0, \"releaseDate\":\"2018-12-16T00:00:00Z\", \"collectionExplicitness\":\"cleaned\", \"trackExplicitness\":\"cleaned\", \"trackCount\":10, \"country\":\"USA\", \"currency\":\"USD\", \"primaryGenreName\":\"Christianity\", \"contentAdvisoryRating\":\"Clean\", \"artworkUrl600\":\"https://is4-ssl.mzstatic.com/image/thumb/Music62/v4/fc/82/99/fc829985-3f12-7bb2-89e5-3964e82907bf/source/600x600bb.jpg\", \"genreIds\":[\"1439\", \"26\", \"1314\"], \"genres\":[\"Christianity\", \"Podcasts\", \"Religion & Spirituality\"]}, \n{\"wrapperType\":\"track\", \"kind\":\"podcast\", \"collectionId\":121092869, \"trackId\":121092869, \"artistName\":\"WCBN-FM Ann Arbor; T Hetzel\", \"collectionName\":\"Living Writers\", \"trackName\":\"Living Writers\", \"collectionCensoredName\":\"Living Writers\", \"trackCensoredName\":\"Living Writers\", \"collectionViewUrl\":\"https://itunes.apple.com/us/podcast/living-writers/id121092869?mt=2&uo=4\", \"feedUrl\":\"http://beanball.wcbn.org/Living_Writers/\", \"trackViewUrl\":\"https://itunes.apple.com/us/podcast/living-writers/id121092869?mt=2&uo=4\", \"artworkUrl30\":\"https://is3-ssl.mzstatic.com/image/thumb/Music71/v4/80/57/b1/8057b172-4bd7-e1a0-45c8-c240c33736ab/source/30x30bb.jpg\", \"artworkUrl60\":\"https://is3-ssl.mzstatic.com/image/thumb/Music71/v4/80/57/b1/8057b172-4bd7-e1a0-45c8-c240c33736ab/source/60x60bb.jpg\", \"artworkUrl100\":\"https://is3-ssl.mzstatic.com/image/thumb/Music71/v4/80/57/b1/8057b172-4bd7-e1a0-45c8-c240c33736ab/source/100x100bb.jpg\", \"collectionPrice\":0.00, \"trackPrice\":0.00, \"trackRentalPrice\":0, \"collectionHdPrice\":0, \"trackHdPrice\":0, \"trackHdRentalPrice\":0, \"releaseDate\":\"2018-12-13T23:02:00Z\", \"collectionExplicitness\":\"cleaned\", \"trackExplicitness\":\"cleaned\", \"trackCount\":294, \"country\":\"USA\", \"currency\":\"USD\", \"primaryGenreName\":\"Literature\", \"contentAdvisoryRating\":\"Clean\", \"artworkUrl600\":\"https://is3-ssl.mzstatic.com/image/thumb/Music71/v4/80/57/b1/8057b172-4bd7-e1a0-45c8-c240c33736ab/source/600x600bb.jpg\", \"genreIds\":[\"1401\", \"26\", \"1301\"], \"genres\":[\"Literature\", \"Podcasts\", \"Arts\"]}, \n{\"wrapperType\":\"track\", \"kind\":\"podcast\", \"collectionId\":719611936, \"trackId\":719611936, \"artistName\":\"Grace Ann Arbor\", \"collectionName\":\"Grace Ann Arbor Church\", \"trackName\":\"Grace Ann Arbor Church\", \"collectionCensoredName\":\"Grace Ann Arbor Church\", \"trackCensoredName\":\"Grace Ann Arbor Church\", \"collectionViewUrl\":\"https://itunes.apple.com/us/podcast/grace-ann-arbor-church/id719611936?mt=2&uo=4\", \"feedUrl\":\"http://graceannarbor.cloversites.com/podcast/85b60233-1049-4bff-ad57-97b021cde2d8.xml\", \"trackViewUrl\":\"https://itunes.apple.com/us/podcast/grace-ann-arbor-church/id719611936?mt=2&uo=4\", \"artworkUrl30\":\"https://is2-ssl.mzstatic.com/image/thumb/Music128/v4/c7/7b/59/c77b5904-2595-50b5-f98d-00aa89f0eb30/source/30x30bb.jpg\", \"artworkUrl60\":\"https://is2-ssl.mzstatic.com/image/thumb/Music128/v4/c7/7b/59/c77b5904-2595-50b5-f98d-00aa89f0eb30/source/60x60bb.jpg\", \"artworkUrl100\":\"https://is2-ssl.mzstatic.com/image/thumb/Music128/v4/c7/7b/59/c77b5904-2595-50b5-f98d-00aa89f0eb30/source/100x100bb.jpg\", \"collectionPrice\":0.00, \"trackPrice\":0.00, \"trackRentalPrice\":0, \"collectionHdPrice\":0, \"trackHdPrice\":0, \"trackHdRentalPrice\":0, \"releaseDate\":\"2018-12-16T00:00:00Z\", \"collectionExplicitness\":\"cleaned\", \"trackExplicitness\":\"cleaned\", \"trackCount\":125, \"country\":\"USA\", \"currency\":\"USD\", \"primaryGenreName\":\"Christianity\", \"contentAdvisoryRating\":\"Clean\", \"artworkUrl600\":\"https://is2-ssl.mzstatic.com/image/thumb/Music128/v4/c7/7b/59/c77b5904-2595-50b5-f98d-00aa89f0eb30/source/600x600bb.jpg\", \"genreIds\":[\"1439\", \"26\", \"1314\"], \"genres\":[\"Christianity\", \"Podcasts\", \"Religion & Spirituality\"]}, \n{\"wrapperType\":\"track\", \"kind\":\"podcast\", \"collectionId\":1156696806, \"trackId\":1156696806, \"artistName\":\"Grace Bible Church Ann Arbor\", \"collectionName\":\"Grace Bible Church Ann Arbor\", \"trackName\":\"Grace Bible Church Ann Arbor\", \"collectionCensoredName\":\"Grace Bible Church Ann Arbor\", \"trackCensoredName\":\"Grace Bible Church Ann Arbor\", \"collectionViewUrl\":\"https://itunes.apple.com/us/podcast/grace-bible-church-ann-arbor/id1156696806?mt=2&uo=4\", \"feedUrl\":\"http://www.gracea2.org/mediafiles/sermons.xml\", \"trackViewUrl\":\"https://itunes.apple.com/us/podcast/grace-bible-church-ann-arbor/id1156696806?mt=2&uo=4\", \"artworkUrl30\":\"https://is1-ssl.mzstatic.com/image/thumb/Music71/v4/18/ec/b2/18ecb265-8b35-716c-5a3a-ca83fa6b3b32/source/30x30bb.jpg\", \"artworkUrl60\":\"https://is1-ssl.mzstatic.com/image/thumb/Music71/v4/18/ec/b2/18ecb265-8b35-716c-5a3a-ca83fa6b3b32/source/60x60bb.jpg\", \"artworkUrl100\":\"https://is1-ssl.mzstatic.com/image/thumb/Music71/v4/18/ec/b2/18ecb265-8b35-716c-5a3a-ca83fa6b3b32/source/100x100bb.jpg\", \"collectionPrice\":0.00, \"trackPrice\":0.00, \"trackRentalPrice\":0, \"collectionHdPrice\":0, \"trackHdPrice\":0, \"trackHdRentalPrice\":0, \"releaseDate\":\"2018-12-09T20:07:00Z\", \"collectionExplicitness\":\"cleaned\", \"trackExplicitness\":\"cleaned\", \"trackCount\":10, \"country\":\"USA\", \"currency\":\"USD\", \"primaryGenreName\":\"Religion & Spirituality\", \"contentAdvisoryRating\":\"Clean\", \"artworkUrl600\":\"https://is1-ssl.mzstatic.com/image/thumb/Music71/v4/18/ec/b2/18ecb265-8b35-716c-5a3a-ca83fa6b3b32/source/600x600bb.jpg\", \"genreIds\":[\"1314\", \"26\"], \"genres\":[\"Religion & Spirituality\", \"Podcasts\"]}, \n{\"wrapperType\":\"track\", \"kind\":\"podcast\", \"collectionId\":1051188622, \"trackId\":1051188622, \"artistName\":\"WCBN-FM Ann Arbor\", \"collectionName\":\"It\u2019s Hot In Here\", \"trackName\":\"It\u2019s Hot In Here\", \"collectionCensoredName\":\"It\u2019s Hot In Here\", \"trackCensoredName\":\"It\u2019s Hot In Here\", \"collectionViewUrl\":\"https://itunes.apple.com/us/podcast/its-hot-in-here/id1051188622?mt=2&uo=4\", \"feedUrl\":\"http://hotinhere.us/feed/podcast\", \"trackViewUrl\":\"https://itunes.apple.com/us/podcast/its-hot-in-here/id1051188622?mt=2&uo=4\", \"artworkUrl30\":\"https://is2-ssl.mzstatic.com/image/thumb/Music71/v4/3e/56/3f/3e563fb1-8e50-8c6e-c5c0-a94c57151f71/source/30x30bb.jpg\", \"artworkUrl60\":\"https://is2-ssl.mzstatic.com/image/thumb/Music71/v4/3e/56/3f/3e563fb1-8e50-8c6e-c5c0-a94c57151f71/source/60x60bb.jpg\", \"artworkUrl100\":\"https://is2-ssl.mzstatic.com/image/thumb/Music71/v4/3e/56/3f/3e563fb1-8e50-8c6e-c5c0-a94c57151f71/source/100x100bb.jpg\", \"collectionPrice\":0.00, \"trackPrice\":0.00, \"trackRentalPrice\":0, \"collectionHdPrice\":0, \"trackHdPrice\":0, \"trackHdRentalPrice\":0, \"releaseDate\":\"2018-11-18T17:26:00Z\", \"collectionExplicitness\":\"cleaned\", \"trackExplicitness\":\"cleaned\", \"trackCount\":96, \"country\":\"USA\", \"currency\":\"USD\", \"primaryGenreName\":\"Natural Sciences\", \"contentAdvisoryRating\":\"Clean\", \"artworkUrl600\":\"https://is2-ssl.mzstatic.com/image/thumb/Music71/v4/3e/56/3f/3e563fb1-8e50-8c6e-c5c0-a94c57151f71/source/600x600bb.jpg\", \"genreIds\":[\"1477\", \"26\", \"1315\"], \"genres\":[\"Natural Sciences\", \"Podcasts\", \"Science & Medicine\"]}, \n{\"wrapperType\":\"track\", \"kind\":\"podcast\", \"collectionId\":1425282628, \"trackId\":1425282628, \"artistName\":\"Ann Arbor District Library\", \"collectionName\":\"Presenting Alfred Hitchcock Presents | Ann Arbor District Library\", \"trackName\":\"Presenting Alfred Hitchcock Presents | Ann Arbor District Library\", \"collectionCensoredName\":\"Presenting Alfred Hitchcock Presents | Ann Arbor District Library\", \"trackCensoredName\":\"Presenting Alfred Hitchcock Presents | Ann Arbor District Library\", \"collectionViewUrl\":\"https://itunes.apple.com/us/podcast/presenting-alfred-hitchcock-presents-ann-arbor-district/id1425282628?mt=2&uo=4\", \"feedUrl\":\"https://aadl.org/hitchcock/feed\", \"trackViewUrl\":\"https://itunes.apple.com/us/podcast/presenting-alfred-hitchcock-presents-ann-arbor-district/id1425282628?mt=2&uo=4\", \"artworkUrl30\":\"https://is2-ssl.mzstatic.com/image/thumb/Music128/v4/65/22/4b/65224b1c-6e5a-b2cb-9c80-8d0e232d26b8/source/30x30bb.jpg\", \"artworkUrl60\":\"https://is2-ssl.mzstatic.com/image/thumb/Music128/v4/65/22/4b/65224b1c-6e5a-b2cb-9c80-8d0e232d26b8/source/60x60bb.jpg\", \"artworkUrl100\":\"https://is2-ssl.mzstatic.com/image/thumb/Music128/v4/65/22/4b/65224b1c-6e5a-b2cb-9c80-8d0e232d26b8/source/100x100bb.jpg\", \"collectionPrice\":0.00, \"trackPrice\":0.00, \"trackRentalPrice\":0, \"collectionHdPrice\":0, \"trackHdPrice\":0, \"trackHdRentalPrice\":0, \"releaseDate\":\"2018-12-03T20:00:00Z\", \"collectionExplicitness\":\"cleaned\", \"trackExplicitness\":\"cleaned\", \"trackCount\":6, \"country\":\"USA\", \"currency\":\"USD\", \"primaryGenreName\":\"Literature\", \"contentAdvisoryRating\":\"Clean\", \"artworkUrl600\":\"https://is2-ssl.mzstatic.com/image/thumb/Music128/v4/65/22/4b/65224b1c-6e5a-b2cb-9c80-8d0e232d26b8/source/600x600bb.jpg\", \"genreIds\":[\"1401\", \"26\", \"1301\", \"1311\", \"1406\", \"1324\", \"1405\"], \"genres\":[\"Literature\", \"Podcasts\", \"Arts\", \"News & Politics\", \"Visual Arts\", \"Society & Culture\", \"Performing Arts\"]}, \n{\"wrapperType\":\"track\", \"kind\":\"podcast\", \"collectionId\":1085131603, \"trackId\":1085131603, \"artistName\":\"Blue Ocean Faith Ann Arbor\", \"collectionName\":\"Blue Ocean Faith Ann Arbor Sunday Sermons\", \"trackName\":\"Blue Ocean Faith Ann Arbor Sunday Sermons\", \"collectionCensoredName\":\"Blue Ocean Faith Ann Arbor Sunday Sermons\", \"trackCensoredName\":\"Blue Ocean Faith Ann Arbor Sunday Sermons\", \"collectionViewUrl\":\"https://itunes.apple.com/us/podcast/blue-ocean-faith-ann-arbor-sunday-sermons/id1085131603?mt=2&uo=4\", \"feedUrl\":\"http://www.a2blue.org/feed/podcast\", \"trackViewUrl\":\"https://itunes.apple.com/us/podcast/blue-ocean-faith-ann-arbor-sunday-sermons/id1085131603?mt=2&uo=4\", \"artworkUrl30\":\"https://is5-ssl.mzstatic.com/image/thumb/Music122/v4/8a/02/72/8a02724f-fa99-4a4a-607c-c76dfe6a9a93/source/30x30bb.jpg\", \"artworkUrl60\":\"https://is5-ssl.mzstatic.com/image/thumb/Music122/v4/8a/02/72/8a02724f-fa99-4a4a-607c-c76dfe6a9a93/source/60x60bb.jpg\", \"artworkUrl100\":\"https://is5-ssl.mzstatic.com/image/thumb/Music122/v4/8a/02/72/8a02724f-fa99-4a4a-607c-c76dfe6a9a93/source/100x100bb.jpg\", \"collectionPrice\":0.00, \"trackPrice\":0.00, \"trackRentalPrice\":0, \"collectionHdPrice\":0, \"trackHdPrice\":0, \"trackHdRentalPrice\":0, \"releaseDate\":\"2018-12-17T05:01:00Z\", \"collectionExplicitness\":\"cleaned\", \"trackExplicitness\":\"cleaned\", \"trackCount\":26, \"country\":\"USA\", \"currency\":\"USD\", \"primaryGenreName\":\"Christianity\", \"contentAdvisoryRating\":\"Clean\", \"artworkUrl600\":\"https://is5-ssl.mzstatic.com/image/thumb/Music122/v4/8a/02/72/8a02724f-fa99-4a4a-607c-c76dfe6a9a93/source/600x600bb.jpg\", \"genreIds\":[\"1439\", \"26\", \"1314\"], \"genres\":[\"Christianity\", \"Podcasts\", \"Religion & Spirituality\"]}, \n{\"wrapperType\":\"track\", \"kind\":\"podcast\", \"collectionId\":1046987958, \"trackId\":1046987958, \"artistName\":\"Radiant Church - Ann Arbor\", \"collectionName\":\"Radiant Church - Ann Arbor: Sermons\", \"trackName\":\"Radiant Church - Ann Arbor: Sermons\", \"collectionCensoredName\":\"Radiant Church - Ann Arbor: Sermons\", \"trackCensoredName\":\"Radiant Church - Ann Arbor: Sermons\", \"collectionViewUrl\":\"https://itunes.apple.com/us/podcast/radiant-church-ann-arbor-sermons/id1046987958?mt=2&uo=4\", \"feedUrl\":\"http://feeds.feedburner.com/RadiantChurch-AnnArbor-Sermons\", \"trackViewUrl\":\"https://itunes.apple.com/us/podcast/radiant-church-ann-arbor-sermons/id1046987958?mt=2&uo=4\", \"artworkUrl30\":\"https://is4-ssl.mzstatic.com/image/thumb/Music118/v4/7f/b5/40/7fb540fb-2500-d775-7e0f-1cfa8bce19ff/source/30x30bb.jpg\", \"artworkUrl60\":\"https://is4-ssl.mzstatic.com/image/thumb/Music118/v4/7f/b5/40/7fb540fb-2500-d775-7e0f-1cfa8bce19ff/source/60x60bb.jpg\", \"artworkUrl100\":\"https://is4-ssl.mzstatic.com/image/thumb/Music118/v4/7f/b5/40/7fb540fb-2500-d775-7e0f-1cfa8bce19ff/source/100x100bb.jpg\", \"collectionPrice\":0.00, \"trackPrice\":0.00, \"trackRentalPrice\":0, \"collectionHdPrice\":0, \"trackHdPrice\":0, \"trackHdRentalPrice\":0, \"releaseDate\":\"2018-11-05T23:53:00Z\", \"collectionExplicitness\":\"cleaned\", \"trackExplicitness\":\"cleaned\", \"trackCount\":50, \"country\":\"USA\", \"currency\":\"USD\", \"primaryGenreName\":\"Christianity\", \"contentAdvisoryRating\":\"Clean\", \"artworkUrl600\":\"https://is4-ssl.mzstatic.com/image/thumb/Music118/v4/7f/b5/40/7fb540fb-2500-d775-7e0f-1cfa8bce19ff/source/600x600bb.jpg\", \"genreIds\":[\"1439\", \"26\", \"1314\", \"1309\"], \"genres\":[\"Christianity\", \"Podcasts\", \"Religion & Spirituality\", \"TV & Film\"]}, \n{\"wrapperType\":\"track\", \"kind\":\"podcast\", \"collectionId\":288905801, \"trackId\":288905801, \"artistName\":\"Calvary Bible Church of Ann Arbor / Ypsilanti\", \"collectionName\":\"Calvary Sunday Messages\", \"trackName\":\"Calvary Sunday Messages\", \"collectionCensoredName\":\"Calvary Sunday Messages\", \"trackCensoredName\":\"Calvary Sunday Messages\", \"collectionViewUrl\":\"https://itunes.apple.com/us/podcast/calvary-sunday-messages/id288905801?mt=2&uo=4\", \"feedUrl\":\"https://podcasts.subsplash.com/rw7sztf/podcast.rss\", \"trackViewUrl\":\"https://itunes.apple.com/us/podcast/calvary-sunday-messages/id288905801?mt=2&uo=4\", \"artworkUrl30\":\"https://is5-ssl.mzstatic.com/image/thumb/Music128/v4/f9/6c/e3/f96ce363-7e0a-b3a2-273a-b16597d8e6a3/source/30x30bb.jpg\", \"artworkUrl60\":\"https://is5-ssl.mzstatic.com/image/thumb/Music128/v4/f9/6c/e3/f96ce363-7e0a-b3a2-273a-b16597d8e6a3/source/60x60bb.jpg\", \"artworkUrl100\":\"https://is5-ssl.mzstatic.com/image/thumb/Music128/v4/f9/6c/e3/f96ce363-7e0a-b3a2-273a-b16597d8e6a3/source/100x100bb.jpg\", \"collectionPrice\":0.00, \"trackPrice\":0.00, \"trackRentalPrice\":0, \"collectionHdPrice\":0, \"trackHdPrice\":0, \"trackHdRentalPrice\":0, \"releaseDate\":\"2018-12-09T10:00:00Z\", \"collectionExplicitness\":\"cleaned\", \"trackExplicitness\":\"cleaned\", \"trackCount\":131, \"country\":\"USA\", \"currency\":\"USD\", \"primaryGenreName\":\"Christianity\", \"contentAdvisoryRating\":\"Clean\", \"artworkUrl600\":\"https://is5-ssl.mzstatic.com/image/thumb/Music128/v4/f9/6c/e3/f96ce363-7e0a-b3a2-273a-b16597d8e6a3/source/600x600bb.jpg\", \"genreIds\":[\"1439\", \"26\", \"1314\"], \"genres\":[\"Christianity\", \"Podcasts\", \"Religion & Spirituality\"]}, \n{\"wrapperType\":\"track\", \"kind\":\"podcast\", \"collectionId\":1233253932, \"trackId\":1233253932, \"artistName\":\"Ann Arbor District Library\", \"collectionName\":\"Fellow Youths | Ann Arbor District Library\", \"trackName\":\"Fellow Youths | Ann Arbor District Library\", \"collectionCensoredName\":\"Fellow Youths | Ann Arbor District Library\", \"trackCensoredName\":\"Fellow Youths | Ann Arbor District Library\", \"collectionViewUrl\":\"https://itunes.apple.com/us/podcast/fellow-youths-ann-arbor-district-library/id1233253932?mt=2&uo=4\", \"feedUrl\":\"https://aadl.org/fellowyouths/feed\", \"trackViewUrl\":\"https://itunes.apple.com/us/podcast/fellow-youths-ann-arbor-district-library/id1233253932?mt=2&uo=4\", \"artworkUrl30\":\"https://is1-ssl.mzstatic.com/image/thumb/Music118/v4/b3/dc/e0/b3dce009-2378-ca7b-a646-ba4e73335e53/source/30x30bb.jpg\", \"artworkUrl60\":\"https://is1-ssl.mzstatic.com/image/thumb/Music118/v4/b3/dc/e0/b3dce009-2378-ca7b-a646-ba4e73335e53/source/60x60bb.jpg\", \"artworkUrl100\":\"https://is1-ssl.mzstatic.com/image/thumb/Music118/v4/b3/dc/e0/b3dce009-2378-ca7b-a646-ba4e73335e53/source/100x100bb.jpg\", \"collectionPrice\":0.00, \"trackPrice\":0.00, \"trackRentalPrice\":0, \"collectionHdPrice\":0, \"trackHdPrice\":0, \"trackHdRentalPrice\":0, \"releaseDate\":\"2018-12-04T10:30:00Z\", \"collectionExplicitness\":\"explicit\", \"trackExplicitness\":\"explicit\", \"trackCount\":17, \"country\":\"USA\", \"currency\":\"USD\", \"primaryGenreName\":\"Literature\", \"contentAdvisoryRating\":\"Explicit\", \"artworkUrl600\":\"https://is1-ssl.mzstatic.com/image/thumb/Music118/v4/b3/dc/e0/b3dce009-2378-ca7b-a646-ba4e73335e53/source/600x600bb.jpg\", \"genreIds\":[\"1401\", \"26\", \"1301\", \"1325\", \"1475\", \"1324\", \"1305\"], \"genres\":[\"Literature\", \"Podcasts\", \"Arts\", \"Government & Organizations\", \"Local\", \"Society & Culture\", \"Kids & Family\"]}, \n{\"wrapperType\":\"track\", \"kind\":\"podcast\", \"collectionId\":1359378554, \"trackId\":1359378554, \"artistName\":\"Ann Arbor District Library\", \"collectionName\":\"Behind The Marquee | Ann Arbor District Library\", \"trackName\":\"Behind The Marquee | Ann Arbor District Library\", \"collectionCensoredName\":\"Behind The Marquee | Ann Arbor District Library\", \"trackCensoredName\":\"Behind The Marquee | Ann Arbor District Library\", \"collectionViewUrl\":\"https://itunes.apple.com/us/podcast/behind-the-marquee-ann-arbor-district-library/id1359378554?mt=2&uo=4\", \"feedUrl\":\"https://aadl.org/marquee/feed\", \"trackViewUrl\":\"https://itunes.apple.com/us/podcast/behind-the-marquee-ann-arbor-district-library/id1359378554?mt=2&uo=4\", \"artworkUrl30\":\"https://is3-ssl.mzstatic.com/image/thumb/Music118/v4/38/e9/ba/38e9bac2-529b-1b01-5b1a-491d91e7c37b/source/30x30bb.jpg\", \"artworkUrl60\":\"https://is3-ssl.mzstatic.com/image/thumb/Music118/v4/38/e9/ba/38e9bac2-529b-1b01-5b1a-491d91e7c37b/source/60x60bb.jpg\", \"artworkUrl100\":\"https://is3-ssl.mzstatic.com/image/thumb/Music118/v4/38/e9/ba/38e9bac2-529b-1b01-5b1a-491d91e7c37b/source/100x100bb.jpg\", \"collectionPrice\":0.00, \"trackPrice\":0.00, \"trackRentalPrice\":0, \"collectionHdPrice\":0, \"trackHdPrice\":0, \"trackHdRentalPrice\":0, \"releaseDate\":\"2018-12-07T10:30:00Z\", \"collectionExplicitness\":\"cleaned\", \"trackExplicitness\":\"cleaned\", \"trackCount\":22, \"country\":\"USA\", \"currency\":\"USD\", \"primaryGenreName\":\"Performing Arts\", \"contentAdvisoryRating\":\"Clean\", \"artworkUrl600\":\"https://is3-ssl.mzstatic.com/image/thumb/Music118/v4/38/e9/ba/38e9bac2-529b-1b01-5b1a-491d91e7c37b/source/600x600bb.jpg\", \"genreIds\":[\"1405\", \"26\", \"1301\", \"1406\", \"1324\", \"1462\", \"1311\"], \"genres\":[\"Performing Arts\", \"Podcasts\", \"Arts\", \"Visual Arts\", \"Society & Culture\", \"History\", \"News & Politics\"]}, \n{\"wrapperType\":\"track\", \"kind\":\"podcast\", \"collectionId\":1142242316, \"trackId\":1142242316, \"artistName\":\"Ann Arbor SPARK\", \"collectionName\":\"Ann Arbor SPARK CEO Podcast\", \"trackName\":\"Ann Arbor SPARK CEO Podcast\", \"collectionCensoredName\":\"Ann Arbor SPARK CEO Podcast\", \"trackCensoredName\":\"Ann Arbor SPARK CEO Podcast\", \"collectionViewUrl\":\"https://itunes.apple.com/us/podcast/ann-arbor-spark-ceo-podcast/id1142242316?mt=2&uo=4\", \"feedUrl\":\"http://feeds.soundcloud.com/users/soundcloud:users:58541313/sounds.rss\", \"trackViewUrl\":\"https://itunes.apple.com/us/podcast/ann-arbor-spark-ceo-podcast/id1142242316?mt=2&uo=4\", \"artworkUrl30\":\"https://is5-ssl.mzstatic.com/image/thumb/Music62/v4/8c/c5/6d/8cc56d2d-17e8-9d2a-50e2-e0352a3d79fc/source/30x30bb.jpg\", \"artworkUrl60\":\"https://is5-ssl.mzstatic.com/image/thumb/Music62/v4/8c/c5/6d/8cc56d2d-17e8-9d2a-50e2-e0352a3d79fc/source/60x60bb.jpg\", \"artworkUrl100\":\"https://is5-ssl.mzstatic.com/image/thumb/Music62/v4/8c/c5/6d/8cc56d2d-17e8-9d2a-50e2-e0352a3d79fc/source/100x100bb.jpg\", \"collectionPrice\":0.00, \"trackPrice\":0.00, \"trackRentalPrice\":0, \"collectionHdPrice\":0, \"trackHdPrice\":0, \"trackHdRentalPrice\":0, \"releaseDate\":\"2018-12-05T19:07:00Z\", \"collectionExplicitness\":\"cleaned\", \"trackExplicitness\":\"cleaned\", \"trackCount\":52, \"country\":\"USA\", \"currency\":\"USD\", \"primaryGenreName\":\"Business\", \"contentAdvisoryRating\":\"Clean\", \"artworkUrl600\":\"https://is5-ssl.mzstatic.com/image/thumb/Music62/v4/8c/c5/6d/8cc56d2d-17e8-9d2a-50e2-e0352a3d79fc/source/600x600bb.jpg\", \"genreIds\":[\"1321\", \"26\"], \"genres\":[\"Business\", \"Podcasts\"]}, \n{\"wrapperType\":\"track\", \"kind\":\"podcast\", \"collectionId\":908482513, \"trackId\":908482513, \"artistName\":\"Bethel AME Church - Ann Arbor\", \"collectionName\":\"Bethel AME - Ann Arbor\", \"trackName\":\"Bethel AME - Ann Arbor\", \"collectionCensoredName\":\"Bethel AME - Ann Arbor\", \"trackCensoredName\":\"Bethel AME - Ann Arbor\", \"collectionViewUrl\":\"https://itunes.apple.com/us/podcast/bethel-ame-ann-arbor/id908482513?mt=2&uo=4\", \"feedUrl\":\"http://www.bethelameannarbor.org/podcast-bethel-ame/\", \"trackViewUrl\":\"https://itunes.apple.com/us/podcast/bethel-ame-ann-arbor/id908482513?mt=2&uo=4\", \"artworkUrl30\":\"https://is5-ssl.mzstatic.com/image/thumb/Music1/v4/0d/96/a8/0d96a8c7-dd77-581f-3b31-e2064b110757/source/30x30bb.jpg\", \"artworkUrl60\":\"https://is5-ssl.mzstatic.com/image/thumb/Music1/v4/0d/96/a8/0d96a8c7-dd77-581f-3b31-e2064b110757/source/60x60bb.jpg\", \"artworkUrl100\":\"https://is5-ssl.mzstatic.com/image/thumb/Music1/v4/0d/96/a8/0d96a8c7-dd77-581f-3b31-e2064b110757/source/100x100bb.jpg\", \"collectionPrice\":0.00, \"trackPrice\":0.00, \"trackRentalPrice\":0, \"collectionHdPrice\":0, \"trackHdPrice\":0, \"trackHdRentalPrice\":0, \"releaseDate\":\"2016-11-09T12:32:00Z\", \"collectionExplicitness\":\"notExplicit\", \"trackExplicitness\":\"notExplicit\", \"trackCount\":46, \"country\":\"USA\", \"currency\":\"USD\", \"primaryGenreName\":\"Christianity\", \"artworkUrl600\":\"https://is5-ssl.mzstatic.com/image/thumb/Music1/v4/0d/96/a8/0d96a8c7-dd77-581f-3b31-e2064b110757/source/600x600bb.jpg\", \"genreIds\":[\"1439\", \"26\", \"1314\"], \"genres\":[\"Christianity\", \"Podcasts\", \"Religion & Spirituality\"]}, \n{\"wrapperType\":\"track\", \"kind\":\"podcast\", \"collectionId\":717203964, \"trackId\":717203964, \"artistName\":\"Ann Arbor West Side UMC\", \"collectionName\":\"Ann Arbor West Side UMC Sermons\", \"trackName\":\"Ann Arbor West Side UMC Sermons\", \"collectionCensoredName\":\"Ann Arbor West Side UMC Sermons\", \"trackCensoredName\":\"Ann Arbor West Side UMC Sermons\", \"collectionViewUrl\":\"https://itunes.apple.com/us/podcast/ann-arbor-west-side-umc-sermons/id717203964?mt=2&uo=4\", \"feedUrl\":\"http://www.rbb.me/wsumc/podcast.xml\", \"trackViewUrl\":\"https://itunes.apple.com/us/podcast/ann-arbor-west-side-umc-sermons/id717203964?mt=2&uo=4\", \"artworkUrl30\":\"https://is2-ssl.mzstatic.com/image/thumb/Music71/v4/8e/78/e7/8e78e713-66b8-2bf0-c80a-ef7d230e5ee4/source/30x30bb.jpg\", \"artworkUrl60\":\"https://is2-ssl.mzstatic.com/image/thumb/Music71/v4/8e/78/e7/8e78e713-66b8-2bf0-c80a-ef7d230e5ee4/source/60x60bb.jpg\", \"artworkUrl100\":\"https://is2-ssl.mzstatic.com/image/thumb/Music71/v4/8e/78/e7/8e78e713-66b8-2bf0-c80a-ef7d230e5ee4/source/100x100bb.jpg\", \"collectionPrice\":0.00, \"trackPrice\":0.00, \"trackRentalPrice\":0, \"collectionHdPrice\":0, \"trackHdPrice\":0, \"trackHdRentalPrice\":0, \"releaseDate\":\"2015-12-13T15:30:00Z\", \"collectionExplicitness\":\"cleaned\", \"trackExplicitness\":\"cleaned\", \"trackCount\":115, \"country\":\"USA\", \"currency\":\"USD\", \"primaryGenreName\":\"Christianity\", \"contentAdvisoryRating\":\"Clean\", \"artworkUrl600\":\"https://is2-ssl.mzstatic.com/image/thumb/Music71/v4/8e/78/e7/8e78e713-66b8-2bf0-c80a-ef7d230e5ee4/source/600x600bb.jpg\", \"genreIds\":[\"1439\", \"26\", \"1314\"], \"genres\":[\"Christianity\", \"Podcasts\", \"Religion & Spirituality\"]}, \n{\"wrapperType\":\"track\", \"kind\":\"podcast\", \"collectionId\":372059430, \"trackId\":372059430, \"artistName\":\"Guest\", \"collectionName\":\"Grace Ann Arbor Podcast\", \"trackName\":\"Grace Ann Arbor Podcast\", \"collectionCensoredName\":\"Grace Ann Arbor Podcast\", \"trackCensoredName\":\"Grace Ann Arbor Podcast\", \"collectionViewUrl\":\"https://itunes.apple.com/us/podcast/grace-ann-arbor-podcast/id372059430?mt=2&uo=4\", \"feedUrl\":\"http://www.sermonconnect.com/podcasts.php?church=118882&feedid=20100510010558283064\", \"trackViewUrl\":\"https://itunes.apple.com/us/podcast/grace-ann-arbor-podcast/id372059430?mt=2&uo=4\", \"artworkUrl30\":\"https://is1-ssl.mzstatic.com/image/thumb/Music/v4/34/49/8f/34498f83-0457-6c55-e638-8e22a4326f2f/source/30x30bb.jpg\", \"artworkUrl60\":\"https://is1-ssl.mzstatic.com/image/thumb/Music/v4/34/49/8f/34498f83-0457-6c55-e638-8e22a4326f2f/source/60x60bb.jpg\", \"artworkUrl100\":\"https://is1-ssl.mzstatic.com/image/thumb/Music/v4/34/49/8f/34498f83-0457-6c55-e638-8e22a4326f2f/source/100x100bb.jpg\", \"collectionPrice\":0.00, \"trackPrice\":0.00, \"trackRentalPrice\":0, \"collectionHdPrice\":0, \"trackHdPrice\":0, \"trackHdRentalPrice\":0, \"releaseDate\":\"2010-05-23T00:00:00Z\", \"collectionExplicitness\":\"cleaned\", \"trackExplicitness\":\"cleaned\", \"trackCount\":6, \"country\":\"USA\", \"currency\":\"USD\", \"primaryGenreName\":\"Christianity\", \"contentAdvisoryRating\":\"Clean\", \"artworkUrl600\":\"https://is1-ssl.mzstatic.com/image/thumb/Music/v4/34/49/8f/34498f83-0457-6c55-e638-8e22a4326f2f/source/600x600bb.jpg\", \"genreIds\":[\"1439\", \"26\", \"1314\"], \"genres\":[\"Christianity\", \"Podcasts\", \"Religion & Spirituality\"]}, \n{\"wrapperType\":\"track\", \"kind\":\"podcast\", \"collectionId\":1434246896, \"trackId\":1434246896, \"artistName\":\"Mosaic Church of Ann Arbor\", \"collectionName\":\"Mosaic Church of Ann Arbor\", \"trackName\":\"Mosaic Church of Ann Arbor\", \"collectionCensoredName\":\"Mosaic Church of Ann Arbor\", \"trackCensoredName\":\"Mosaic Church of Ann Arbor\", \"collectionViewUrl\":\"https://itunes.apple.com/us/podcast/mosaic-church-of-ann-arbor/id1434246896?mt=2&uo=4\", \"feedUrl\":\"https://mosaica2.squarespace.com/podcast?format=rss\", \"trackViewUrl\":\"https://itunes.apple.com/us/podcast/mosaic-church-of-ann-arbor/id1434246896?mt=2&uo=4\", \"artworkUrl30\":\"https://is2-ssl.mzstatic.com/image/thumb/Music118/v4/eb/d9/cb/ebd9cb93-6f62-5010-660a-8634fbbb82d8/source/30x30bb.jpg\", \"artworkUrl60\":\"https://is2-ssl.mzstatic.com/image/thumb/Music118/v4/eb/d9/cb/ebd9cb93-6f62-5010-660a-8634fbbb82d8/source/60x60bb.jpg\", \"artworkUrl100\":\"https://is2-ssl.mzstatic.com/image/thumb/Music118/v4/eb/d9/cb/ebd9cb93-6f62-5010-660a-8634fbbb82d8/source/100x100bb.jpg\", \"collectionPrice\":0.00, \"trackPrice\":0.00, \"trackRentalPrice\":0, \"collectionHdPrice\":0, \"trackHdPrice\":0, \"trackHdRentalPrice\":0, \"releaseDate\":\"2018-11-05T01:36:00Z\", \"collectionExplicitness\":\"cleaned\", \"trackExplicitness\":\"cleaned\", \"trackCount\":39, \"country\":\"USA\", \"currency\":\"USD\", \"primaryGenreName\":\"Christianity\", \"contentAdvisoryRating\":\"Clean\", \"artworkUrl600\":\"https://is2-ssl.mzstatic.com/image/thumb/Music118/v4/eb/d9/cb/ebd9cb93-6f62-5010-660a-8634fbbb82d8/source/600x600bb.jpg\", \"genreIds\":[\"1439\", \"26\", \"1314\"], \"genres\":[\"Christianity\", \"Podcasts\", \"Religion & Spirituality\"]}, \n{\"wrapperType\":\"track\", \"kind\":\"podcast\", \"collectionId\":434700786, \"trackId\":434700786, \"artistName\":\"Rebecca Rhodes for the City of Ann Arbor\", \"collectionName\":\"A2 City News\", \"trackName\":\"A2 City News\", \"collectionCensoredName\":\"A2 City News\", \"trackCensoredName\":\"A2 City News\", \"collectionViewUrl\":\"https://itunes.apple.com/us/podcast/a2-city-news/id434700786?mt=2&uo=4\", \"feedUrl\":\"http://feeds.feedburner.com/A2CityNews\", \"trackViewUrl\":\"https://itunes.apple.com/us/podcast/a2-city-news/id434700786?mt=2&uo=4\", \"artworkUrl30\":\"https://is2-ssl.mzstatic.com/image/thumb/Music6/v4/3b/93/64/3b9364d4-e06b-49f3-1811-6f26bbf08447/source/30x30bb.jpg\", \"artworkUrl60\":\"https://is2-ssl.mzstatic.com/image/thumb/Music6/v4/3b/93/64/3b9364d4-e06b-49f3-1811-6f26bbf08447/source/60x60bb.jpg\", \"artworkUrl100\":\"https://is2-ssl.mzstatic.com/image/thumb/Music6/v4/3b/93/64/3b9364d4-e06b-49f3-1811-6f26bbf08447/source/100x100bb.jpg\", \"collectionPrice\":0.00, \"trackPrice\":0.00, \"trackRentalPrice\":0, \"collectionHdPrice\":0, \"trackHdPrice\":0, \"trackHdRentalPrice\":0, \"releaseDate\":\"2012-03-08T18:54:00Z\", \"collectionExplicitness\":\"cleaned\", \"trackExplicitness\":\"cleaned\", \"trackCount\":13, \"country\":\"USA\", \"currency\":\"USD\", \"primaryGenreName\":\"Local\", \"contentAdvisoryRating\":\"Clean\", \"artworkUrl600\":\"https://is2-ssl.mzstatic.com/image/thumb/Music6/v4/3b/93/64/3b9364d4-e06b-49f3-1811-6f26bbf08447/source/600x600bb.jpg\", \"genreIds\":[\"1475\", \"26\", \"1325\"], \"genres\":[\"Local\", \"Podcasts\", \"Government & Organizations\"]}, \n{\"wrapperType\":\"track\", \"kind\":\"podcast\", \"collectionId\":79905145, \"trackId\":79905145, \"artistName\":\"First Presbyterian Church of Ann Arbor\", \"collectionName\":\"Sermons from First Pres\", \"trackName\":\"Sermons from First Pres\", \"collectionCensoredName\":\"Sermons from First Pres\", \"trackCensoredName\":\"Sermons from First Pres\", \"collectionViewUrl\":\"https://itunes.apple.com/us/podcast/sermons-from-first-pres/id79905145?mt=2&uo=4\", \"feedUrl\":\"http://feeds.feedburner.com/sermonsfromfirstpres?q=podcasts.xml\", \"trackViewUrl\":\"https://itunes.apple.com/us/podcast/sermons-from-first-pres/id79905145?mt=2&uo=4\", \"artworkUrl30\":\"https://is4-ssl.mzstatic.com/image/thumb/Music62/v4/2b/5b/24/2b5b2400-8251-bdca-f02f-70473380d2fc/source/30x30bb.jpg\", \"artworkUrl60\":\"https://is4-ssl.mzstatic.com/image/thumb/Music62/v4/2b/5b/24/2b5b2400-8251-bdca-f02f-70473380d2fc/source/60x60bb.jpg\", \"artworkUrl100\":\"https://is4-ssl.mzstatic.com/image/thumb/Music62/v4/2b/5b/24/2b5b2400-8251-bdca-f02f-70473380d2fc/source/100x100bb.jpg\", \"collectionPrice\":0.00, \"trackPrice\":0.00, \"trackRentalPrice\":0, \"collectionHdPrice\":0, \"trackHdPrice\":0, \"trackHdRentalPrice\":0, \"releaseDate\":\"2018-11-25T05:00:00Z\", \"collectionExplicitness\":\"cleaned\", \"trackExplicitness\":\"cleaned\", \"trackCount\":279, \"country\":\"USA\", \"currency\":\"USD\", \"primaryGenreName\":\"Christianity\", \"contentAdvisoryRating\":\"Clean\", \"artworkUrl600\":\"https://is4-ssl.mzstatic.com/image/thumb/Music62/v4/2b/5b/24/2b5b2400-8251-bdca-f02f-70473380d2fc/source/600x600bb.jpg\", \"genreIds\":[\"1439\", \"26\", \"1314\"], \"genres\":[\"Christianity\", \"Podcasts\", \"Religion & Spirituality\"]}, \n{\"wrapperType\":\"track\", \"kind\":\"podcast\", \"collectionId\":1105933828, \"trackId\":1105933828, \"artistName\":\"Redeemer Ann Arbor\", \"collectionName\":\"Redeemer Ann Arbor\", \"trackName\":\"Redeemer Ann Arbor\", \"collectionCensoredName\":\"Redeemer Ann Arbor\", \"trackCensoredName\":\"Redeemer Ann Arbor\", \"collectionViewUrl\":\"https://itunes.apple.com/us/podcast/redeemer-ann-arbor/id1105933828?mt=2&uo=4\", \"feedUrl\":\"https://www.redeemera2.org/archive/feed/\", \"trackViewUrl\":\"https://itunes.apple.com/us/podcast/redeemer-ann-arbor/id1105933828?mt=2&uo=4\", \"artworkUrl30\":\"https://is5-ssl.mzstatic.com/image/thumb/Music71/v4/64/42/69/6442695f-df87-3332-52d3-fce473e3c1d4/source/30x30bb.jpg\", \"artworkUrl60\":\"https://is5-ssl.mzstatic.com/image/thumb/Music71/v4/64/42/69/6442695f-df87-3332-52d3-fce473e3c1d4/source/60x60bb.jpg\", \"artworkUrl100\":\"https://is5-ssl.mzstatic.com/image/thumb/Music71/v4/64/42/69/6442695f-df87-3332-52d3-fce473e3c1d4/source/100x100bb.jpg\", \"collectionPrice\":0.00, \"trackPrice\":0.00, \"trackRentalPrice\":0, \"collectionHdPrice\":0, \"trackHdPrice\":0, \"trackHdRentalPrice\":0, \"releaseDate\":\"2018-12-16T00:00:00Z\", \"collectionExplicitness\":\"cleaned\", \"trackExplicitness\":\"cleaned\", \"trackCount\":10, \"country\":\"USA\", \"currency\":\"USD\", \"primaryGenreName\":\"Christianity\", \"contentAdvisoryRating\":\"Clean\", \"artworkUrl600\":\"https://is5-ssl.mzstatic.com/image/thumb/Music71/v4/64/42/69/6442695f-df87-3332-52d3-fce473e3c1d4/source/600x600bb.jpg\", \"genreIds\":[\"1439\", \"26\", \"1314\"], \"genres\":[\"Christianity\", \"Podcasts\", \"Religion & Spirituality\"]}, \n{\"wrapperType\":\"track\", \"kind\":\"podcast\", \"collectionId\":1089272384, \"trackId\":1089272384, \"artistName\":\"Ann Arbor District Library\", \"collectionName\":\"Ann Arbor Stories | Ann Arbor District Library\", \"trackName\":\"Ann Arbor Stories | Ann Arbor District Library\", \"collectionCensoredName\":\"Ann Arbor Stories | Ann Arbor District Library\", \"trackCensoredName\":\"Ann Arbor Stories | Ann Arbor District Library\", \"collectionViewUrl\":\"https://itunes.apple.com/us/podcast/ann-arbor-stories-ann-arbor-district-library/id1089272384?mt=2&uo=4\", \"feedUrl\":\"https://aadl.org/annarborstories/feed\", \"trackViewUrl\":\"https://itunes.apple.com/us/podcast/ann-arbor-stories-ann-arbor-district-library/id1089272384?mt=2&uo=4\", \"artworkUrl30\":\"https://is2-ssl.mzstatic.com/image/thumb/Music128/v4/83/dd/44/83dd4483-afda-e94d-95f7-78e0fc8e139b/source/30x30bb.jpg\", \"artworkUrl60\":\"https://is2-ssl.mzstatic.com/image/thumb/Music128/v4/83/dd/44/83dd4483-afda-e94d-95f7-78e0fc8e139b/source/60x60bb.jpg\", \"artworkUrl100\":\"https://is2-ssl.mzstatic.com/image/thumb/Music128/v4/83/dd/44/83dd4483-afda-e94d-95f7-78e0fc8e139b/source/100x100bb.jpg\", \"collectionPrice\":0.00, \"trackPrice\":0.00, \"trackRentalPrice\":0, \"collectionHdPrice\":0, \"trackHdPrice\":0, \"trackHdRentalPrice\":0, \"releaseDate\":\"2018-04-12T10:30:00Z\", \"collectionExplicitness\":\"cleaned\", \"trackExplicitness\":\"cleaned\", \"trackCount\":47, \"country\":\"USA\", \"currency\":\"USD\", \"primaryGenreName\":\"Arts\", \"contentAdvisoryRating\":\"Clean\", \"artworkUrl600\":\"https://is2-ssl.mzstatic.com/image/thumb/Music128/v4/83/dd/44/83dd4483-afda-e94d-95f7-78e0fc8e139b/source/600x600bb.jpg\", \"genreIds\":[\"1301\", \"26\", \"1324\", \"1462\", \"1320\", \"1325\", \"1474\"], \"genres\":[\"Arts\", \"Podcasts\", \"Society & Culture\", \"History\", \"Places & Travel\", \"Government & Organizations\", \"Regional\"]}, \n{\"wrapperType\":\"track\", \"kind\":\"podcast\", \"collectionId\":704835029, \"trackId\":704835029, \"artistName\":\"admin@antiocha2.org (admin@antiocha2.org)\", \"collectionName\":\"Antioch Ann Arbor\", \"trackName\":\"Antioch Ann Arbor\", \"collectionCensoredName\":\"Antioch Ann Arbor\", \"trackCensoredName\":\"Antioch Ann Arbor\", \"collectionViewUrl\":\"https://itunes.apple.com/us/podcast/antioch-ann-arbor/id704835029?mt=2&uo=4\", \"feedUrl\":\"http://antiocha2.libsyn.com/rss\", \"trackViewUrl\":\"https://itunes.apple.com/us/podcast/antioch-ann-arbor/id704835029?mt=2&uo=4\", \"artworkUrl30\":\"https://is3-ssl.mzstatic.com/image/thumb/Music128/v4/4d/ec/6a/4dec6afc-2713-d837-b5fd-bf841ac7de11/source/30x30bb.jpg\", \"artworkUrl60\":\"https://is3-ssl.mzstatic.com/image/thumb/Music128/v4/4d/ec/6a/4dec6afc-2713-d837-b5fd-bf841ac7de11/source/60x60bb.jpg\", \"artworkUrl100\":\"https://is3-ssl.mzstatic.com/image/thumb/Music128/v4/4d/ec/6a/4dec6afc-2713-d837-b5fd-bf841ac7de11/source/100x100bb.jpg\", \"collectionPrice\":0.00, \"trackPrice\":0.00, \"trackRentalPrice\":0, \"collectionHdPrice\":0, \"trackHdPrice\":0, \"trackHdRentalPrice\":0, \"releaseDate\":\"2018-12-09T17:00:00Z\", \"collectionExplicitness\":\"cleaned\", \"trackExplicitness\":\"cleaned\", \"trackCount\":20, \"country\":\"USA\", \"currency\":\"USD\", \"primaryGenreName\":\"Christianity\", \"contentAdvisoryRating\":\"Clean\", \"artworkUrl600\":\"https://is3-ssl.mzstatic.com/image/thumb/Music128/v4/4d/ec/6a/4dec6afc-2713-d837-b5fd-bf841ac7de11/source/600x600bb.jpg\", \"genreIds\":[\"1439\", \"26\", \"1314\", \"1444\"], \"genres\":[\"Christianity\", \"Podcasts\", \"Religion & Spirituality\", \"Spirituality\"]}, \n{\"wrapperType\":\"track\", \"kind\":\"podcast\", \"collectionId\":1137457133, \"trackId\":1137457133, \"artistName\":\"Unknown\", \"collectionName\":\"Zion Lutheran Ann Arbor\", \"trackName\":\"Zion Lutheran Ann Arbor\", \"collectionCensoredName\":\"Zion Lutheran Ann Arbor\", \"trackCensoredName\":\"Zion Lutheran Ann Arbor\", \"collectionViewUrl\":\"https://itunes.apple.com/us/podcast/zion-lutheran-ann-arbor/id1137457133?mt=2&uo=4\", \"feedUrl\":\"http://recordings.zlc-aa.org/rss.php\", \"trackViewUrl\":\"https://itunes.apple.com/us/podcast/zion-lutheran-ann-arbor/id1137457133?mt=2&uo=4\", \"artworkUrl30\":\"https://is2-ssl.mzstatic.com/image/thumb/Music71/v4/f9/83/1b/f9831b28-5ee7-56a4-4ef8-415b14dd6d00/source/30x30bb.jpg\", \"artworkUrl60\":\"https://is2-ssl.mzstatic.com/image/thumb/Music71/v4/f9/83/1b/f9831b28-5ee7-56a4-4ef8-415b14dd6d00/source/60x60bb.jpg\", \"artworkUrl100\":\"https://is2-ssl.mzstatic.com/image/thumb/Music71/v4/f9/83/1b/f9831b28-5ee7-56a4-4ef8-415b14dd6d00/source/100x100bb.jpg\", \"collectionPrice\":0.00, \"trackPrice\":0.00, \"trackRentalPrice\":0, \"collectionHdPrice\":0, \"trackHdPrice\":0, \"trackHdRentalPrice\":0, \"releaseDate\":\"2018-03-11T19:17:00Z\", \"collectionExplicitness\":\"cleaned\", \"trackExplicitness\":\"cleaned\", \"trackCount\":116, \"country\":\"USA\", \"currency\":\"USD\", \"primaryGenreName\":\"Christianity\", \"contentAdvisoryRating\":\"Clean\", \"artworkUrl600\":\"https://is2-ssl.mzstatic.com/image/thumb/Music71/v4/f9/83/1b/f9831b28-5ee7-56a4-4ef8-415b14dd6d00/source/600x600bb.jpg\", \"genreIds\":[\"1439\", \"26\", \"1314\"], \"genres\":[\"Christianity\", \"Podcasts\", \"Religion & Spirituality\"]}, \n{\"wrapperType\":\"track\", \"kind\":\"podcast\", \"collectionId\":1085248390, \"trackId\":1085248390, \"artistName\":\"Ann Arbor District Library\", \"collectionName\":\"Martin Bandyke Under Covers | Ann Arbor District Library\", \"trackName\":\"Martin Bandyke Under Covers | Ann Arbor District Library\", \"collectionCensoredName\":\"Martin Bandyke Under Covers | Ann Arbor District Library\", \"trackCensoredName\":\"Martin Bandyke Under Covers | Ann Arbor District Library\", \"collectionViewUrl\":\"https://itunes.apple.com/us/podcast/martin-bandyke-under-covers-ann-arbor-district-library/id1085248390?mt=2&uo=4\", \"feedUrl\":\"https://aadl.org/undercovers/feed\", \"trackViewUrl\":\"https://itunes.apple.com/us/podcast/martin-bandyke-under-covers-ann-arbor-district-library/id1085248390?mt=2&uo=4\", \"artworkUrl30\":\"https://is1-ssl.mzstatic.com/image/thumb/Music125/v4/13/33/ce/1333ce50-b06c-821d-2266-7cc30476a53e/source/30x30bb.jpg\", \"artworkUrl60\":\"https://is1-ssl.mzstatic.com/image/thumb/Music125/v4/13/33/ce/1333ce50-b06c-821d-2266-7cc30476a53e/source/60x60bb.jpg\", \"artworkUrl100\":\"https://is1-ssl.mzstatic.com/image/thumb/Music125/v4/13/33/ce/1333ce50-b06c-821d-2266-7cc30476a53e/source/100x100bb.jpg\", \"collectionPrice\":0.00, \"trackPrice\":0.00, \"trackRentalPrice\":0, \"collectionHdPrice\":0, \"trackHdPrice\":0, \"trackHdRentalPrice\":0, \"releaseDate\":\"2018-12-04T20:00:00Z\", \"collectionExplicitness\":\"cleaned\", \"trackExplicitness\":\"cleaned\", \"trackCount\":53, \"country\":\"USA\", \"currency\":\"USD\", \"primaryGenreName\":\"Literature\", \"contentAdvisoryRating\":\"Clean\", \"artworkUrl600\":\"https://is1-ssl.mzstatic.com/image/thumb/Music125/v4/13/33/ce/1333ce50-b06c-821d-2266-7cc30476a53e/source/600x600bb.jpg\", \"genreIds\":[\"1401\", \"26\", \"1301\", \"1324\", \"1406\", \"1405\", \"1311\"], \"genres\":[\"Literature\", \"Podcasts\", \"Arts\", \"Society & Culture\", \"Visual Arts\", \"Performing Arts\", \"News & Politics\"]}, \n{\"wrapperType\":\"track\", \"kind\":\"podcast\", \"collectionId\":1440591293, \"trackId\":1440591293, \"artistName\":\"Rev. Jeremy Byrd\", \"collectionName\":\"Christ Church Ann Arbor\", \"trackName\":\"Christ Church Ann Arbor\", \"collectionCensoredName\":\"Christ Church Ann Arbor\", \"trackCensoredName\":\"Christ Church Ann Arbor\", \"collectionViewUrl\":\"https://itunes.apple.com/us/podcast/christ-church-ann-arbor/id1440591293?mt=2&uo=4\", \"feedUrl\":\"http://www.sermonaudio.com/rss_source.asp?sourceid=christchurcha2\", \"trackViewUrl\":\"https://itunes.apple.com/us/podcast/christ-church-ann-arbor/id1440591293?mt=2&uo=4\", \"artworkUrl30\":\"https://is4-ssl.mzstatic.com/image/thumb/Music118/v4/1f/ba/28/1fba287a-b84f-5beb-b945-6e06b40325f7/source/30x30bb.jpg\", \"artworkUrl60\":\"https://is4-ssl.mzstatic.com/image/thumb/Music118/v4/1f/ba/28/1fba287a-b84f-5beb-b945-6e06b40325f7/source/60x60bb.jpg\", \"artworkUrl100\":\"https://is4-ssl.mzstatic.com/image/thumb/Music118/v4/1f/ba/28/1fba287a-b84f-5beb-b945-6e06b40325f7/source/100x100bb.jpg\", \"collectionPrice\":0.00, \"trackPrice\":0.00, \"trackRentalPrice\":0, \"collectionHdPrice\":0, \"trackHdPrice\":0, \"trackHdRentalPrice\":0, \"releaseDate\":\"2018-12-10T21:23:00Z\", \"collectionExplicitness\":\"cleaned\", \"trackExplicitness\":\"cleaned\", \"trackCount\":8, \"country\":\"USA\", \"currency\":\"USD\", \"primaryGenreName\":\"Christianity\", \"contentAdvisoryRating\":\"Clean\", \"artworkUrl600\":\"https://is4-ssl.mzstatic.com/image/thumb/Music118/v4/1f/ba/28/1fba287a-b84f-5beb-b945-6e06b40325f7/source/600x600bb.jpg\", \"genreIds\":[\"1439\", \"26\", \"1314\"], \"genres\":[\"Christianity\", \"Podcasts\", \"Religion & Spirituality\"]}]\n}",
      "http://www.omdbapi.com/r-json_t-Date Night": "{\"Title\":\"Date Night\",\"Year\":\"2010\",\"Rated\":\"PG-13\",\"Released\":\"09 Apr 2010\",\"Runtime\":\"88 min\",\"Genre\":\"Comedy, Crime, Romance, Thriller\",\"Director\":\"Shawn Levy\",\"Writer\":\"Josh Klausner\",\"Actors\":\"Steve Carell, Tina Fey, Mark Wahlberg, Taraji P. Henson\",\"Plot\":\"In New York City, a case of mistaken identity turns a bored married couple's attempt at a glamorous and romantic evening into something more thrilling and dangerous.\",\"Language\":\"English, Hebrew\",\"Country\":\"USA\",\"Awards\":\"4 wins & 8 nominations.\",\"Poster\":\"https://m.media-amazon.com/images/M/MV5BODgwMjM2ODE4M15BMl5BanBnXkFtZTcwMTU2MDcyMw@@._V1_SX300.jpg\",\"Ratings\":[{\"Source\":\"Internet Movie Database\",\"Value\":\"6.3/10\"},{\"Source\":\"Rotten Tomatoes\",\"Value\":\"67%\"},{\"Source\":\"Metacritic\",\"Value\":\"56/100\"}],\"Metascore\":\"56\",\"imdbRating\":\"6.3\",\"imdbVotes\":\"141,482\",\"imdbID\":\"tt1279935\",\"Type\":\"movie\",\"DVD\":\"10 Aug 2010\",\"BoxOffice\":\"$98,710,290\",\"Production\":\"20th Century Fox\",\"Website\":\"http://www.datenight-movie.com/\",\"Response\":\"True\"}",
      "https://tastedive.com/api/similarlimit-5_q-Black Panther": "{\"Similar\": {\"Info\": [{\"Name\": \"Black Panther\", \"Type\": \"movie\"}], \"Results\": [{\"Name\": \"Captain Marvel\", \"Type\": \"movie\"}, {\"Name\": \"Avengers: Infinity War\", \"Type\": \"movie\"}, {\"Name\": \"Ant-Man And The Wasp\", \"Type\": \"movie\"}, {\"Name\": \"The Fate Of The Furious\", \"Type\": \"movie\"}, {\"Name\": \"Deadpool 2\", \"Type\": \"movie\"}]}}",
      "http://www.omdbapi.com/r-json_t-Avengers: Infinity War": "{\"Title\":\"Avengers: Infinity War\",\"Year\":\"2018\",\"Rated\":\"PG-13\",\"Released\":\"27 Apr 2018\",\"Runtime\":\"149 min\",\"Genre\":\"Action, Adventure, Fantasy, Sci-Fi\",\"Director\":\"Anthony Russo, Joe Russo\",\"Writer\":\"Christopher Markus (screenplay by), Stephen McFeely (screenplay by), Stan Lee (based on the Marvel comics by), Jack Kirby (based on the Marvel comics by), Joe Simon (Captain America created by), Jack Kirby (Captain America created by), Steve Englehart (Star-Lord created by), Steve Gan (Star-Lord created by), Bill Mantlo (Rocket Raccoon created by), Keith Giffen (Rocket Raccoon created by), Jim Starlin (Thanos,  Gamora and Drax created by), Stan Lee (Groot created by), Larry Lieber (Groot created by), Jack Kirby (Groot created by), Steve Englehart (Mantis created by), Don Heck (Mantis created by)\",\"Actors\":\"Robert Downey Jr., Chris Hemsworth, Mark Ruffalo, Chris Evans\",\"Plot\":\"The Avengers and their allies must be willing to sacrifice all in an attempt to defeat the powerful Thanos before his blitz of devastation and ruin puts an end to the universe.\",\"Language\":\"English\",\"Country\":\"USA\",\"Awards\":\"N/A\",\"Poster\":\"https://m.media-amazon.com/images/M/MV5BMjMxNjY2MDU1OV5BMl5BanBnXkFtZTgwNzY1MTUwNTM@._V1_SX300.jpg\",\"Ratings\":[{\"Source\":\"Internet Movie Database\",\"Value\":\"8.5/10\"},{\"Source\":\"Rotten Tomatoes\",\"Value\":\"84%\"},{\"Source\":\"Metacritic\",\"Value\":\"68/100\"}],\"Metascore\":\"68\",\"imdbRating\":\"8.5\",\"imdbVotes\":\"525,758\",\"imdbID\":\"tt4154756\",\"Type\":\"movie\",\"DVD\":\"14 Aug 2018\",\"BoxOffice\":\"$664,987,816\",\"Production\":\"Walt Disney Pictures\",\"Website\":\"http://marvel.com/movies/movie/223/avengers_infinity_war\",\"Response\":\"True\"}",
      "http://www.omdbapi.com/r-json_t-Captain Marvel": "{\"Title\":\"Captain Marvel\",\"Year\":\"1962\",\"Rated\":\"N/A\",\"Released\":\"N/A\",\"Runtime\":\"10 min\",\"Genre\":\"Short, Action, Adventure\",\"Director\":\"Donald F. Glut\",\"Writer\":\"Donald F. Glut\",\"Actors\":\"Rich Hagopian, Dick Andersen, Bill Bungert, Donald F. Glut\",\"Plot\":\"Captain Marvel saves the world from an oncoming asteroid, but in doing so, comes to the attention of criminals.\",\"Language\":\"English\",\"Country\":\"USA\",\"Awards\":\"N/A\",\"Poster\":\"https://images-na.ssl-images-amazon.com/images/M/MV5BYWQwNTE2YzItNjViYy00ODdmLWI1MDctMGEyODNlYzI5ZDcyXkEyXkFqcGdeQXVyMzI2MDEwNA@@._V1_SX300.jpg\",\"Ratings\":[{\"Source\":\"Internet Movie Database\",\"Value\":\"6.1/10\"}],\"Metascore\":\"N/A\",\"imdbRating\":\"6.1\",\"imdbVotes\":\"27\",\"imdbID\":\"tt3959390\",\"Type\":\"movie\",\"DVD\":\"N/A\",\"BoxOffice\":\"N/A\",\"Production\":\"N/A\",\"Website\":\"N/A\",\"Response\":\"True\"}",
      "http://www.omdbapi.com/r-json_t-Pirates Of The Caribbean: On Stranger Tides": "{\"Title\":\"Pirates of the Caribbean: On Stranger Tides\",\"Year\":\"2011\",\"Rated\":\"PG-13\",\"Released\":\"20 May 2011\",\"Runtime\":\"136 min\",\"Genre\":\"Action, Adventure, Fantasy\",\"Director\":\"Rob Marshall\",\"Writer\":\"Ted Elliott (screenplay), Terry Rossio (screenplay), Ted Elliott (screen story), Terry Rossio (screen story), Ted Elliott (characters), Terry Rossio (characters), Stuart Beattie (characters), Jay Wolpert (characters), Tim Powers (novel)\",\"Actors\":\"Johnny Depp, Pen\u00e9lope Cruz, Geoffrey Rush, Ian McShane\",\"Plot\":\"Jack Sparrow and Barbossa embark on a quest to find the elusive fountain of youth, only to discover that Blackbeard and his daughter are after it too.\",\"Language\":\"English, Spanish\",\"Country\":\"USA, UK\",\"Awards\":\"2 wins & 31 nominations.\",\"Poster\":\"https://m.media-amazon.com/images/M/MV5BMjE5MjkwODI3Nl5BMl5BanBnXkFtZTcwNjcwMDk4NA@@._V1_SX300.jpg\",\"Ratings\":[{\"Source\":\"Internet Movie Database\",\"Value\":\"6.6/10\"},{\"Source\":\"Rotten Tomatoes\",\"Value\":\"32%\"},{\"Source\":\"Metacritic\",\"Value\":\"45/100\"}],\"Metascore\":\"45\",\"imdbRating\":\"6.6\",\"imdbVotes\":\"434,049\",\"imdbID\":\"tt1298650\",\"Type\":\"movie\",\"DVD\":\"18 Oct 2011\",\"BoxOffice\":\"$241,063,875\",\"Production\":\"Walt Disney Pictures\",\"Website\":\"http://Disney.com/Pirates\",\"Response\":\"True\"}",
      "http://www.omdbapi.com/r-json_t-Venom": "{\"Title\":\"Venom\",\"Year\":\"2018\",\"Rated\":\"N/A\",\"Released\":\"05 Oct 2018\",\"Runtime\":\"112 min\",\"Genre\":\"Action, Sci-Fi\",\"Director\":\"Ruben Fleischer\",\"Writer\":\"Jeff Pinkner (screenplay by), Scott Rosenberg (screenplay by), Kelly Marcel (screenplay by), Jeff Pinkner (screen story by), Scott Rosenberg (screen story by), Todd McFarlane (Marvel's Venom Character created by), David Michelinie (Marvel's Venom Character created by)\",\"Actors\":\"Tom Hardy, Michelle Williams, Riz Ahmed, Scott Haze\",\"Plot\":\"When Eddie Brock acquires the powers of a symbiote, he will have to release his alter-ego \\\"Venom\\\" to save his life.\",\"Language\":\"English\",\"Country\":\"USA, China\",\"Awards\":\"N/A\",\"Poster\":\"https://m.media-amazon.com/images/M/MV5BNzAwNzUzNjY4MV5BMl5BanBnXkFtZTgwMTQ5MzM0NjM@._V1_SX300.jpg\",\"Ratings\":[{\"Source\":\"Internet Movie Database\",\"Value\":\"6.9/10\"},{\"Source\":\"Metacritic\",\"Value\":\"35/100\"}],\"Metascore\":\"35\",\"imdbRating\":\"6.9\",\"imdbVotes\":\"154,911\",\"imdbID\":\"tt1270797\",\"Type\":\"movie\",\"DVD\":\"18 Jun 2013\",\"BoxOffice\":\"N/A\",\"Production\":\"Vis\",\"Website\":\"N/A\",\"Response\":\"True\"}",
      "https://tastedive.com/api/similarlimit-5_q-Tony Bennett_type-movies": "{\"Similar\": {\"Info\": [{\"Name\": \"Tony Bennett\", \"Type\": \"music\"}], \"Results\": [{\"Name\": \"The Startup Kids\", \"Type\": \"movie\"}, {\"Name\": \"Charlie Chaplin\", \"Type\": \"movie\"}, {\"Name\": \"Venus In Fur\", \"Type\": \"movie\"}, {\"Name\": \"Loving\", \"Type\": \"movie\"}, {\"Name\": \"The African Queen\", \"Type\": \"movie\"}]}}",
      "http://www.omdbapi.com/r-json_t-Eyyvah Eyvah": "{\"Title\":\"Eyyvah Eyvah\",\"Year\":\"2010\",\"Rated\":\"N/A\",\"Released\":\"26 Feb 2010\",\"Runtime\":\"104 min\",\"Genre\":\"Comedy\",\"Director\":\"Hakan Alg\u00fcl\",\"Writer\":\"Ata Demirer\",\"Actors\":\"Demet Akbag, Ata Demirer, Tanju Tuncel, Sehsuvar Aktas\",\"Plot\":\"A clarinet player goes on a search for his real father, teaming up with an eccentric nightclub singer.\",\"Language\":\"Turkish\",\"Country\":\"Turkey\",\"Awards\":\"N/A\",\"Poster\":\"https://images-na.ssl-images-amazon.com/images/M/MV5BMjA3MDM3OTAxNF5BMl5BanBnXkFtZTcwNTcyODY5Ng@@._V1_SX300.jpg\",\"Ratings\":[{\"Source\":\"Internet Movie Database\",\"Value\":\"7.1/10\"}],\"Metascore\":\"N/A\",\"imdbRating\":\"7.1\",\"imdbVotes\":\"15,395\",\"imdbID\":\"tt1594918\",\"Type\":\"movie\",\"DVD\":\"N/A\",\"BoxOffice\":\"N/A\",\"Production\":\"Kinostar\",\"Website\":\"N/A\",\"Response\":\"True\"}",
      "http://www.omdbapi.com/r-json_t-Deadpool 2": "{\"Title\":\"Deadpool 2\",\"Year\":\"2018\",\"Rated\":\"R\",\"Released\":\"18 May 2018\",\"Runtime\":\"119 min\",\"Genre\":\"Action, Adventure, Comedy, Sci-Fi\",\"Director\":\"David Leitch\",\"Writer\":\"Rhett Reese, Paul Wernick, Ryan Reynolds\",\"Actors\":\"Ryan Reynolds, Josh Brolin, Morena Baccarin, Julian Dennison\",\"Plot\":\"Foul-mouthed mutant mercenary Wade Wilson (AKA. Deadpool), brings together a team of fellow mutant rogues to protect a young boy with supernatural abilities from the brutal, time-traveling cyborg, Cable.\",\"Language\":\"English, Cantonese, Spanish, Russian\",\"Country\":\"USA\",\"Awards\":\"N/A\",\"Poster\":\"https://m.media-amazon.com/images/M/MV5BNjk1Njk3YjctMmMyYS00Y2I4LThhMzktN2U0MTMyZTFlYWQ5XkEyXkFqcGdeQXVyODM2ODEzMDA@._V1_SX300.jpg\",\"Ratings\":[{\"Source\":\"Internet Movie Database\",\"Value\":\"7.8/10\"},{\"Source\":\"Rotten Tomatoes\",\"Value\":\"83%\"},{\"Source\":\"Metacritic\",\"Value\":\"66/100\"}],\"Metascore\":\"66\",\"imdbRating\":\"7.8\",\"imdbVotes\":\"311,435\",\"imdbID\":\"tt5463162\",\"Type\":\"movie\",\"DVD\":\"21 Aug 2018\",\"BoxOffice\":\"N/A\",\"Production\":\"20th Century Fox\",\"Website\":\"https://www.foxmovies.com/movies/deadpool-2\",\"Response\":\"True\"}",
      "http://www.omdbapi.com/r-json_t-Sherlock Holmes: A Game Of Shadows": "{\"Title\":\"Sherlock Holmes: A Game of Shadows\",\"Year\":\"2011\",\"Rated\":\"PG-13\",\"Released\":\"16 Dec 2011\",\"Runtime\":\"129 min\",\"Genre\":\"Action, Adventure, Crime, Mystery, Thriller\",\"Director\":\"Guy Ritchie\",\"Writer\":\"Michele Mulroney, Kieran Mulroney, Arthur Conan Doyle (characters)\",\"Actors\":\"Robert Downey Jr., Jude Law, Noomi Rapace, Rachel McAdams\",\"Plot\":\"Sherlock Holmes and his sidekick Dr. Watson join forces to outwit and bring down their fiercest adversary, Professor Moriarty.\",\"Language\":\"English, French, Italian, German, Romany\",\"Country\":\"USA\",\"Awards\":\"3 wins & 11 nominations.\",\"Poster\":\"https://m.media-amazon.com/images/M/MV5BMTQwMzQ5Njk1MF5BMl5BanBnXkFtZTcwNjIxNzIxNw@@._V1_SX300.jpg\",\"Ratings\":[{\"Source\":\"Internet Movie Database\",\"Value\":\"7.5/10\"},{\"Source\":\"Rotten Tomatoes\",\"Value\":\"60%\"},{\"Source\":\"Metacritic\",\"Value\":\"48/100\"}],\"Metascore\":\"48\",\"imdbRating\":\"7.5\",\"imdbVotes\":\"384,886\",\"imdbID\":\"tt1515091\",\"Type\":\"movie\",\"DVD\":\"12 Jun 2012\",\"BoxOffice\":\"$186,830,669\",\"Production\":\"Warner Bros. Pictures\",\"Website\":\"http://www.sherlockholmes2.com\",\"Response\":\"True\"}",
      "http://www.omdbapi.com/r-json_t-The Five-Year Engagement": "{\"Title\":\"The Five-Year Engagement\",\"Year\":\"2012\",\"Rated\":\"R\",\"Released\":\"27 Apr 2012\",\"Runtime\":\"124 min\",\"Genre\":\"Comedy, Romance\",\"Director\":\"Nicholas Stoller\",\"Writer\":\"Jason Segel, Nicholas Stoller\",\"Actors\":\"Jason Segel, Emily Blunt, Chris Pratt, Alison Brie\",\"Plot\":\"One year after meeting, Tom proposes to his girlfriend, Violet, but unexpected events keep tripping them up as they look to walk down the aisle together.\",\"Language\":\"English, Spanish\",\"Country\":\"USA, Japan\",\"Awards\":\"N/A\",\"Poster\":\"https://m.media-amazon.com/images/M/MV5BMTk3MDE1NTA3M15BMl5BanBnXkFtZTcwNDcwNTMxNw@@._V1_SX300.jpg\",\"Ratings\":[{\"Source\":\"Internet Movie Database\",\"Value\":\"6.2/10\"},{\"Source\":\"Rotten Tomatoes\",\"Value\":\"63%\"},{\"Source\":\"Metacritic\",\"Value\":\"62/100\"}],\"Metascore\":\"62\",\"imdbRating\":\"6.2\",\"imdbVotes\":\"86,279\",\"imdbID\":\"tt1195478\",\"Type\":\"movie\",\"DVD\":\"04 Sep 2012\",\"BoxOffice\":\"$28,644,770\",\"Production\":\"Universal Pictures\",\"Website\":\"http://www.thefiveyearengagementmovie.com/\",\"Response\":\"True\"}",
      "http://www.omdbapi.com/r-json_t-American Assassin": "{\"Title\":\"American Assassin\",\"Year\":\"2017\",\"Rated\":\"R\",\"Released\":\"15 Sep 2017\",\"Runtime\":\"112 min\",\"Genre\":\"Action, Thriller\",\"Director\":\"Michael Cuesta\",\"Writer\":\"Stephen Schiff (screenplay by), Michael Finch (screenplay by), Edward Zwick (screenplay by), Marshall Herskovitz (screenplay by), Vince Flynn (based on the novel by)\",\"Actors\":\"Dylan O'Brien, Charlotte Vega, Christopher Bomford, Chris Webster\",\"Plot\":\"After the death of his girlfriend at the hands of terrorists, Mitch Rapp is drawn into the world of counterterrorism, mentored by tough-as-nails former U.S. Navy S.E.A.L. Stan Hurley.\",\"Language\":\"English, Spanish, Arabic, Polish, Turkish, Russian, Italian, Persian\",\"Country\":\"USA\",\"Awards\":\"1 nomination.\",\"Poster\":\"https://m.media-amazon.com/images/M/MV5BMjM0MjY4MTk5NV5BMl5BanBnXkFtZTgwNDE2NTQyMzI@._V1_SX300.jpg\",\"Ratings\":[{\"Source\":\"Internet Movie Database\",\"Value\":\"6.2/10\"},{\"Source\":\"Rotten Tomatoes\",\"Value\":\"34%\"},{\"Source\":\"Metacritic\",\"Value\":\"45/100\"}],\"Metascore\":\"45\",\"imdbRating\":\"6.2\",\"imdbVotes\":\"53,627\",\"imdbID\":\"tt1961175\",\"Type\":\"movie\",\"DVD\":\"05 Dec 2017\",\"BoxOffice\":\"$34,449,582\",\"Production\":\"Lionsgate and CBS Films\",\"Website\":\"http://www.cbsfilms.com/american-assassin/\",\"Response\":\"True\"}",
      "http://www.omdbapi.com/r-json_t-Baby Mama": "{\"Title\":\"Baby Mama\",\"Year\":\"2008\",\"Rated\":\"PG-13\",\"Released\":\"25 Apr 2008\",\"Runtime\":\"99 min\",\"Genre\":\"Comedy, Romance\",\"Director\":\"Michael McCullers\",\"Writer\":\"Michael McCullers\",\"Actors\":\"Amy Poehler, Tina Fey, Greg Kinnear, Dax Shepard\",\"Plot\":\"A successful, single businesswoman who dreams of having a baby discovers she is infertile and hires a working class woman to be her unlikely surrogate.\",\"Language\":\"English\",\"Country\":\"USA\",\"Awards\":\"1 win & 5 nominations.\",\"Poster\":\"https://m.media-amazon.com/images/M/MV5BMTYwNTc1Nzk1N15BMl5BanBnXkFtZTcwNjE4OTI2MQ@@._V1_SX300.jpg\",\"Ratings\":[{\"Source\":\"Internet Movie Database\",\"Value\":\"6.0/10\"},{\"Source\":\"Rotten Tomatoes\",\"Value\":\"64%\"},{\"Source\":\"Metacritic\",\"Value\":\"55/100\"}],\"Metascore\":\"55\",\"imdbRating\":\"6.0\",\"imdbVotes\":\"37,454\",\"imdbID\":\"tt0871426\",\"Type\":\"movie\",\"DVD\":\"09 Sep 2008\",\"BoxOffice\":\"$60,269,340\",\"Production\":\"Universal Pictures\",\"Website\":\"http://www.babymamamovie.net/\",\"Response\":\"True\"}",
      "https://tastedive.com/api/similarlimit-5_q-Bridesmaids_type-movies": "{\"Similar\": {\"Info\": [{\"Name\": \"Bridesmaids\", \"Type\": \"movie\"}], \"Results\": [{\"Name\": \"Baby Mama\", \"Type\": \"movie\"}, {\"Name\": \"The Five-Year Engagement\", \"Type\": \"movie\"}, {\"Name\": \"Bachelorette\", \"Type\": \"movie\"}, {\"Name\": \"The Heat\", \"Type\": \"movie\"}, {\"Name\": \"Date Night\", \"Type\": \"movie\"}]}}",
      "http://www.omdbapi.com/r-json_t-The Fate Of The Furious": "{\"Title\":\"The Fate of the Furious\",\"Year\":\"2017\",\"Rated\":\"PG-13\",\"Released\":\"14 Apr 2017\",\"Runtime\":\"136 min\",\"Genre\":\"Action, Crime, Thriller\",\"Director\":\"F. Gary Gray\",\"Writer\":\"Gary Scott Thompson (based on characters created by), Chris Morgan\",\"Actors\":\"Vin Diesel, Jason Statham, Dwayne Johnson, Michelle Rodriguez\",\"Plot\":\"When a mysterious woman seduces Dom into the world of terrorism and a betrayal of those closest to him, the crew face trials that will test them as never before.\",\"Language\":\"English, Russian\",\"Country\":\"China, USA, Japan\",\"Awards\":\"1 win & 9 nominations.\",\"Poster\":\"https://m.media-amazon.com/images/M/MV5BMjMxODI2NDM5Nl5BMl5BanBnXkFtZTgwNjgzOTk1MTI@._V1_SX300.jpg\",\"Ratings\":[{\"Source\":\"Internet Movie Database\",\"Value\":\"6.7/10\"},{\"Source\":\"Rotten Tomatoes\",\"Value\":\"67%\"},{\"Source\":\"Metacritic\",\"Value\":\"56/100\"}],\"Metascore\":\"56\",\"imdbRating\":\"6.7\",\"imdbVotes\":\"169,577\",\"imdbID\":\"tt4630562\",\"Type\":\"movie\",\"DVD\":\"11 Jul 2017\",\"BoxOffice\":\"$225,697,400\",\"Production\":\"Universal Pictures\",\"Website\":\"N/A\",\"Response\":\"True\"}",
      "http://www.omdbapi.com/r-json_t-Ant-Man And The Wasp": "{\"Title\":\"Ant-Man and the Wasp\",\"Year\":\"2018\",\"Rated\":\"PG-13\",\"Released\":\"06 Jul 2018\",\"Runtime\":\"118 min\",\"Genre\":\"Action, Adventure, Comedy, Sci-Fi\",\"Director\":\"Peyton Reed\",\"Writer\":\"Chris McKenna, Erik Sommers, Paul Rudd, Andrew Barrer, Gabriel Ferrari, Stan Lee (based on the Marvel Comics by), Larry Lieber (based on the Marvel Comics by), Jack Kirby (based on the Marvel Comics by)\",\"Actors\":\"Paul Rudd, Evangeline Lilly, Michael Pe\u00f1a, Walton Goggins\",\"Plot\":\"As Scott Lang balances being both a Super Hero and a father, Hope van Dyne and Dr. Hank Pym present an urgent new mission that finds the Ant-Man fighting alongside The Wasp to uncover secrets from their past.\",\"Language\":\"English\",\"Country\":\"USA\",\"Awards\":\"N/A\",\"Poster\":\"https://m.media-amazon.com/images/M/MV5BYjcyYTk0N2YtMzc4ZC00Y2E0LWFkNDgtNjE1MzZmMGE1YjY1XkEyXkFqcGdeQXVyMTMxODk2OTU@._V1_SX300.jpg\",\"Ratings\":[{\"Source\":\"Internet Movie Database\",\"Value\":\"7.2/10\"},{\"Source\":\"Rotten Tomatoes\",\"Value\":\"88%\"},{\"Source\":\"Metacritic\",\"Value\":\"70/100\"}],\"Metascore\":\"70\",\"imdbRating\":\"7.2\",\"imdbVotes\":\"158,453\",\"imdbID\":\"tt5095030\",\"Type\":\"movie\",\"DVD\":\"16 Oct 2018\",\"BoxOffice\":\"N/A\",\"Production\":\"Walt Disney Pictures\",\"Website\":\"https://marvel.com/antman\",\"Response\":\"True\"}",
      "https://tastedive.com/api/similarlimit-5_q-Captain Marvel_type-movies": "{\"Similar\": {\"Info\": [{\"Name\": \"Captain Marvel\", \"Type\": \"movie\"}], \"Results\": [{\"Name\": \"Inhumans\", \"Type\": \"movie\"}, {\"Name\": \"The Fate Of The Furious\", \"Type\": \"movie\"}, {\"Name\": \"Venom\", \"Type\": \"movie\"}, {\"Name\": \"American Assassin\", \"Type\": \"movie\"}, {\"Name\": \"Black Panther\", \"Type\": \"movie\"}]}}",
      "http://www.omdbapi.com/r-json_t-Inhumans": "{\"Title\":\"Inhumans\",\"Year\":\"2017\",\"Rated\":\"TV-PG\",\"Released\":\"29 Sep 2017\",\"Runtime\":\"43 min\",\"Genre\":\"Action, Adventure, Sci-Fi\",\"Director\":\"N/A\",\"Writer\":\"Scott Buck\",\"Actors\":\"Anson Mount, Serinda Swan, Ken Leung, Eme Ikwuakor\",\"Plot\":\"An isolated community of superhumans fight to protect themselves.\",\"Language\":\"English\",\"Country\":\"USA\",\"Awards\":\"N/A\",\"Poster\":\"https://m.media-amazon.com/images/M/MV5BZjI0ODcwODAtYzZjOS00ZmQ3LWE4ZGItOWMwYTI0Y2YyYWMyXkEyXkFqcGdeQXVyMzAzNTY3MDM@._V1_SX300.jpg\",\"Ratings\":[{\"Source\":\"Internet Movie Database\",\"Value\":\"5.1/10\"}],\"Metascore\":\"N/A\",\"imdbRating\":\"5.1\",\"imdbVotes\":\"17,590\",\"imdbID\":\"tt4154858\",\"Type\":\"series\",\"totalSeasons\":\"1\",\"Response\":\"True\"}",
      "https://tastedive.com/api/similarq-Black Panther": "{\"Similar\": {\"Info\": [{\"Name\": \"Black Panther\", \"Type\": \"movie\"}], \"Results\": [{\"Name\": \"Captain Marvel\", \"Type\": \"movie\"}, {\"Name\": \"Avengers: Infinity War\", \"Type\": \"movie\"}, {\"Name\": \"Ant-Man And The Wasp\", \"Type\": \"movie\"}, {\"Name\": \"The Fate Of The Furious\", \"Type\": \"movie\"}, {\"Name\": \"Deadpool 2\", \"Type\": \"movie\"}, {\"Name\": \"Jumanji: Welcome To The Jungle\", \"Type\": \"movie\"}, {\"Name\": \"Ready Player One\", \"Type\": \"movie\"}, {\"Name\": \"Ocean's 8\", \"Type\": \"movie\"}, {\"Name\": \"Solo: A Star Wars Story\", \"Type\": \"movie\"}, {\"Name\": \"Tag\", \"Type\": \"movie\"}, {\"Name\": \"Thor: Ragnarok\", \"Type\": \"movie\"}, {\"Name\": \"Venom\", \"Type\": \"movie\"}, {\"Name\": \"Pitch Perfect 3\", \"Type\": \"movie\"}, {\"Name\": \"Independence Day: Resurgence\", \"Type\": \"movie\"}, {\"Name\": \"The Divergent Series: Allegiant\", \"Type\": \"movie\"}, {\"Name\": \"Daddy's Home 2\", \"Type\": \"movie\"}, {\"Name\": \"The Greatest Showman\", \"Type\": \"movie\"}, {\"Name\": \"A Quiet Place\", \"Type\": \"movie\"}, {\"Name\": \"Game Night\", \"Type\": \"movie\"}, {\"Name\": \"Power Rangers\", \"Type\": \"movie\"}]}}",
      "http://www.omdbapi.com/r-json_t-Black Panther": "{\"Title\":\"Black Panther\",\"Year\":\"2018\",\"Rated\":\"PG-13\",\"Released\":\"16 Feb 2018\",\"Runtime\":\"134 min\",\"Genre\":\"Action, Adventure, Sci-Fi\",\"Director\":\"Ryan Coogler\",\"Writer\":\"Ryan Coogler, Joe Robert Cole, Stan Lee (based on the Marvel comics by), Jack Kirby (based on the Marvel Comics by)\",\"Actors\":\"Chadwick Boseman, Michael B. Jordan, Lupita Nyong'o, Danai Gurira\",\"Plot\":\"After the death of his father, Prince T'Challa ascends to the throne of the small African country of Wakanda. T'Challa is also Black Panther, superhero. Now not only must he govern his country but defeat a dangerous adversary.\",\"Language\":\"Swahili, Nama, English, Xhosa, Korean\",\"Country\":\"USA, South Africa, South Korea, Australia\",\"Awards\":\"14 nominations.\",\"Poster\":\"https://m.media-amazon.com/images/M/MV5BMTg1MTY2MjYzNV5BMl5BanBnXkFtZTgwMTc4NTMwNDI@._V1_SX300.jpg\",\"Ratings\":[{\"Source\":\"Internet Movie Database\",\"Value\":\"7.4/10\"},{\"Source\":\"Rotten Tomatoes\",\"Value\":\"97%\"},{\"Source\":\"Metacritic\",\"Value\":\"88/100\"}],\"Metascore\":\"88\",\"imdbRating\":\"7.4\",\"imdbVotes\":\"415,699\",\"imdbID\":\"tt1825683\",\"Type\":\"movie\",\"DVD\":\"15 May 2018\",\"BoxOffice\":\"$501,105,037\",\"Production\":\"Marvel Studios\",\"Website\":\"https://www.facebook.com/BlackPantherMovie/\",\"Response\":\"True\"}",
      "http://www.omdbapi.com/r-json_t-Bachelorette": "{\"Title\":\"Bachelorette\",\"Year\":\"2012\",\"Rated\":\"R\",\"Released\":\"06 Sep 2012\",\"Runtime\":\"87 min\",\"Genre\":\"Comedy, Romance\",\"Director\":\"Leslye Headland\",\"Writer\":\"Leslye Headland (screenplay), Leslye Headland (based on the play by)\",\"Actors\":\"Kirsten Dunst, Rebel Wilson, Lizzy Caplan, Paul Corning\",\"Plot\":\"Three friends are asked to be bridesmaids at a wedding of a woman they used to ridicule back in high school.\",\"Language\":\"English\",\"Country\":\"USA\",\"Awards\":\"1 win & 2 nominations.\",\"Poster\":\"https://m.media-amazon.com/images/M/MV5BZGQ1YjkyZGQtZDk0ZS00ZDgwLTlhYjgtNzQxNjkwMmQ0Mzk5XkEyXkFqcGdeQXVyMjA5MTIzMjQ@._V1_SX300.jpg\",\"Ratings\":[{\"Source\":\"Internet Movie Database\",\"Value\":\"5.3/10\"},{\"Source\":\"Rotten Tomatoes\",\"Value\":\"56%\"},{\"Source\":\"Metacritic\",\"Value\":\"53/100\"}],\"Metascore\":\"53\",\"imdbRating\":\"5.3\",\"imdbVotes\":\"42,043\",\"imdbID\":\"tt1920849\",\"Type\":\"movie\",\"DVD\":\"19 Mar 2013\",\"BoxOffice\":\"$400,000\",\"Production\":\"Radius-TWC\",\"Website\":\"http://www.facebook.com/bachelorettemovie\",\"Response\":\"True\"}",
      "http://www.omdbapi.com/r-json_t-Prince Of Persia: The Sands Of Time": "{\"Title\":\"Prince of Persia: The Sands of Time\",\"Year\":\"2010\",\"Rated\":\"PG-13\",\"Released\":\"28 May 2010\",\"Runtime\":\"116 min\",\"Genre\":\"Action, Adventure, Fantasy\",\"Director\":\"Mike Newell\",\"Writer\":\"Boaz Yakin (screenplay), Doug Miro (screenplay), Carlo Bernard (screenplay), Jordan Mechner (screen story), Jordan Mechner (video game series \\\"Prince of Persia\\\")\",\"Actors\":\"Jake Gyllenhaal, Gemma Arterton, Ben Kingsley, Alfred Molina\",\"Plot\":\"A young fugitive Prince and Princess must stop a villain who unknowingly threatens to destroy the world with a special dagger that enables the magic sand inside to reverse time.\",\"Language\":\"English\",\"Country\":\"USA\",\"Awards\":\"1 win & 10 nominations.\",\"Poster\":\"https://m.media-amazon.com/images/M/MV5BMTMwNDg0NzcyMV5BMl5BanBnXkFtZTcwNjg4MjQyMw@@._V1_SX300.jpg\",\"Ratings\":[{\"Source\":\"Internet Movie Database\",\"Value\":\"6.6/10\"},{\"Source\":\"Rotten Tomatoes\",\"Value\":\"36%\"},{\"Source\":\"Metacritic\",\"Value\":\"50/100\"}],\"Metascore\":\"50\",\"imdbRating\":\"6.6\",\"imdbVotes\":\"248,627\",\"imdbID\":\"tt0473075\",\"Type\":\"movie\",\"DVD\":\"14 Sep 2010\",\"BoxOffice\":\"$90,755,643\",\"Production\":\"Walt Disney Pictures\",\"Website\":\"http://Disney.com/PrinceOfPersia\",\"Response\":\"True\"}",
      "https://api.datamuse.com/wordsmax-3_rel_rhy-funny": "[{\"word\":\"money\",\"score\":4507,\"numSyllables\":2},{\"word\":\"honey\",\"score\":1250,\"numSyllables\":2},{\"word\":\"sunny\",\"score\":763,\"numSyllables\":2}]"
    }
