<?xml version="1.0"?>
<section xml:id="functions_returning-a-value-from-a-function">
  <title>Returning a value from a function</title>
  <image source="Functions/Figures/function_call.gif" width="50%" alt="gif of a box labeled function with three spaces on the top for input and a space on the bottom for output. Three arrows enter the top and are labeled as input or arguments. The function box shakes, and then one arrow leaves the bottom of the function box."/>
  <p>Not only can you pass a parameter value into a function, a function can also produce a value. You have already
            seen this in some previous functions that you have used. For example, <c>len</c> takes a list or string as a parameter
            value and returns a number, the length of that list or string. <c>range</c> takes an integer as a parameter value and
            returns a list containing all the numbers from 0 up to that parameter value.</p>
  <p>Functions that return values are sometimes called <term>fruitful functions</term>. In many other languages, a function that
            doesn't return a value is called a <term>procedure</term>, but we will stick here with the Python way of also calling it a
            function, or if we want to stress it, a <em>non-fruitful</em> function.</p>
  <video xml:id="goog_return_values" youtube="LGOZyrRCJ1o" width="auto"/>
  <image source="Functions/Figures/blackboxfun.png" width="50%"/>
  <p>How do we write our own fruitful function? Let's start by creating a very simple mathematical function that we will
            call <c>square</c>. The square function will take one number as a parameter and return the result of squaring that
            number. Here is the black-box diagram with the Python code following.</p>
  <image source="Functions/Figures/squarefun.png" width="50%"/>
  <program xml:id="ac11_4_1" interactive="activecode" language="python">
    <input>
def square(x):
    y = x * x
    return y

toSquare = 10
result = square(toSquare)
print("The result of {} squared is {}.".format(toSquare, result))
        </input>
  </program>
  <p>The <term>return</term> statement is followed by an expression which is evaluated. Its result is returned to the caller as the
            <q>fruit</q> of calling this function. Because the return statement can contain any Python expression we could have
            avoided creating the <term>temporary variable</term> <c>y</c> and simply used <c>return x*x</c>. Try modifying the square function
            above to see that this works just the same. On the other hand, using <term>temporary variables</term> like <c>y</c> in the program
            above makes debugging easier. These temporary variables are referred to as <term>local variables</term>.</p>
  <p>Notice something important here. The name of the variable we pass as an argument &#x2014; <c>toSquare</c> &#x2014; has nothing to
            do with the name of the formal parameter &#x2014; <c>x</c>.  It is as if  <c>x = toSquare</c> is executed when <c>square</c> is
            called. It doesn't matter what the value was named in the caller (the place where the function was invoked). Inside
            <c>square</c>, it's name is <c>x</c>.  You can see this very clearly in codelens, where the global variables and the local
            variables for the square function are in separate boxes.</p>
  <program xml:id="clens11_4_1" interactive="codelens" language="python">
    <input>
def square(x):
    y = x * x
    return y

toSquare = 10
squareResult = square(toSquare)
    </input>
  </program>
  <p>There is one more aspect of function return values that should be noted. All Python functions return the special value
            <c>None</c> unless there is an explicit return statement with a value other than <c>None</c>. Consider the following common
            mistake made by beginning Python programmers. As you step through this example, pay very close attention to the return
            value in the local variables listing. Then look at what is printed when the function is over.</p>
  <program xml:id="clens11_4_2" interactive="codelens" language="python">
    <input>
def square(x):
    y = x * x
    print(y)   # Bad! This is confusing! Should use return instead!

toSquare = 10
squareResult = square(toSquare)
print("The result of {} squared is {}.".format(toSquare, squareResult))
    </input>
  </program>
  <p>The problem with this function is that even though it prints the value of the squared input, that value will not be
            returned to the place where the call was done. Instead, the value <c>None</c> will be returned. Since line 6 uses the
            return value as the right hand side of an assignment statement, <c>squareResult</c> will have <c>None</c> as its value and
            the result printed in line 7 is incorrect. Typically, functions will return values that can be printed or processed in
            some other way by the caller.</p>
  <p>A return statement, once executed, immediately terminates execution of a function, even if it is not the last
            statement in the function. In the following code, when line 3 executes, the value 5 is returned and assigned to the
            variable x, then printed. Lines 4 and 5 never execute. Run the following code and try making some modifications of
            it to make sure you understand why <q>there</q> and 10 never print out.</p>
  <program xml:id="ac11_4_2" interactive="activecode" language="python">
    <input>
def weird():
    print("here")
    return 5
    print("there")
    return 10

x = weird()
print(x)
        </input>
  </program>
  <p>The fact that a return statement immediately ends execution of the code block inside a function is important to
            understand for writing complex programs, and it can also be very useful. The following example is a situation where
            you can use this to your advantage &#x2013; and understanding this will help you understand other people's code better, and
            be able to walk through code more confidently.</p>
  <p>Consider a situation where you want to write a function to find out, from a class attendance list, whether anyone's
            first name is longer than five letters, called <c>longer_than_five</c>. If there is anyone in class whose first name is
            longer than 5 letters, the function should return <c>True</c>. Otherwise, it should return <c>False</c>.</p>
  <p>In this case, you'll be using conditional statements in the code that exists in the <term>function body</term>, the code block
            indented underneath the function definition statement (just like the code that starts with the line <c>print("here")</c>
            in the example above &#x2013; that's the body of the function <c>weird</c>, above).</p>
  <p><term>Bonus challenge for studying:</term> After you look at the explanation below, stop looking at the code &#x2013; just the
            description of the function above it, and try to write the code yourself! Then test it on different lists and make
            sure that it works. But read the explanation first, so you can be sure you have a solid grasp on these function
            mechanics.</p>
  <p>First, an English plan for this new function to define called <c>longer_than_five</c>:</p>
  <p>
    <ul>
      <li>
        <p>You'll want to pass in a list of strings (representing people's first names) to the function.</p>
      </li>
      <li>
        <p>You'll want to iterate over all the items in the list, each of the strings.</p>
      </li>
      <li>
        <p>As soon as you get to one name that is longer than five letters, you know the function should return <c>True</c> &#x2013; yes, there is at least one name longer than five letters!</p>
      </li>
      <li>
        <p>And if you go through the whole list and there was no name longer than five letters, then the function should return <c>False</c>.</p>
      </li>
    </ul>
  </p>
  <p>Now, the code:</p>
  <program xml:id="ac11_4_3" interactive="activecode" language="python">
    <input>
def longer_than_five(list_of_names):
    for name in list_of_names: # iterate over the list to look at each name
        if len(name) &gt; 5: # as soon as you see a name longer than 5 letters,
            return True # then return True!
            # If Python executes that return statement, the function is over and the rest of the code will not run -- you already have your answer!
    return False # You will only get to this line if you
    # iterated over the whole list and did not get a name where
    # the if expression evaluated to True, so at this point, it's correct to return False!

# Here are a couple sample calls to the function with different lists of names. Try running this code in Codelens a few times and make sure you understand exactly what is happening.

list1 = ["Sam","Tera","Sal","Amita"]
list2 = ["Rey","Ayo","Lauren","Natalie"]

print(longer_than_five(list1))
print(longer_than_five(list2))
        </input>
  </program>
  <p>So far, we have just seen return values being assigned to variables. For example, we had the line
            <c>squareResult = square(toSquare)</c>. As with all assignment statements, the right hand side is executed first. It
            invokes the <c>square</c> function, passing in a parameter value 10 (the current value of <c>toSquare</c>). That returns a
            value 100, which completes the evaluation of the right-hand side of the assignment. 100 is then assigned to the
            variable <c>squareResult</c>. In this case, the function invocation was the entire expression that was evaluated.</p>
  <p>Function invocations, however, can also be used as part of more complicated expressions. For example,
            <c>squareResult = 2 * square(toSquare)</c>. In this case, the value 100 is returned and is then multiplied by 2 to
            produce the value 200. When python evaluates an expression like <c>x * 3</c>, it substitutes the current value of x into
            the expression and then does the multiplication. When python evaluates an expression like <c>2 * square(toSquare)</c>, it
            substitutes the return value 100 for entire function invocation and then does the multiplication.</p>
  <p>To reiterate, when executing a line of code <c>squareResult = 2 * square(toSquare)</c>, the python
            interpreter does these steps:</p>
  <p>
    <ol label="1">
      <li>
        <p>It's an assignment statement, so evaluate the right-hand side expression <c>2 * square(toSquare)</c>.</p>
      </li>
      <li>
        <p>Look up the values of the variables square and toSquare: square is a function object and toSquare is 10</p>
      </li>
      <li>
        <p>Pass 10 as a parameter value to the function, get back the return value 100</p>
      </li>
      <li>
        <p>Substitute 100 for square(toSquare), so that the expression now reads <c>2 * 100</c></p>
      </li>
      <li>
        <p>Assign 200 to variable <c>squareResult</c></p>
      </li>
    </ol>
  </p>
  <p>
    <term>Check your understanding</term>
  </p>
  <exercise label="question11_4_1">
    <statement>
      <p>What is wrong with the following function definition:</p>
      <program language="python">
        <input>
def addEm(x, y, z):
    return x+y+z
    print('the answer is', x+y+z)
</input>
      </program>
    </statement>
    <choices>
      <choice>
        <statement>
          <p>You should never use a print statement in a function definition.</p>
        </statement>
        <feedback>
          <p>Although you should not mistake print for return, you may include print statements inside your functions.</p>
        </feedback>
      </choice>
      <choice correct="yes">
        <statement>
          <p>You should not have any statements in a function after the return statement.  Once the function gets to the return statement it will immediately stop executing the function.</p>
        </statement>
        <feedback>
          <p>This is a very common mistake so be sure to watch out for it when you write your code!</p>
        </feedback>
      </choice>
      <choice>
        <statement>
          <p>You must calculate the value of x+y+z before you return it.</p>
        </statement>
        <feedback>
          <p>Python will automatically calculate the value x+y+z and then return it in the statement as it is written</p>
        </feedback>
      </choice>
      <choice>
        <statement>
          <p>A function cannot return a number.</p>
        </statement>
        <feedback>
          <p>Functions can return any legal data, including (but not limited to) numbers, strings, lists, dictionaries, etc.</p>
        </feedback>
      </choice>
    </choices>
  </exercise>
  <exercise label="question11_4_2">
    <statement>
      <p>What will the following function return?</p>
      <program language="python">
        <input>
def addEm(x, y, z):
    print(x+y+z)
</input>
      </program>
    </statement>
    <choices>
      <choice correct="yes">
        <statement>
          <p>The value None</p>
        </statement>
        <feedback>
          <p>We have accidentally used print where we mean return.  Therefore, the function will return the value None by default.  This is a VERY COMMON mistake so watch out!  This mistake is also particularly difficult to find because when you run the function the output looks the same.  It is not until you try to assign its value to a variable that you can notice a difference.</p>
        </feedback>
      </choice>
      <choice>
        <statement>
          <p>The value of x+y+z</p>
        </statement>
        <feedback>
          <p>Careful!  This is a very common mistake.  Here we have printed the value x+y+z but we have not returned it.  To return a value we MUST use the return keyword.</p>
        </feedback>
      </choice>
      <choice>
        <statement>
          <p>The string 'x+y+z'</p>
        </statement>
        <feedback>
          <p>x+y+z calculates a number (assuming x+y+z are numbers) which represents the sum of the values x, y and z.</p>
        </feedback>
      </choice>
    </choices>
  </exercise>
  <exercise label="question11_4_3">
    <statement>
      <p>What will the following code output?</p>
      <program language="python">
        <input>
def square(x):
    y = x * x
    return y

print(square(5) + square(5))
</input>
      </program>
    </statement>
    <choices>
      <choice>
        <statement>
          <p>25</p>
        </statement>
        <feedback>
          <p>It squares 5 twice, and adds them together.</p>
        </feedback>
      </choice>
      <choice correct="yes">
        <statement>
          <p>50</p>
        </statement>
        <feedback>
          <p>The two return values are added together.</p>
        </feedback>
      </choice>
      <choice>
        <statement>
          <p>25 + 25</p>
        </statement>
        <feedback>
          <p>The two results are substituted into the expression and then it is evaluated. The returned values are integers in this case, not strings.</p>
        </feedback>
      </choice>
    </choices>
  </exercise>
  <exercise label="question11_4_4">
    <statement>
      <p>What will the following code output?</p>
      <program language="python">
        <input>
def square(x):
    y = x * x
    return y

print(square(square(2)))
</input>
      </program>
    </statement>
    <choices>
      <choice>
        <statement>
          <p>8</p>
        </statement>
        <feedback>
          <p>It squares 2, yielding the value 4. But that doesn't mean the next value multiplies 2 and 4.</p>
        </feedback>
      </choice>
      <choice correct="yes">
        <statement>
          <p>16</p>
        </statement>
        <feedback>
          <p>It squares 2, yielding the value 4. 4 is then passed as a value to square again, yeilding 16.</p>
        </feedback>
      </choice>
      <choice>
        <statement>
          <p>Error: can't put a function invocation inside parentheses</p>
        </statement>
        <feedback>
          <p>This is a more complicated expression, but still valid. The expression square(2) is evaluated, and the return value 4 substitutes for square(2) in the expression.</p>
        </feedback>
      </choice>
    </choices>
  </exercise>
  <exercise label="question11_4_5">
    <statement>
      <p>What will the following code output?</p>
      <program language="python">
        <input>
def cyu2(s1, s2):
    x = len(s1)
    y = len(s2)
    return x-y

z = cyu2("Yes", "no")
if z &gt; 0:
    print("First one was longer")
else:
    print("Second one was at least as long")
</input>
      </program>
    </statement>
    <choices>
      <choice>
        <statement>
          <p>1</p>
        </statement>
        <feedback>
          <p>cyu2 returns the value 1, but that's not what prints.</p>
        </feedback>
      </choice>
      <choice>
        <statement>
          <p>Yes</p>
        </statement>
        <feedback>
          <p>"Yes" is longer, but that's not what prints.</p>
        </feedback>
      </choice>
      <choice correct="yes">
        <statement>
          <p>First one was longer</p>
        </statement>
        <feedback>
          <p>cyu2 returns the value 1, which is assigned to z.</p>
        </feedback>
      </choice>
      <choice>
        <statement>
          <p>Second one was at least as long</p>
        </statement>
        <feedback>
          <p>cyu2 returns the value 1, which is assigned to z.</p>
        </feedback>
      </choice>
      <choice>
        <statement>
          <p>Error</p>
        </statement>
        <feedback>
          <p>what do you think will cause an error.</p>
        </feedback>
      </choice>
    </choices>
  </exercise>
  <exercise label="question11_4_6">
    <statement>
      <p>Which will print out first, square, g, or a number?</p>
      <program language="python">
        <input>
def square(x):
    print("square")
    return x*x

def g(y):
    print("g")
    return y + 3

print(square(g(2)))
</input>
      </program>
    </statement>
    <choices>
      <choice>
        <statement>
          <p>square</p>
        </statement>
        <feedback>
          <p>Before executing square, it has to figure out what value to pass in, so g is executed first</p>
        </feedback>
      </choice>
      <choice correct="yes">
        <statement>
          <p>g</p>
        </statement>
        <feedback>
          <p>g has to be executed and return a value in order to know what paramater value to provide to x.</p>
        </feedback>
      </choice>
      <choice>
        <statement>
          <p>a number</p>
        </statement>
        <feedback>
          <p>square and g both have to execute before the number is printed.</p>
        </feedback>
      </choice>
    </choices>
  </exercise>
  <exercise label="question11_4_7">
    <statement>
      <p>How many lines will the following code print?</p>
      <program language="python">
        <input>
def show_me_numbers(list_of_ints):
    print(10)
    print("Next we'll accumulate the sum")
    accum = 0
    for num in list_of_ints:
        accum = accum + num
    return accum
    print("All done with accumulation!")

show_me_numbers([4,2,3])
</input>
      </program>
    </statement>
    <choices>
      <choice>
        <statement>
          <p>3</p>
        </statement>
        <feedback>
          <p>The function gets to a return statement after 2 lines are printed, so the third print statement will not run.</p>
        </feedback>
      </choice>
      <choice correct="yes">
        <statement>
          <p>2</p>
        </statement>
        <feedback>
          <p>Yes! Two printed lines, and then the function body execution reaches a return statement.</p>
        </feedback>
      </choice>
      <choice>
        <statement>
          <p>None</p>
        </statement>
        <feedback>
          <p>The function returns an integer value! However, this code does not print out the result of the function invocation, so you can't see it (print is for people). The only lines you see printed are the ones that occur in the print statements before the return statement.</p>
        </feedback>
      </choice>
    </choices>
  </exercise>
  <exercise label="ac11_4_4">
    <statement>
      <p><term>8.</term> Write a function named <c>same</c> that takes a string as input, and simply returns that string.</p>
    </statement>
    <program xml:id="ac11_4_4_editor" interactive="activecode" language="python">
      <input>

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(same('hello'), 'hello', "Testing the same function on input 'hello'.")

myTests().main()
        </input>
    </program>
  </exercise>
  <exercise label="ac11_4_5">
    <statement>
      <p><term>9.</term> Write a function called <c>same_thing</c> that returns the parameter, unchanged.</p>
    </statement>
    <program xml:id="ac11_4_5_editor" interactive="activecode" language="python">
      <input>

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(same_thing(5), 5,"Testing the function same_thing with input 5")
      self.assertEqual(same_thing("Welcome"), "Welcome", "Testing the function same_thing with input 'Welcome'")

myTests().main()
        </input>
    </program>
  </exercise>
  <exercise label="ac11_4_6">
    <statement>
      <p><term>10.</term> Write a function called <c>subtract_three</c> that takes an integer or any number as input, and returns that number minus three.</p>
    </statement>
    <program xml:id="ac11_4_6_editor" interactive="activecode" language="python">
      <input>

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(subtract_three(9), 6, "Testing the subtract_three function on input 9.")
      self.assertEqual(subtract_three(-5), -8, "Testing the subtract_three function on input -5.")

myTests().main()
        </input>
    </program>
  </exercise>
  <exercise label="ac11_4_7">
    <statement>
      <p><term>11.</term> Write a function called <c>change</c> that takes one number as its input and returns that number, plus 7.</p>
    </statement>
    <program xml:id="ac11_4_7_editor" interactive="activecode" language="python">
      <input>

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(change(5), 12,"Testing the function change with input 5")
      self.assertEqual(change(-10), -3, "Testing the function change with input -10")

myTests().main()
        </input>
    </program>
  </exercise>
  <exercise label="ac11_4_8">
    <statement>
      <p><term>12.</term> Write a function named <c>intro</c> that takes a string as input. This string ist intended to be a person's name and the output is a standardized greeting. For example, given the string <q>Becky</q> as input, the function should return: <q>Hello, my name is Becky and I love SI 106.</q></p>
    </statement>
    <program xml:id="ac11_4_8_editor" interactive="activecode" language="python">
      <input>

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(intro("Mike"), "Hello, my name is Mike and I love SI 106.", "Testing the intro function on input 'Mike'.")

myTests().main()
        </input>
    </program>
  </exercise>
  <exercise label="ac11_4_9">
    <statement>
      <p><term>13.</term> Write a function called <c>s_change</c> that takes one string as input and returns that string, concatenated with the string <q> for fun.</q>.</p>
    </statement>
    <program xml:id="ac11_4_9_editor" interactive="activecode" language="python">
      <input>

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(s_change("Coding"), "Coding for fun." ,"Testing the function s_change with input coding")
      self.assertEqual(s_change("We go to the beach"), "We go to the beach for fun." , "Testing the function s_change with input We go to the beach")

myTests().main()
        </input>
    </program>
  </exercise>
  <exercise label="ac11_4_10">
    <statement>
      <p><term>14.</term> Write a function called <c>decision</c> that takes a string as input, and then checks the number of characters. If it has over 17 characters, return <q>This is a long string</q>, if it is shorter or has 17 characters, return <q>This is a short string</q>.</p>
    </statement>
    <program xml:id="ac11_4_10_editor" interactive="activecode" language="python">
      <input>

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(decision("Well hello dolly"), "This is a short string", "Testing the function decision with input 'Well hello dolly'")
      self.assertEqual(decision("In olden days a glimps of stocking was looked on a something shocking but heaven knows, anything goes"), "This is a long string", "Testing the function decision with input 'In olden days a glimps of stocking was looked on a something shocking but heaven knows, anything goes'")
      self.assertEqual(decision("how do you do sir"), "This is a short string", "Testing the function decision with input 'how do you do sir'")

myTests().main()
        </input>
    </program>
  </exercise>
</section>
