<!DOCTYPE html>

<html xmlns="http://www.w3.org/1999/xhtml" lang="en">
  <head>
    <meta http-equiv="X-UA-Compatible" content="IE=Edge" />
    <meta http-equiv="Content-Type" content="text/html; charset=utf-8" />
    <title>7.6. The Accumulator Pattern &#8212; Runestone Interactive Overview</title>
    <link rel="stylesheet" href="../_static/basic.css" type="text/css" />
    <link rel="stylesheet" href="../_static/pygments.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-3.0.0/css/bootstrap.min.css" type="text/css" />
    <link rel="stylesheet" type="text/css" href="../_static/matrixeq.css" />
    <link rel="stylesheet" type="text/css" href="../_static/clickable.css" />
    <link rel="stylesheet" type="text/css" href="../_static/tabbedstuff.css" />
    <link rel="stylesheet" type="text/css" href="../_static/fitb.css" />
    <link rel="stylesheet" type="text/css" href="../_static/video.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webgldemo.css" />
    <link rel="stylesheet" type="text/css" href="../_static/webglinteractive.css" />
    <link rel="stylesheet" type="text/css" href="../_static/codemirror.css" />
    <link rel="stylesheet" type="text/css" href="../_static/reveal.css" />
    <link rel="stylesheet" type="text/css" href="../_static/parsons.css" />
    <link rel="stylesheet" type="text/css" href="../_static/js_lib/prettify.css" />
    <link rel="stylesheet" type="text/css" href="../_static/pytutor.css" />
    <link rel="stylesheet" type="text/css" href="../_static/modal-basic.css" />
    <link rel="stylesheet" type="text/css" href="../_static/activecode.css" />
    <link rel="stylesheet" type="text/css" href="../_static/datafile.css" />
    <link rel="stylesheet" type="text/css" href="../_static/showEval.css" />
    <link rel="stylesheet" type="text/css" href="../_static/accessibility.css" />
    <link rel="stylesheet" type="text/css" href="../_static/dragndrop.css" />
    <link rel="stylesheet" type="text/css" href="../_static/poll.css" />
    <link rel="stylesheet" href="../_static/jquery-ui-1.10.3.custom.min.css" type="text/css" />
    <link rel="stylesheet" href="../_static/bootstrap-sphinx.css" type="text/css" />
    <link rel="stylesheet" href="../_static/user-highlights.css" type="text/css" />
    <link rel="stylesheet" href="../_static/runestone-custom-sphinx-bootstrap.css" type="text/css" />
    <script type="text/javascript" id="documentation_options" data-url_root="../" src="../_static/documentation_options.js"></script>
    <script type="text/javascript" src="../_static/runestonebase.js"></script>
    <script type="text/javascript" src="../_static/jquery.js"></script>
    <script type="text/javascript" src="../_static/underscore.js"></script>
    <script type="text/javascript" src="../_static/doctools.js"></script>
    <script type="text/javascript" src="../_static/matrixeq.js"></script>
    <script type="text/javascript" src="../_static/clickable.js"></script>
    <script type="text/javascript" src="../_static/timedclickable.js"></script>
    <script type="text/javascript" src="../_static/tabbedstuff.js"></script>
    <script type="text/javascript" src="../_static/fitb.js"></script>
    <script type="text/javascript" src="../_static/timedfitb.js"></script>
    <script type="text/javascript" src="../_static/runestonevideo.js"></script>
    <script type="text/javascript" src="../_static/codemirror.js"></script>
    <script type="text/javascript" src="../_static/xml.js"></script>
    <script type="text/javascript" src="../_static/css.js"></script>
    <script type="text/javascript" src="../_static/htmlmixed.js"></script>
    <script type="text/javascript" src="../_static/javascript.js"></script>
    <script type="text/javascript" src="../_static/webglinteractive.js"></script>
    <script type="text/javascript" src="../_static/FileSaver.min.js"></script>
    <script type="text/javascript" src="../_static/Blob.js"></script>
    <script type="text/javascript" src="../_static/reveal.js"></script>
    <script type="text/javascript" src="../_static/js_lib/prettify.js"></script>
    <script type="text/javascript" src="../_static/js_lib/hammer.min.js"></script>
    <script type="text/javascript" src="../_static/parsons.js"></script>
    <script type="text/javascript" src="../_static/timedparsons.js"></script>
    <script type="text/javascript" src="../_static/animationbase.js"></script>
    <script type="text/javascript" src="../_static/mchoice.js"></script>
    <script type="text/javascript" src="../_static/timedmc.js"></script>
    <script type="text/javascript" src="../_static/timed.js"></script>
    <script type="text/javascript" src="../_static/shortanswer.js"></script>
    <script type="text/javascript" src="../_static/timed_shortanswer.js"></script>
    <script type="text/javascript" src="../_static/d3.v2.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.ba-bbq.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.jsPlumb-1.3.10-all-min.js"></script>
    <script type="text/javascript" src="../_static/pytutor.js"></script>
    <script type="text/javascript" src="../_static/codelens.js"></script>
    <script type="text/javascript" src="../_static/jquery.highlight.js"></script>
    <script type="text/javascript" src="../_static/bookfuncs.js"></script>
    <script type="text/javascript" src="../_static/python.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/CLDRPluralRuleParser.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.messagestore.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.fallbacks.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.language.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.parser.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.emitter.js"></script>
    <script type="text/javascript" src="../_static/jquery_i18n/jquery.i18n.emitter.bidi.js"></script>
    <script type="text/javascript" src="../_static/activecode-i18n.en.js"></script>
    <script type="text/javascript" src="../_static/skulpt.min.js"></script>
    <script type="text/javascript" src="../_static/skulpt-stdlib.js"></script>
    <script type="text/javascript" src="https://cdn.jsdelivr.net/npm/vega@4.0.0-rc.2/build/vega.js"></script>
    <script type="text/javascript" src="https://cdn.jsdelivr.net/npm/vega-lite@2.5.0/build/vega-lite.js"></script>
    <script type="text/javascript" src="https://cdn.jsdelivr.net/npm/vega-embed@3.14.0/build/vega-embed.js"></script>
    <script type="text/javascript" src="../_static/activecode.js"></script>
    <script type="text/javascript" src="../_static/clike.js"></script>
    <script type="text/javascript" src="../_static/timed_activecode.js"></script>
    <script type="text/javascript" src="../_static/datafile.js"></script>
    <script type="text/javascript" src="../_static/showEval.js"></script>
    <script type="text/javascript" src="../_static/dragndrop.js"></script>
    <script type="text/javascript" src="../_static/timeddnd.js"></script>
    <script type="text/javascript" src="../_static/poll.js"></script>
    <script async="async" type="text/javascript" src="https://cdnjs.cloudflare.com/ajax/libs/mathjax/2.7.1/MathJax.js?config=TeX-AMS-MML_HTMLorMML"></script>
    <script type="text/javascript" src="../_static/jquery-ui-1.10.3.custom.min.js"></script>
    <script type="text/javascript" src="../_static/jquery-fix.js"></script>
    <script type="text/javascript" src="../_static/bootstrap-3.0.0/js/bootstrap.min.js"></script>
    <script type="text/javascript" src="../_static/bootstrap-sphinx.js"></script>
    <script type="text/javascript" src="../_static/waypoints.min.js"></script>
    <script type="text/javascript" src="../_static/user-highlights.js"></script>
    <script type="text/javascript" src="../_static/jquery.idle-timer.js"></script>
    <script type="text/javascript" src="../_static/processing-1.4.1.min.js"></script>
    <script type="text/javascript" src="../_static/jquery.hotkey.js"></script>
    <script type="text/javascript" src="../_static/jquery-migrate-1.2.1.min.js"></script>
    <script type="text/javascript" src="../_static/presenter_mode.js"></script>
    <link rel="index" title="Index" href="../genindex.html" />
    <link rel="search" title="Search" href="../search.html" />
    <link rel="next" title="7.7. Traversal and the for Loop: By Index" href="TraversalandtheforLoopByIndex.html" />
    <link rel="prev" title="7.5. Lists and for loops" href="Listsandforloops.html" />
<meta charset='utf-8'>
<meta http-equiv='X-UA-Compatible' content='IE=edge,chrome=1'>
<meta content='width=device-width, initial-scale=1.0, maximum-scale=1.0, user-scalable=0' name='viewport' />
<link rel="shortcut icon" href="/runestone/static/favicon.ico" type="image/ico" />

<script type="text/javascript">
  eBookConfig = {};
  eBookConfig.host = 'http://127.0.0.1:8000' ? 'http://127.0.0.1:8000' : 'http://127.0.0.1:8000';
  eBookConfig.app = eBookConfig.host+'/runestone';
  eBookConfig.ajaxURL = eBookConfig.app+'/ajax/';
  eBookConfig.course = 'fopp';
  eBookConfig.logLevel = 10;
  eBookConfig.loginRequired = false;
  eBookConfig.build_info = "unknown";
  eBookConfig.isLoggedIn = false;
  eBookConfig.useRunestoneServices = true;
  eBookConfig.python3 = true;
  eBookConfig.basecourse = 'fopp';
  eBookConfig.acDefaultLanguage = 'python' ? 'python' : 'python'
  eBookConfig.runestone_version = '3.2.11';
  eBookConfig.jobehost = 'http://jobe2.cosc.canterbury.ac.nz';
  eBookConfig.proxyuri_runs = '/jobe/index.php/restapi/runs/';
  eBookConfig.proxyuri_files = '/jobe/index.php/restapi/files/';
  eBookConfig.downloadsEnabled = false;
  eBookConfig.enable_chatcodes = false ? false : false;
</script>

<div id="fb-root"></div>
<script>
  (function(d, s, id) {
    var js, fjs = d.getElementsByTagName(s)[0];
    if (d.getElementById(id)) return;
    js = d.createElement(s); js.id = id;
    js.src = "//connect.facebook.net/en_US/all.js#xfbml=1&status=0";
    fjs.parentNode.insertBefore(js, fjs);
  }(document, 'script', 'facebook-jssdk'));
</script>


  </head><body>


<!-- Begin navbar -->
<div id="navbar" class="navbar navbar-default navbar-fixed-top" role="navigation">

  <div class="container">

    <div class="navbar-header">
      <!-- .btn-navbar is used as the toggle for collapsed navbar content -->
      <button type='button' class='navbar-toggle' data-toggle="collapse" data-target=".navbar-ex1-collapse">
        <span class="icon-bar" aria-label="navbar toggle" class="visuallyhidden">toggle navbar</span>
        <span class="icon-bar"></span>
        <span class="icon-bar"></span>
      </button>

      <div>
        <a class="brand-logo" style='background: transparent url("../_static/img/logo_small.png") no-repeat 0px 0px;' href="/runestone/default/user/login">&nbsp;<span aria-label="logo" class="visuallyhidden">Runestone Logo</span> </a>
        <a class="navbar-brand" href="../index.html">Foundations of Python Programming</a>
      </div>
    </div>

     

    <div class="navbar-collapse collapse navbar-ex1-collapse">
      <ul class="nav navbar-nav navbar-right" >

        <li class="divider-vertical"></li>

        <!-- social media dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-globe" style="opacity: 0.9"><span aria-label="Social"></span></i>
          </a>
          <ul class="dropdown-menu social-menu">
              <li>
                <div>
                  <b>Runestone in social media:</b>
                </div>
                <a href="https://twitter.com/iRunestone" class="twitter-follow-button" data-show-count="true">Follow @iRunestone</a><script>!function(d,s,id){var js,fjs=d.getElementsByTagName(s)[0],p=/^http:/.test(d.location)?'http':'https';if(!d.getElementById(id)){js=d.createElement(s);js.id=id;js.src=p+'://platform.twitter.com/widgets.js';fjs.parentNode.insertBefore(js,fjs);}}(document, 'script', 'twitter-wjs');</script>
                <div class="fb-like" data-href="https://www.facebook.com/RunestoneInteractive" data-send="false" data-layout="button_count" data-width="300" data-show-faces="false" data-font="arial"></div>
              </li>

              <li class="divider"></li>
              <li>
                <div>
                  <b>Help support us:</b>
                </div>
                <div>
                    <form action="https://www.paypal.com/cgi-bin/webscr" method="post" target="_top">
                  <input type="hidden" name="cmd" value="_s-xclick">
                  <input type="hidden" name="encrypted" value="-----BEGIN PKCS7-----MIIHNwYJKoZIhvcNAQcEoIIHKDCCByQCAQExggEwMIIBLAIBADCBlDCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20CAQAwDQYJKoZIhvcNAQEBBQAEgYAcrkqh1hn3lYqIpfXxNqe1T82EhXzCJGy1yMAmklpyZshyMkfDGe1Bhx+iwyGeoYRTTyphFmP+9M3NyO0+Q5XdHxgZPx/zYjjBxlZHgEV6jhE8bN2fHkkPf0VHfz0a0QQylQOPlKiOTZV7B37Jpk6yM47oVZ1tG/KNm0NkfmB76DELMAkGBSsOAwIaBQAwgbQGCSqGSIb3DQEHATAUBggqhkiG9w0DBwQIi0GmFfOlcjuAgZBbYOo9UO+CpMQa+PYqwsUmUnJvXIImeMeNI3KVTUx5Cfk9gNMo3WzPeiB5IqZo9nRAQ0mf1qL2ecLeB5tidM+lgBUhOxfj3/FecpnVFa0263gp4g+PLw8jzhvVRdUon1K3SeO1Rzh23fIRKwnrD6btt73uwtj0sl3tGd8qz+6GIcwPDdRk9VcUffiBJT/ZagKgggOHMIIDgzCCAuygAwIBAgIBADANBgkqhkiG9w0BAQUFADCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wHhcNMDQwMjEzMTAxMzE1WhcNMzUwMjEzMTAxMzE1WjCBjjELMAkGA1UEBhMCVVMxCzAJBgNVBAgTAkNBMRYwFAYDVQQHEw1Nb3VudGFpbiBWaWV3MRQwEgYDVQQKEwtQYXlQYWwgSW5jLjETMBEGA1UECxQKbGl2ZV9jZXJ0czERMA8GA1UEAxQIbGl2ZV9hcGkxHDAaBgkqhkiG9w0BCQEWDXJlQHBheXBhbC5jb20wgZ8wDQYJKoZIhvcNAQEBBQADgY0AMIGJAoGBAMFHTt38RMxLXJyO2SmS+Ndl72T7oKJ4u4uw+6awntALWh03PewmIJuzbALScsTS4sZoS1fKciBGoh11gIfHzylvkdNe/hJl66/RGqrj5rFb08sAABNTzDTiqqNpJeBsYs/c2aiGozptX2RlnBktH+SUNpAajW724Nv2Wvhif6sFAgMBAAGjge4wgeswHQYDVR0OBBYEFJaffLvGbxe9WT9S1wob7BDWZJRrMIG7BgNVHSMEgbMwgbCAFJaffLvGbxe9WT9S1wob7BDWZJRroYGUpIGRMIGOMQswCQYDVQQGEwJVUzELMAkGA1UECBMCQ0ExFjAUBgNVBAcTDU1vdW50YWluIFZpZXcxFDASBgNVBAoTC1BheVBhbCBJbmMuMRMwEQYDVQQLFApsaXZlX2NlcnRzMREwDwYDVQQDFAhsaXZlX2FwaTEcMBoGCSqGSIb3DQEJARYNcmVAcGF5cGFsLmNvbYIBADAMBgNVHRMEBTADAQH/MA0GCSqGSIb3DQEBBQUAA4GBAIFfOlaagFrl71+jq6OKidbWFSE+Q4FqROvdgIONth+8kSK//Y/4ihuE4Ymvzn5ceE3S/iBSQQMjyvb+s2TWbQYDwcp129OPIbD9epdr4tJOUNiSojw7BHwYRiPh58S1xGlFgHFXwrEBb3dgNbMUa+u4qectsMAXpVHnD9wIyfmHMYIBmjCCAZYCAQEwgZQwgY4xCzAJBgNVBAYTAlVTMQswCQYDVQQIEwJDQTEWMBQGA1UEBxMNTW91bnRhaW4gVmlldzEUMBIGA1UEChMLUGF5UGFsIEluYy4xEzARBgNVBAsUCmxpdmVfY2VydHMxETAPBgNVBAMUCGxpdmVfYXBpMRwwGgYJKoZIhvcNAQkBFg1yZUBwYXlwYWwuY29tAgEAMAkGBSsOAwIaBQCgXTAYBgkqhkiG9w0BCQMxCwYJKoZIhvcNAQcBMBwGCSqGSIb3DQEJBTEPFw0xMzExMDMxMzQxMzFaMCMGCSqGSIb3DQEJBDEWBBRDJF8w+zsMr7FSk+pwinB5f5D4rzANBgkqhkiG9w0BAQEFAASBgHw1LMHpkpaqHIvDGdFE0eG+2mZlmMnUeDCBhQlbc7QMzFQYKTV94NfaebBO4PmNdADe1rq4WidSRZZbE7CzkX9IGENYnBTWY0hb2l0lGdGrJdGeWyV3ekg9WVaFMMumrekds96h3Cx7dGz2kWDzIai2iEXE/qoE+xpkyXAYZNV3-----END PKCS7-----
                  ">
                  <input type="image" src="https://www.paypalobjects.com/en_US/i/btn/btn_donate_SM.gif" border="0" name="submit" alt="PayPal - The safer, easier way to pay online!">
                  <img alt="" border="0" src="https://www.paypalobjects.com/en_US/i/scr/pixel.gif" width="1" height="1">
                  </form>

                </div>
              </li>
          </ul>
        </li>
        <!-- end social media dropdown -->

        <li class="divider-vertical"></li>

        <!-- search dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-search" style='opacity:0.9;'><span aria-label="Search" class="visuallyhidden">Search</span></i>
          </a>
          <ul class='dropdown-menu'>
            
                <li><a href='../index.html'>Table of Contents</a></li>
            
            <li><a href='/runestone/static/fopp/genindex.html'>Book Index</a></li>
            <li class="divider"></li>
            <li style="width: 240px;">
              <form class="navbar-search" style="margin:10px;" action="../search.html" method="get">
                <div class="input-group">
                  <input type="text" class="form-control" name="q" placeholder="Search this book" />
                  <span class="input-group-btn">
                    <button class="btn btn-primary" style="margin:0;" type="submit">
                      <i class="glyphicon glyphicon-search"></i>
                    </button>
                  </span>
                </div><!-- /input-group -->
                <input type="hidden" name="check_keywords" value="yes" />
                <input type="hidden" name="area" value="default" />
              </form>
            </li>
          </ul>
        </li>
        <!-- end search dropdown -->

        <li class="divider-vertical"></li>

        
        <!-- user account dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-user" style="opacity:0.9;"><span aria-label="User" class="visuallyhidden">User</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><span class='loggedinuser'></span></li>
            <li class="divider"></li>
            <li><a href='http://127.0.0.1:8000/runestone/assignments/practice'>Practice</a></li>
            <li class="divider"></li>
            <li><a href='/runestone/default/courses'>Change Course</a></li>
            <li class="divider"></li>
            <li><a href='http://127.0.0.1:8000/runestone/mygroup/initiateGroup' id="joinGroupLink">Join a Study Group</a></li>
            <li><a href='http://127.0.0.1:8000/runestone/mygroup/schedule' id="groupScheduleLink">Group Schedule</a></li>
            <li><a href='http://127.0.0.1:8000/runestone/mygroup/newschedule' id="newChapterLink">Schedule New Chapter</a></li>
            <li><a href='http://127.0.0.1:8000/runestone/mygroup/manageGroup' id="manageGroupLink">Manage Group</a></li>
			<li class="divider"></li>
            <li><a href='http://127.0.0.1:8000/runestone/admin/index'>Instructor's Page</a></li>
            <li><a href='http://127.0.0.1:8000/runestone/assignments/index'>Progress Page</a></li>
            <li class="divider"></li>
            <li><a href="http://127.0.0.1:8000/runestone/default/user/profile" id="profilelink">Edit Profile</a></li>
            <li><a href="http://127.0.0.1:8000/runestone/default/user/change_password" id="passwordlink">Change Password</a></li>
            <li><a href="http://127.0.0.1:8000/runestone/default/user/register" id="registerlink">Register</a></li>
            <li class='loginout'><a href='#'>Login</a></li> <!-- correct link populated by addNavbarLoginLink() -->
          </ul>
        </li>
        <!-- end user account dropdown -->
        
        <li class="divider-vertical"></li>
        <!-- <li id="scratch_ac_link"><a href="javascript:ACFactory.toggleScratchActivecode()">Scratch ActiveCode</a></li> -->

        <!-- <li class="dropdown">             -->
          <li id="scratch_ac_link" class="dropdown"><a href="javascript:ACFactory.toggleScratchActivecode()">
              <i class="glyphicon glyphicon-pencil" style="opacity:0.9;"><span aria-label="Scratch Activecode" class="visuallyhidden">Scratch Activecode</span></i></a></li>
        <!-- </li> -->

        <li class="divider-vertical"></li>

        <!-- help menu dropdown -->
        <li class="dropdown">
          <a class="dropdown-toggle" href="#" data-toggle="dropdown">
            <i class="glyphicon glyphicon-question-sign" style="opacity:0.9;"><span aria-label="Help" class="visuallyhidden">Help</span></i>
          </a>
          <ul class="dropdown-menu user-menu">
            <li><a href='/runestone/static/fopp/navhelp.html'>Navigation Help</a></li>
            <li><a href='/runestone/static/overview/instructor.html'>Help for Instructors</a></li>
            <li class="divider"></li>
            <li><a href='http://runestoneinteractive.org'>About Runestone</a></li>
            <li><a href='/runestone/default/reportabug?course=fopp&page=TheAccumulatorPattern'>Report A Problem</a></li>
          </ul>
        </li>
        <!-- end help menu dropdown -->

        <li class="divider-vertical"></li>

      </ul>

      <ul class="nav navbar-nav">
        <li class="divider-vertical"></li>
        
          <script>
</script>
<li class="dropdown globaltoc-container">
  <a href="../index.html"
     class="dropdown-toggle"
     data-toggle="dropdown">Chapters <b class="caret"></b></a>
  <ul class="dropdown-menu globaltoc"
    ><ul>
<li class="toctree-l1"><a class="reference internal" href="../FrontBackMatter/preface.html">Preface</a></li>
</ul>
<ul class="current">
<li class="toctree-l1"><a class="reference internal" href="../GeneralIntro/toctree.html">1. General Introduction</a></li>
<li class="toctree-l1"><a class="reference internal" href="../SimplePythonData/toctree.html">2. Variables, Statements, and Expressions</a></li>
<li class="toctree-l1"><a class="reference internal" href="../Debugging/toctree.html">3. Debugging</a></li>
<li class="toctree-l1"><a class="reference internal" href="../PythonModules/toctree.html">4. Python Modules</a></li>
<li class="toctree-l1"><a class="reference internal" href="../PythonTurtle/toctree.html">5. Python Turtle</a></li>
<li class="toctree-l1"><a class="reference internal" href="../Sequences/toctree.html">6. Sequences</a></li>
<li class="toctree-l1 current"><a class="reference internal" href="toctree.html">7. Iteration</a></li>
<li class="toctree-l1"><a class="reference internal" href="../Conditionals/toctree.html">8. Conditionals</a></li>
<li class="toctree-l1"><a class="reference internal" href="../TransformingSequences/toctree.html">9. Transforming Sequences</a></li>
<li class="toctree-l1"><a class="reference internal" href="../Files/toctree.html">10. Files</a></li>
<li class="toctree-l1"><a class="reference internal" href="../Dictionaries/toctree.html">11. Dictionaries</a></li>
<li class="toctree-l1"><a class="reference internal" href="../Functions/toctree.html">12. Functions</a></li>
<li class="toctree-l1"><a class="reference internal" href="../Tuples/toctree.html">13. Tuple Packing and Unpacking</a></li>
<li class="toctree-l1"><a class="reference internal" href="../MoreAboutIteration/toctree.html">14. More About Iteration</a></li>
<li class="toctree-l1"><a class="reference internal" href="../AdvancedFunctions/toctree.html">15. Advanced Functions</a></li>
<li class="toctree-l1"><a class="reference internal" href="../Sorting/toctree.html">16. Sorting</a></li>
<li class="toctree-l1"><a class="reference internal" href="../NestedData/toctree.html">17. Nested Data and Nested Iteration</a></li>
<li class="toctree-l1"><a class="reference internal" href="../TestCases/toctree.html">18. Test Cases</a></li>
<li class="toctree-l1"><a class="reference internal" href="../Exceptions/toctree.html">19. Exceptions</a></li>
<li class="toctree-l1"><a class="reference internal" href="../Classes/toctree.html">20. Defining your own Classes</a></li>
<li class="toctree-l1"><a class="reference internal" href="../BuildingPrograms/toctree.html">21. Building Programs</a></li>
<li class="toctree-l1"><a class="reference internal" href="../Inheritance/toctree.html">22. Inheritance</a></li>
<li class="toctree-l1"><a class="reference internal" href="../AdvancedAccumulation/toctree.html">23. More on Accumulation: Map, Filter, List Comprehension, and Zip</a></li>
</ul>
<p class="caption"><span class="caption-text">Acknowledgements</span></p>
<ul>
<li class="toctree-l1"><a class="reference internal" href="../FrontBackMatter/copyright.html">Copyright Notice</a></li>
<li class="toctree-l1"><a class="reference internal" href="../FrontBackMatter/foreword.html">Foreword</a></li>
<li class="toctree-l1"><a class="reference internal" href="../FrontBackMatter/contrib.html">Contributor List</a></li>
<li class="toctree-l1"><a class="reference internal" href="../FrontBackMatter/fdl-1.3.html">GNU Free Documentation License</a></li>
<li class="toctree-l1"><a class="reference internal" href="../FrontBackMatter/fdl-1.3.html#addendum-how-to-use-this-license-for-your-documents">ADDENDUM: How to use this License for your documents</a></li>
</ul>
</ul>
</li>
          <li class="divider-vertical"></li>
        
        <!-- 
          
  <li id="relations-prev" title="Previous Chapter - 7.5. Lists and <code class="docutils literal notranslate"><span class="pre">for</span></code> loops" data-toggle="tooltip">
    <a href="Listsandforloops.html" >
      <i class='glyphicon glyphicon-backward' style='opacity:0.9;'><span aria-label="previous page" class="visuallyhidden">previoius page</span></i>
    </a>
  </li>
  
  <li id="relations-next" title='Next Chapter - 7.7. Traversal and the <code class="docutils literal notranslate"><span class="pre">for</span></code> Loop: By Index' data-toggle="tooltip" >              
    <a href="TraversalandtheforLoopByIndex.html" >
        <i class='glyphicon glyphicon-forward' style='opacity:0.9;'><span aria-label="next page" class="visuallyhidden">next page</span></i>
    </a>
  </li>
  <li class="divider-vertical"></li>

<script type="text/javascript">
  opts = {'placement':'bottom',
          'selector': '',
          'delay': { show: 100, hide: 50}
         };

  $('#relations-prev').tooltip(opts);
  $('#relations-next').tooltip(opts);
</script>
        -->
        
          <li></li>
        
      </ul>

      </div> <!-- navbar -->






      
    </div> <!-- container -->
  </div>
</div>


<div class="container" id="continue-reading"></div>

<div class="container" id="main-content">
  
  <div class="section" id="the-accumulator-pattern">
<span id="accum-pattern"></span><h1>7.6. The Accumulator Pattern<a class="headerlink" href="#the-accumulator-pattern" title="Permalink to this headline">¶</a></h1>
<p>One common programming “pattern” is to traverse a sequence, <strong>accumulating</strong> a value as we go,
such as the sum-so-far or the maximum-so-far. That way, at the end of the traversal we have
accumulated a single value, such as the sum total of all the items or the largest item.</p>
<dl class="docutils">
<dt>The anatomy of the accumulation pattern includes:</dt>
<dd><ul class="first last simple">
<li><strong>initializing</strong> an “accumulator” variable to an initial value (such as 0 if accumulating a sum)</li>
<li><strong>iterating</strong> (e.g., traversing the items in a sequence)</li>
<li><strong>updating</strong> the accumulator variable on each iteration (i.e., when processing each item in the sequence)</li>
</ul>
</dd>
</dl>
<p>For example, consider the following code, which computes the sum of the numbers in a list.</p>

<div data-childcomponent="ac6_6_1" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ac6_6_1 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >
nums = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10]
accum = 0
for w in nums:
    accum = accum + w
print(accum)
</textarea>
</div>
<p>In the program above, notice that the variable <code class="docutils literal notranslate"><span class="pre">accum</span></code> starts out with a value of 0.
Next, the iteration is performed 10 times.  Inside the for loop, the update occurs.
<code class="docutils literal notranslate"><span class="pre">w</span></code> has the value of current item (1 the first time, then 2, then 3, etc.).
<code class="docutils literal notranslate"><span class="pre">accum</span></code> is reassigned a new value which is the old value plus the current value of <code class="docutils literal notranslate"><span class="pre">w</span></code>.</p>
<p>This pattern of iterating the updating of a variable is commonly referred to as the
<strong>accumulator pattern</strong>. We refer to the variable as the <strong>accumulator</strong>. This pattern will come up
over and over again. Remember that the key to making it work successfully is to be sure to
initialize the variable before you start the iteration. Once inside the iteration, it is required
that you update the accumulator.</p>
<p>Here is the same program in codelens.  Step through the function and watch the “running total”
accumulate the result.</p>

<div class="runestone" style="max-width: none;">
<div class="alert alert-warning cd_section">
<div id="clens6_6_1"></div>
<p class="cl_caption"><span class="cl_caption_text"> (clens6_6_1)</span> </p>
</div>
<script type="text/javascript">
var clens6_6_1_trace = {"code": "nums = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10]\naccum = 0\nfor w in nums:\n   accum = accum + w\nprint(accum)", "trace": [{"ordered_globals": [], "stdout": "", "func_name": "<module>", "stack_to_render": [], "globals": {}, "heap": {}, "line": 1, "event": "step_line"}, {"ordered_globals": ["nums"], "stdout": "", "func_name": "<module>", "stack_to_render": [], "globals": {"nums": ["REF", 1]}, "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "line": 2, "event": "step_line"}, {"ordered_globals": ["nums", "accum"], "stdout": "", "func_name": "<module>", "stack_to_render": [], "globals": {"accum": 0, "nums": ["REF", 1]}, "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "line": 3, "event": "step_line"}, {"ordered_globals": ["nums", "accum", "w"], "stdout": "", "func_name": "<module>", "stack_to_render": [], "globals": {"accum": 0, "w": 1, "nums": ["REF", 1]}, "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "line": 4, "event": "step_line"}, {"ordered_globals": ["nums", "accum", "w"], "stdout": "", "func_name": "<module>", "stack_to_render": [], "globals": {"accum": 1, "w": 1, "nums": ["REF", 1]}, "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "line": 3, "event": "step_line"}, {"ordered_globals": ["nums", "accum", "w"], "stdout": "", "func_name": "<module>", "stack_to_render": [], "globals": {"accum": 1, "w": 2, "nums": ["REF", 1]}, "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "line": 4, "event": "step_line"}, {"ordered_globals": ["nums", "accum", "w"], "stdout": "", "func_name": "<module>", "stack_to_render": [], "globals": {"accum": 3, "w": 2, "nums": ["REF", 1]}, "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "line": 3, "event": "step_line"}, {"ordered_globals": ["nums", "accum", "w"], "stdout": "", "func_name": "<module>", "stack_to_render": [], "globals": {"accum": 3, "w": 3, "nums": ["REF", 1]}, "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "line": 4, "event": "step_line"}, {"ordered_globals": ["nums", "accum", "w"], "stdout": "", "func_name": "<module>", "stack_to_render": [], "globals": {"accum": 6, "w": 3, "nums": ["REF", 1]}, "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "line": 3, "event": "step_line"}, {"ordered_globals": ["nums", "accum", "w"], "stdout": "", "func_name": "<module>", "stack_to_render": [], "globals": {"accum": 6, "w": 4, "nums": ["REF", 1]}, "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "line": 4, "event": "step_line"}, {"ordered_globals": ["nums", "accum", "w"], "stdout": "", "func_name": "<module>", "stack_to_render": [], "globals": {"accum": 10, "w": 4, "nums": ["REF", 1]}, "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "line": 3, "event": "step_line"}, {"ordered_globals": ["nums", "accum", "w"], "stdout": "", "func_name": "<module>", "stack_to_render": [], "globals": {"accum": 10, "w": 5, "nums": ["REF", 1]}, "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "line": 4, "event": "step_line"}, {"ordered_globals": ["nums", "accum", "w"], "stdout": "", "func_name": "<module>", "stack_to_render": [], "globals": {"accum": 15, "w": 5, "nums": ["REF", 1]}, "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "line": 3, "event": "step_line"}, {"ordered_globals": ["nums", "accum", "w"], "stdout": "", "func_name": "<module>", "stack_to_render": [], "globals": {"accum": 15, "w": 6, "nums": ["REF", 1]}, "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "line": 4, "event": "step_line"}, {"ordered_globals": ["nums", "accum", "w"], "stdout": "", "func_name": "<module>", "stack_to_render": [], "globals": {"accum": 21, "w": 6, "nums": ["REF", 1]}, "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "line": 3, "event": "step_line"}, {"ordered_globals": ["nums", "accum", "w"], "stdout": "", "func_name": "<module>", "stack_to_render": [], "globals": {"accum": 21, "w": 7, "nums": ["REF", 1]}, "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "line": 4, "event": "step_line"}, {"ordered_globals": ["nums", "accum", "w"], "stdout": "", "func_name": "<module>", "stack_to_render": [], "globals": {"accum": 28, "w": 7, "nums": ["REF", 1]}, "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "line": 3, "event": "step_line"}, {"ordered_globals": ["nums", "accum", "w"], "stdout": "", "func_name": "<module>", "stack_to_render": [], "globals": {"accum": 28, "w": 8, "nums": ["REF", 1]}, "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "line": 4, "event": "step_line"}, {"ordered_globals": ["nums", "accum", "w"], "stdout": "", "func_name": "<module>", "stack_to_render": [], "globals": {"accum": 36, "w": 8, "nums": ["REF", 1]}, "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "line": 3, "event": "step_line"}, {"ordered_globals": ["nums", "accum", "w"], "stdout": "", "func_name": "<module>", "stack_to_render": [], "globals": {"accum": 36, "w": 9, "nums": ["REF", 1]}, "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "line": 4, "event": "step_line"}, {"ordered_globals": ["nums", "accum", "w"], "stdout": "", "func_name": "<module>", "stack_to_render": [], "globals": {"accum": 45, "w": 9, "nums": ["REF", 1]}, "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "line": 3, "event": "step_line"}, {"ordered_globals": ["nums", "accum", "w"], "stdout": "", "func_name": "<module>", "stack_to_render": [], "globals": {"accum": 45, "w": 10, "nums": ["REF", 1]}, "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "line": 4, "event": "step_line"}, {"ordered_globals": ["nums", "accum", "w"], "stdout": "", "func_name": "<module>", "stack_to_render": [], "globals": {"accum": 55, "w": 10, "nums": ["REF", 1]}, "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "line": 3, "event": "step_line"}, {"ordered_globals": ["nums", "accum", "w"], "stdout": "", "func_name": "<module>", "stack_to_render": [], "globals": {"accum": 55, "w": 10, "nums": ["REF", 1]}, "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "line": 5, "event": "step_line"}, {"ordered_globals": ["nums", "accum", "w"], "stdout": "55\n", "func_name": "<module>", "stack_to_render": [], "globals": {"accum": 55, "w": 10, "nums": ["REF", 1]}, "heap": {"1": ["LIST", 1, 2, 3, 4, 5, 6, 7, 8, 9, 10]}, "line": 5, "event": "return"}]};
var clens6_6_1_vis;

$(document).ready(function() {
    try {
        clens6_6_1_vis = new ExecutionVisualizer('clens6_6_1',clens6_6_1_trace,
                                    {embeddedMode: true,
                                    verticalStack: false,
                                    heightChangeCallback: redrawAllVisualizerArrows,
                                    codeDivWidth: 500,
                                    lang : 'py3'
                                    });
        attachLoggers(clens6_6_1_vis,'clens6_6_1');
        styleButtons('clens6_6_1');
        allVisualizers.push(clens6_6_1_vis);
    } catch (e) {
        console.log("Failed to Initialize CodeLens component clens6_6_1_vis" );
        console.log(e.toString());
    }

});

$(document).ready(function() {
    $("#clens6_6_1_tracecheck").click(function() {
        logBookEvent({'event':'codelens', 'act': 'check', 'div_id':'clens6_6_1'});
    });
});

if (allVisualizers === undefined) {
   var allVisualizers = [];
}


$(window).resize(function() {
    if (clens6_6_1_vis) {
        clens6_6_1_vis.redrawConnectors();
    }
});
</script>
</div>
<div class="admonition note">
<p class="first admonition-title">Note</p>
<p class="last">What would happen if we indented the print accum statement? Not sure? Make a prediction, then try it and find out.</p>
</div>
<p>We can utilize the range function in this situation as well. Previously, you’ve seen it used when we wanted to draw in
turtle. There we used it to iterate a certain number of times. We can do more than that though. The <code class="docutils literal notranslate"><span class="pre">range</span></code> function
takes at least one input - which should be an integer - and returns a list as long as your input. While you can provide
two inputs, we will focus on using range with just one input. With one input, range will start at zero and go up to - but
not include - the input.</p>

<div class="highlight-default notranslate"><div class="highlight"><pre><span></span>
  # Notice the casting of `range` to the `list`
  print(list(range(5))) # This will print [0, 1, 2, 3, 4]

  print(list(range(0,5))) # This will print [0, 1, 2, 3, 4]

  # Note: `range` function is already casted as `list` in the textbook.
  print(range(5)) # This will also print [0, 1, 2, 3, 4]
  </pre></div>
</div>

<p>One important thing to know about the range function in python3 is that if we want to use it outside of iteration, we
have to cast it as a list using <code class="docutils literal notranslate"><span class="pre">list()</span></code>. Inside the textbook you’ll notice that <code class="docutils literal notranslate"><span class="pre">range</span></code> works with or without
casting it as a list but it is best for you to try and get into the habit of casting it as a list. Here’s how you could use the range function in the previous problem.</p>

<div data-childcomponent="ac6_6_2" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ac6_6_2 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >
accum = 0
for w in range(11):
    accum = accum + w
print(accum)

# or, if you use two inputs for the range function

sec_accum = 0
for w in range(1,11):
    sec_accum = sec_accum + w
print(sec_accum)
</textarea>
</div>
<p>Because the the range function is exclusive of the ending number, we have to use 11 as the function input.</p>
<p>We can use the accumulation pattern is count the number of something or to sum up a total. The
above examples only covered how to get the sum for a list, but we can also count how many items are
in the list if we wanted to.</p>

<div data-childcomponent="ac6_6_3" class="runestone explainer ac_section alert alert-warning">

<textarea data-component="activecode" id=ac6_6_3 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >
nums = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10]
count = 0
for w in nums:
    count = count + 1
print(count)
</textarea>
</div>
<p>In this example we don’t make use of <code class="docutils literal notranslate"><span class="pre">w</span></code> even though the iterator variable is a necessary part of
constructing a for loop. Instead of adding the value of <code class="docutils literal notranslate"><span class="pre">w</span></code> to <code class="docutils literal notranslate"><span class="pre">count</span></code> we add a 1 to it,
because we’re incrementing the value of count when we iterate each time through the loop. Though in
this scenario we could have used the <code class="docutils literal notranslate"><span class="pre">len</span></code> function, there are other cases later on where len
won’t be useful but we will still need to count.</p>
<p><strong>Check your understanding</strong></p>

            <div class="runestone alert alert-warning">
            <ul data-component="multiplechoice" data-multipleanswers="false"  id="question6_6_1">
            <p>iter-6-1: Consider the following code:</p>
<div class="highlight-python notranslate"><div class="highlight"><pre><span></span><span class="n">nums</span> <span class="o">=</span> <span class="p">[</span><span class="mi">1</span><span class="p">,</span> <span class="mi">2</span><span class="p">,</span> <span class="mi">3</span><span class="p">,</span> <span class="mi">4</span><span class="p">,</span> <span class="mi">5</span><span class="p">,</span> <span class="mi">6</span><span class="p">,</span> <span class="mi">7</span><span class="p">,</span> <span class="mi">8</span><span class="p">,</span> <span class="mi">9</span><span class="p">,</span> <span class="mi">10</span><span class="p">]</span>
<span class="k">for</span> <span class="n">w</span> <span class="ow">in</span> <span class="n">nums</span><span class="p">:</span>
   <span class="n">accum</span> <span class="o">=</span> <span class="mi">0</span>
   <span class="n">accum</span> <span class="o">=</span> <span class="n">accum</span> <span class="o">+</span> <span class="n">w</span>
<span class="k">print</span><span class="p">(</span><span class="n">accum</span><span class="p">)</span>
</pre></div>
</div>
<p>What happens if you put the initialization of accum inside the for loop as the first
instruction in the loop?</p>

            <li data-component="answer" data-correct id="question6_6_1_opt_a">It will print out 10 instead of 55</li><li data-component="feedback" id="question6_6_1_opt_a">The variable accum will be reset to 0 each time through the loop. Then it will add the current item. Only the last item will count.</li>
            
            <li data-component="answer"  id="question6_6_1_opt_b">It will cause a run-time error</li><li data-component="feedback" id="question6_6_1_opt_b">Assignment statements are perfectly legal inside loops and will not cause an error.</li>
            
            <li data-component="answer"  id="question6_6_1_opt_c">It will print out 0 instead of 55</li><li data-component="feedback" id="question6_6_1_opt_c">Good thought: the variable accum will be reset to 0 each time through the loop. But then it adds the current item.</li>
            

            </ul>
            </div>
            
        <div class="runestone" style="max-width: none;">
        <pre data-component="parsons" id="pp6_6_1"      >
        <span data-question>iter-6-2: Rearrange the code statements so that the program will add up the first n odd numbers where n is provided by the user.</span>n = int(input('How many odd numbers would you like to add together?'))
thesum = 0
oddnumber = 1
---
for counter in range(n):
---
   thesum = thesum + oddnumber
   oddnumber = oddnumber + 2
---
print(thesum)
        </pre>
        </div>
    
<div data-childcomponent="ac6_6_4" class="runestone explainer ac_section alert alert-warning">
Write code to create a list of integers from 0 through 52 and assign that list to the variable <code class="docutils literal notranslate"><span class="pre">numbers</span></code>. You should use a special Python function – do not type out the whole list yourself. HINT: You can do this in one line of code!
<textarea data-component="activecode" id=ac6_6_4 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(numbers, range(53), "Testing that numbers is a list that contains the correct elements.")

myTests().main()
</textarea>
</div>

<div data-childcomponent="ac6_6_10" class="runestone explainer ac_section alert alert-warning">
Count the number of characters in string <code class="docutils literal notranslate"><span class="pre">str1</span></code>. Do not use <code class="docutils literal notranslate"><span class="pre">len()</span></code>. Save the number in variable <code class="docutils literal notranslate"><span class="pre">numbs</span></code>.
<textarea data-component="activecode" id=ac6_6_10 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >
str1 = "I like nonsense, it wakes up the brain cells. Fantasy is a necessary ingredient in living."

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testEight(self):
      self.assertEqual(numbs, 90, "Testing that numbs is assigned to correct values.")
      self.assertNotIn("len(", self.getEditorText(), "Testing your code (Don't worry about actual and expected values).")

myTests().main()
</textarea>
</div>

<div data-childcomponent="ac6_8_9" class="runestone explainer ac_section alert alert-warning">
Create a list of numbers 0 through 40 and assign this list to the variable <code class="docutils literal notranslate"><span class="pre">numbers</span></code>. Then, accumulate the total of the list’s values and assign that sum to the variable <code class="docutils literal notranslate"><span class="pre">sum1</span></code>.
<textarea data-component="activecode" id=ac6_8_9 data-lang="python" 
      data-timelimit=25000  data-codelens="true"  
    data-audio=''   
          >

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testNineA(self):
      self.assertEqual(numbers, [0, 1, 2, 3, 4, 5, 6, 7, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22, 23, 24, 25, 26, 27, 28, 29, 30, 31, 32, 33, 34, 35, 36, 37, 38, 39, 40], "Testing that numbers is assigned to correct values.")

   def testNineB(self):
      self.assertEqual(sum1, 820, "Testing that sum1 has the correct value.")

myTests().main()
</textarea>
</div>
</div>


  
      
  <li id="relations-prev" class="navLink" title='Previous Section - 7.5. Lists and <code class="docutils literal notranslate"><span class="pre">for</span></code> loops' data-toggle="tooltip">
    <a href="Listsandforloops.html" >
        <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">            
            <i class='prevNav glyphicon glyphicon-chevron-left' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%;"></i>
        </div>
    </a>
  </li>
  <a class="navLinkBg" id="navLinkBgLeft"  href="Listsandforloops.html" ></a>
  
  <li id="relations-next" class="navLink" title='Next Section - 7.7. Traversal and the <code class="docutils literal notranslate"><span class="pre">for</span></code> Loop: By Index' data-toggle="tooltip" >
    <a href="TraversalandtheforLoopByIndex.html" >
      <div style="background-color: white; border-style:solid; border-color:lightgrey; border-width:2px; width:100px; height:50px">            
          <i id="relationsNextIcon" class='nextNav glyphicon glyphicon-chevron-right' style="top:50%; transform:translateY(-50%)translateX(-50%); left: 50%;"></i>
      </div>
    </a>
  </li>
  <a class="navLinkBg navLink" id="navLinkBgRight" href="TraversalandtheforLoopByIndex.html" >Next Section - 7.7. Traversal and the <code class="docutils literal notranslate"><span class="pre">for</span></code> Loop: By Index</a>

<script type="text/javascript">

  $('#relations-prev').tooltip({'placement':'right', 'selector': '', 'delay': { show: 100, hide: 50}});
  $('#relations-next').tooltip({'placement':'left', 'selector': '', 'delay': { show: 100, hide: 50}});
  
</script>
  
</div>
<footer class="footer">
  <div class="container">
    <p class="pull-right">
      
      <span id='numuserspan'></span> readers online now | <span class='loggedinuser'></span>
      
      | <a href="#">Back to top</a>
      
    </p>
    <p>
        &copy; Copyright 2017 bradleymiller.
      Created using <a href="http://runestoneinteractive.org/">Runestone</a> 3.2.11.
    </p>
  </div>
</footer>



<script type="text/javascript">
  var _gaq = _gaq || [];
  _gaq.push(['_setAccount', 'UA-32029811-1']);
  _gaq.push(['_trackPageview']);

  (function() {
    var ga = document.createElement('script'); ga.type = 'text/javascript'; ga.async = true;
    ga.src = ('https:' == document.location.protocol ? 'https://ssl' : 'http://www') + '.google-analytics.com/ga.js';
    var s = document.getElementsByTagName('script')[0]; s.parentNode.insertBefore(ga, s);
  })();
</script>
<script async src="https://hypothes.is/embed.js"></script>




  </body>
</html>