<?xml version="1.0"?>
<section xml:id="more-about-iteration_the-while-statement">
  <title>The <c>while</c> Statement</title>
  <video xml:id="whileloop" youtube="iUbToOoAykE" width="auto"/>
  <p>There is another Python statement that can also be used to build an iteration. It is called the <c>while</c> statement.
            The <c>while</c> statement provides a much more general mechanism for iterating. Similar to the <c>if</c> statement, it uses
            a boolean expression to control the flow of execution. The body of while will be repeated as long as the controlling
            boolean expression evaluates to <c>True</c>.</p>
  <p>The following two figures show the flow of control. The first focuses on the flow inside the while loop and the second
            shows the while loop in context.</p>
  <image source="MoreAboutIteration/Figures/while_flow.png" width="50%" alt="a diamond at the top has the phrase &quot;Is the condition True?&quot;. Two arrows come out it with either the phrase yes or no on the arrows. The yes arrow points to a box that says &quot;evaluate the statemenets in the body of the loop&quot;. It then has an arrow that unconditionally points back to &quot;Is the condition True?&quot; diamond. The no arrow escapes the loop and points down past the &quot;evaluate&quot; square."/>
  <image source="MoreAboutIteration/Figures/while_loop.png" width="50%" alt="image showing a rectangle with &quot;code block&quot; written on it on top. Then, text that read &quot;while {condition}&quot;: followed by an indented block with &quot;run if {condition} is True&quot; written on it. An arrow points from the bottom of the indented block to the top of the while loop and says &quot;run loop again&quot;. At the bottom of the image is an unindented block that says the phrase &quot;code block.&quot;"/>
  <p>We can use the <c>while</c> loop to create any type of iteration we wish, including anything that we have previously
            done with a <c>for</c> loop. For example, the program in the previous section could be rewritten using <c>while</c>.
            Instead of relying on the <c>range</c> function to produce the numbers for our summation, we will need to produce them
            ourselves. To do this, we will create a variable called <c>aNumber</c> and initialize it to 1, the first number in the
            summation. Every iteration will add <c>aNumber</c> to the running total until all the values have been used. In order to
            control the iteration, we must create a boolean expression that evaluates to <c>True</c> as long as we want to keep
            adding values to our running total. In this case, as long as <c>aNumber</c> is less than or equal to the bound, we
            should keep going.</p>
  <p>Here is a new version of the summation program that uses a while statement.</p>
  <program xml:id="ac14_2_1" interactive="activecode" language="python">
    <input>
def sumTo(aBound):
    """ Return the sum of 1+2+3 ... n """

    theSum  = 0
    aNumber = 1
    while aNumber &lt;= aBound:
        theSum = theSum + aNumber
        aNumber = aNumber + 1
    return theSum

print(sumTo(4))

print(sumTo(1000))
        </input>
  </program>
  <p>You can almost read the <c>while</c> statement as if it were in natural language. It means, while <c>aNumber</c> is less
            than or equal to <c>aBound</c>, continue executing the body of the loop. Within the body, each time, update <c>theSum</c>
            using the accumulator pattern and increment <c>aNumber</c>. After the body of the loop, we go back up to the condition
            of the <c>while</c> and reevaluate it. When <c>aNumber</c> becomes greater than <c>aBound</c>, the condition fails and flow
            of control continues to the <c>return</c> statement.</p>
  <p>The same program in codelens will allow you to observe the flow of execution.</p>
  <program xml:id="clens14_2_1" interactive="codelens" language="python">
    <input>
def sumTo(aBound):
    """ Return the sum of 1+2+3 ... n """

    theSum  = 0
    aNumber = 1
    while aNumber &lt;= aBound:
        theSum = theSum + aNumber
        aNumber = aNumber + 1
    return theSum

print(sumTo(4))
    </input>
  </program>
  <note>
    <p>The names of the variables have been chosen to help readability.</p>
  </note>
  <p>More formally, here is the flow of execution for a <c>while</c> statement:</p>
  <p>
    <ol label="1">
      <li>
        <p>Evaluate the condition, yielding <c>False</c> or <c>True</c>.</p>
      </li>
      <li>
        <p>If the condition is <c>False</c>, exit the <c>while</c> statement and continue
                    execution at the next statement.</p>
      </li>
      <li>
        <p>If the condition is <c>True</c>, execute each of the statements in the body and
                    then go back to step 1.</p>
      </li>
    </ol>
  </p>
  <p>The body consists of all of the statements below the header with the same indentation.</p>
  <p>This type of flow is called a <term>loop</term> because the third step loops back around to the top. Notice that if the
            condition is <c>False</c> the first time through the loop, the statements inside the loop are never executed.</p>
  <p>The body of the loop should change the value of one or more variables so that eventually the condition becomes
            <c>False</c> and the loop terminates. Otherwise the loop will repeat forever. This is called an <term>infinite loop</term>.
            An endless source of amusement for computer scientists is the observation that the directions written on the back of
            the shampoo bottle (lather, rinse, repeat) create an infinite loop.</p>
  <p>In the case shown above, we can prove that the loop terminates because we know that the value of <c>aBound</c> is
            finite, and we can see that the value of <c>aNumber</c> increments each time through the loop, so eventually it will
            have to exceed <c>aBound</c>. In other cases, it is not so easy to tell.</p>
  <note>
    <p>Introduction of the while statement causes us to think about the types of iteration we have seen. The <c>for</c>
                statement will always iterate through a sequence of values like the list of names for the party or the list of
                numbers created by <c>range</c>. Since we know that it will iterate once for each value in the collection, it is
                often said that a <c>for</c> loop creates a <term>definite iteration</term> because we definitely know how many times we are
                going to iterate.  On the other hand, the <c>while</c> statement is dependent on a condition that needs to evaluate
                to <c>False</c> in order for the loop to terminate. Since we do not necessarily know when this will happen, it
                creates what we call <term>indefinite iteration</term>. Indefinite iteration simply means that we don't know how many
                times we will repeat but eventually the condition controlling the iteration will fail and the iteration will
                stop. (Unless we have an infinite loop which is of course a problem)</p>
  </note>
  <p>What you will notice here is that the <c>while</c> loop is more work for you &#x2014; the programmer &#x2014; than the equivalent
            <c>for</c> loop.  When using a <c>while</c> loop you have to control the loop variable yourself. You give it an initial
            value, test for completion, and then make sure you change something in the body so that the loop terminates. That
            also makes a while loop harder to read and understand than the equivalent for loop. So, while you <em>can</em> implement
            definite iteration with a while loop, it's not a good idea to do that. Use a for loop whenever it will be known at
            the beginning of the iteration process how many times the block of code needs to be executed.</p>
  <p>
    <term>Check your understanding</term>
  </p>
  <exercise label="question14_2_1">
    <statement>
      <p>True or False: You can rewrite any for-loop as a while-loop.</p>
    </statement>
    <choices>
      <choice correct="yes">
        <statement>
          <p>True</p>
        </statement>
        <feedback>
          <p>Although the while loop uses a different syntax, it is just as powerful as a for-loop and often more flexible.</p>
        </feedback>
      </choice>
      <choice>
        <statement>
          <p>False</p>
        </statement>
        <feedback>
          <p>Often a for-loop is more natural and convenient for a task, but that same task can always be expressed using a while loop.</p>
        </feedback>
      </choice>
    </choices>
  </exercise>
  <exercise label="question14_2_2">
    <statement>
      <p>The following code contains an infinite loop. Which is the best explanation for why the loop does not terminate?</p>
      <program language="python">
        <input>
n = 10
answer = 1
while ( n &gt; 0 ):
  answer = answer + n
  n = n + 1
print(answer)
</input>
      </program>
    </statement>
    <choices>
      <choice correct="yes">
        <statement>
          <p>n starts at 10 and is incremented by 1 each time through the loop, so it will always be positive</p>
        </statement>
        <feedback>
          <p>The loop will run as long as n is positive.  In this case, we can see that n will never become non-positive.</p>
        </feedback>
      </choice>
      <choice>
        <statement>
          <p>answer starts at 1 and is incremented by n each time, so it will always be positive</p>
        </statement>
        <feedback>
          <p>While it is true that answer will always be positive, answer is not considered in the loop condition.</p>
        </feedback>
      </choice>
      <choice>
        <statement>
          <p>You cannot compare n to 0 in while loop.  You must compare it to another variable.</p>
        </statement>
        <feedback>
          <p>It is perfectly valid to compare n to 0.  Though indirectly, this is what causes the infinite loop.</p>
        </feedback>
      </choice>
      <choice>
        <statement>
          <p>In the while loop body, we must set n to False, and this code does not do that.</p>
        </statement>
        <feedback>
          <p>The loop condition must become False for the loop to terminate, but n by itself is not the condition in this case.</p>
        </feedback>
      </choice>
    </choices>
  </exercise>
  <exercise label="question14_2_3">
    <statement>
      <p>Which type of loop can be used to perform the following iteration: You choose a positive integer at random and then print the numbers from 1 up to and including the selected integer.</p>
    </statement>
    <choices>
      <choice correct="yes">
        <statement>
          <p>a for-loop or a while-loop</p>
        </statement>
        <feedback>
          <p>Although you do not know how many iterations you loop will run before the program starts running, once you have chosen your random integer, Python knows exactly how many iterations the loop will run, so either a for-loop or a while-loop will work.</p>
        </feedback>
      </choice>
      <choice>
        <statement>
          <p>only a for-loop</p>
        </statement>
        <feedback>
          <p>As you learned in section 7.2, a while-loop can always be used for anything a for-loop can be used for.</p>
        </feedback>
      </choice>
      <choice>
        <statement>
          <p>only a while-loop</p>
        </statement>
        <feedback>
          <p>Although you do not know how many iterations you loop will run before the program starts running, once you have chosen your random integer, Python knows exactly how many iterations the loop will run, so this is an example of definite iteration.</p>
        </feedback>
      </choice>
    </choices>
  </exercise>
  <exercise label="ac14_2_2">
    <statement>
      <p>Write a while loop that is initialized at 0 and stops at 15. If the counter is an even number, append the counter to a list called <c>eve_nums</c>.</p>
    </statement>
    <program xml:id="ac14_2_2_editor" interactive="activecode" language="python">
      <input>

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

  def testOne(self):
     self.assertEqual(eve_nums, [0,2,4,6,8,10,12,14], "Testing that eve_nums has been assigned the correct elements")

myTests().main()
        </input>
    </program>
  </exercise>
  <exercise label="ac14_2_3">
    <statement>
      <p>Below, we've provided a for loop that sums all the elements of <c>list1</c>. Write code that accomplishes the same task, but instead uses a while loop. Assign the accumulator variable to the name <c>accum</c>.</p>
    </statement>
    <program xml:id="ac14_2_3_editor" interactive="activecode" language="python">
      <input>

list1 = [8, 3, 4, 5, 6, 7, 9]

tot = 0
for elem in list1:
    tot = tot + elem

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(accum, 42, "Testing that accum has the correct value.")
      self.assertIn('while', self.getEditorText(), "Testing your code (Don't worry about actual and expected values).")


myTests().main()
        </input>
    </program>
  </exercise>
  <exercise label="ac14_2_4">
    <statement>
      <p>Write a function called <c>stop_at_four</c> that iterates through a list of numbers. Using a while loop, append each number to a new list until the number 4 appears. The function should return the new list.</p>
    </statement>
    <program xml:id="ac14_2_4_editor" interactive="activecode" language="python">
      <input>

def stop_at_four():



=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(stop_at_four([0, 9, 4.5, 1, 7, 4, 8, 9, 3]), [0, 9, 4.5, 1, 7], "Testing the function stop_at_four on the input [0, 9, 4.5, 1, 7, 4, 8, 9, 3].")
      self.assertEqual(stop_at_four([4, 1, 2, 8]), [], "Testing the function stop_at_four on the input [4, 1, 2, 8].")
      self.assertEqual(stop_at_four([4]), [], "Testing the function stop_at_four on the input [4].")
      self.assertEqual(stop_at_four([1, 6, 2, 3, 9]), [1, 6, 2, 3, 9], "Testing that stop_at_four([1, 6, 2, 3, 9]) returns ([1, 6, 2, 3, 9])")

myTests().main()
        </input>
    </program>
  </exercise>
</section>
