<?xml version="1.0"?>
<section xml:id="advanced-functions_introduction-optional-parameters">
  <title>Introduction: Optional Parameters</title>
  <p>In the treatment of functions so far, each function definition specifies zero or more formal parameters
            and each function invocation provides exactly that many values. Sometimes it is convenient to have
            <term>optional parameters</term> that can be specified or omitted. When an optional parameter is omitted from a function invocation, the formal parameter is bound to a <term>default value</term>. When the optional parameter is included, then
            the formal parameter is bound to the value provided. Optional parameters are convenient when a function is
            almost always used in a simple way, but it's nice to allow it to be used in a more complex way, with non-default
            values specified for the optional parameters.</p>
  <p>Consider, for example, the <c>int</c> function, which you have used previously. Its first parameter,
            which is required, specifies the object that you wish to convert to an integer. For example, if you
            call in on a string, <c>int("100")</c>, the return value will be the integer 100.</p>
  <p>That's the most common way programmers want to convert strings to integers. Sometimes, however, they
            are working with numbers in some other <q>base</q> rather than base 10. For example, in base 8, the rightmost
            digit is ones, the next digit to the left is 8s, and the one to the left of that is the 64s place (8**2).</p>
  <p>The int function provides an optional parameter for the base. When it is not specified, the number is
            converted to an integer assuming the original number was in base 10. We say that 10 is the default value.
            So <c>int("100")</c> is the same as invoking <c>int("100", 10)</c>. We can override the default of 10 by
            supplying a different value.</p>
  <program xml:id="ac15_1_1" interactive="activecode" language="python">
    <input>
print(int("100"))
print(int("100", 10))   # same thing, 10 is the default value for the base
print(int("100", 8))     # now the base is 8, so the result is 1*64 = 64
        </input>
  </program>
  <note>
    <p>Tom Lehrer's New Math</p>
    <p>Some math educators believe that elementary school students will get a much deeper understanding
                of the place-value system, and set a foundation for learning algebra later, if they learn to do
                arithmetic not only in base-10 but also in base-8 and other bases. This was part of a movement
                called <q>The New Math</q>, though it's not so new now (I had it when I was in elementary school!) Tom
                Lehrer made a really funny song about it, and it's set with visuals in several YouTube renditions
                now. Try this very nice <url href="http://www.youtube.com/watch?v=DfCJgC2zezw" visual="http://www.youtube.com/watch?v=DfCJgC2zezw">lip-synched version</url>.</p>
  </note>
  <p>When defining a function, you can specify a default value for a parameter. That parameter then becomes an
            optional parameter when the function is called. The way to specify a default value is with an assignment
            statement inside the parameter list. Consider the following code, for example.</p>
  <program xml:id="clens15_1_1" interactive="codelens" language="python">
    <input>
initial = 7
def f(x, y =3, z=initial):
    print("x, y, z, are: " + str(x) + ", " + str(y) + ", " + str(z))

f(2)
f(2, 5)
f(2, 5, 8)
    </input>
  </program>
  <p>Notice the different bindings of x, y, and z on the three invocations of f. The first time, y and z have
            their default values, 3 and 7. The second time, y gets the value 5 that is passed in, but z still gets the
            default value of 7. The last time, z gets the value 8 that is passed in.</p>
  <p>If you want to provide a non-default value for the third parameter (z), you also need to provide a value
            for the second item (y). We will see in the next section a mechanism called keyword parameters that lets you
            specify a value for z without specifying a value for y.</p>
  <note>
    <p>This is a second, related but slightly different use of = than we have seen previously. In a stand-alone assignment statement, not part of a function definition, <c>x=3</c> assigns 3 to the variable x. As part of specifying the parameters in a function definition, <c>x=3</c> says that 3 is the <em>default</em> value for x, used <em>only when</em> no value is provided during the function invocation.</p>
  </note>
  <p>There are two tricky things that can confuse you with default values. The first is that the default
            value is determined at the time that the function is defined, not at the time that it is invoked. So
            in the example above, if we wanted to invoke the function f with a value of 10 for z, we cannot simply
            set initial = 10 right before invoking f. See what happens in the code below, where z still gets the
            value 7 when f is invoked without specifying a value for z.</p>
  <program xml:id="clens15_1_2" interactive="codelens" language="python">
    <input>
initial = 7
def f(x, y =3, z=initial):
    print("x, y, z, are: " + str(x) + ", " + str(y) + ", " + str(z))

initial = 10
f(2)
    </input>
  </program>
  <p>The second tricky thing is that if the default value is set to a mutable object, such as a list or a dictionary,
            that object will be shared in all invocations of the function. This can get very confusing, so I suggest that you
            never set a default value that is a mutable object. For example, follow the exceution of this one carefully.</p>
  <program xml:id="opt_params_4" interactive="codelens" language="python">
    <input>
def f(a, L=[]):
    L.append(a)
    return L

print(f(1))
print(f(2))
print(f(3))
print(f(4, ["Hello"]))
print(f(5, ["Hello"]))
    </input>
  </program>
  <p>When the default value is used, the same list is shared. But on lines 8 and 9 two different copies of the list
            [<q>Hello</q>] are provided, so the 4 that is appended is not present in the list that is printed on line 9.</p>
  <p>
    <term>Check your understanding</term>
  </p>
  <exercise label="question15_1_1">
    <statement>
      <p>What will the following code print?</p>
      <program language="python">
        <input>
def f(x = 0, y = 1):
    return x * y

print(f())
</input>
      </program>
    </statement>
    <choices>
      <choice correct="yes">
        <statement>
          <p>0</p>
        </statement>
        <feedback>
          <p>Since no parameters are specified, x is 0 and y is 1, so 0 is returned.</p>
        </feedback>
      </choice>
      <choice>
        <statement>
          <p>1</p>
        </statement>
        <feedback>
          <p>0 * 1 is 0.</p>
        </feedback>
      </choice>
      <choice>
        <statement>
          <p>None</p>
        </statement>
        <feedback>
          <p>The function does return a value.</p>
        </feedback>
      </choice>
      <choice>
        <statement>
          <p>Runtime error since no parameters are passed in the call to f.</p>
        </statement>
        <feedback>
          <p>Because both parameters have default values specified in the definition, they are both optional.</p>
        </feedback>
      </choice>
    </choices>
  </exercise>
  <exercise label="question15_1_2">
    <statement>
      <p>What will the following code print?</p>
      <program language="python">
        <input>
def f(x = 0, y = 1):
    return x * y

print(f(1))
</input>
      </program>
    </statement>
    <choices>
      <choice>
        <statement>
          <p>0</p>
        </statement>
        <feedback>
          <p>Since one parameter value is specified, it is bound to x; y gets the default value of 1.</p>
        </feedback>
      </choice>
      <choice correct="yes">
        <statement>
          <p>1</p>
        </statement>
        <feedback>
          <p>Since one parameter value is specified, it is bound to x; y gets the default value of 1.</p>
        </feedback>
      </choice>
      <choice>
        <statement>
          <p>None</p>
        </statement>
        <feedback>
          <p>The function does return a value.</p>
        </feedback>
      </choice>
      <choice>
        <statement>
          <p>Runtime error since the second parameter value is missing.</p>
        </statement>
        <feedback>
          <p>Because both parameters have default values specified in the definition, they are both optional.</p>
        </feedback>
      </choice>
    </choices>
  </exercise>
  <exercise label="ac15_1_2">
    <statement>
      <p><term>3.</term> Write a function called <c>str_mult</c> that takes in a required string parameter and an optional integer parameter. The default value for the integer parameter should be 3. The function should return the string multiplied by the integer parameter.</p>
    </statement>
    <program xml:id="ac15_1_2_editor" interactive="activecode" language="python">
      <input>

=====

from unittest.gui import TestCaseGui

class myTests(TestCaseGui):

   def testOne(self):
      self.assertEqual(str_mult("ha"), "hahaha", "Testing that str_mult('ha') returns 'hahaha'")
      self.assertEqual(str_mult("ha", 10), "hahahahahahahahahaha", "Testing that str_mult('ha') returns 'hahahahahahahahahaha'")
      self.assertEqual(str_mult("ha", 0), "", "Testing that str_mult('ha', 0) returns ''")

myTests().main()
        </input>
    </program>
  </exercise>
</section>
